[1582272747] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average per CPU: 0.36, 0.24, 0.11
[1582272807] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average per CPU: 0.15, 0.20, 0.11
[1582272867] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average per CPU: 0.09, 0.17, 0.10
[1582272927] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average per CPU: 0.08, 0.16, 0.10
[1582273149] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 12.34, 11.40, 6.53
[1582273209] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 9.53, 10.79, 6.63
[1582273269] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 10.74, 10.91, 6.93
[1582273329] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.53, 10.88, 7.17
[1582274127] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1582274127] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average per CPU: 0.57, 0.19, 0.10
[1582274427] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average per CPU: 0.08, 0.22, 0.16
[1582275400] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Flapping state cleared for service. (25.7% change < 5.0% threshold)
[1582276751] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.37, 5.87, 5.38
[1582276811] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.51, 5.58, 5.31
[1582276871] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 3.19, 5.03, 5.14
[1582276931] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.17, 4.98, 5.11
[1582277231] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.66, 4.39, 4.84
[1582277427] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average per CPU: 0.14, 0.27, 0.17
[1582277727] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average per CPU: 0.13, 0.16, 0.14
[1582277955] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.74, 10.86, 10.07
[1582278015] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.49, 10.76, 10.08
[1582278075] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 9.97, 10.55, 10.05
[1582278119] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582278135] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.33, 10.50, 10.07
[1582278138] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1582278138] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582278220] SERVICE ALERT: vonly-backup;PING;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582278276] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1582278324] SERVICE ALERT: vonly-backup;Root Partition;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582278326] SERVICE ALERT: vonly-backup;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582278386] SERVICE ALERT: vonly-backup;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582278391] SERVICE ALERT: vonly-backup;Swap Usage;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582278435] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.40, 9.63, 9.79
[1582278446] SERVICE ALERT: vonly-backup;Current Users;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582278506] SERVICE ALERT: vonly-backup;Current Users;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582279031] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.94, 6.08, 5.37
[1582279035] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.92, 10.72, 10.02
[1582279091] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.93, 6.04, 5.40
[1582279095] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.99, 11.24, 10.24
[1582279151] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.49, 5.88, 5.38
[1582279155] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.15, 11.38, 10.35
[1582279211] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.31, 5.75, 5.37
[1582279215] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.21, 11.83, 10.57
[1582279320] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1582279442] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1582279645] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1582279648] SERVICE ALERT: vonly-backup;/media/data Partition;OK;HARD;4;/media/data is at 89% capacity, 6.1T of 7.3T
[1582279670] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 246 processes with STATE = RSZDT
[1582279706] SERVICE ALERT: vonly-backup;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1582279724] SERVICE ALERT: vonly-backup;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.06 ms
[1582279760] SERVICE ALERT: vonly-backup;Swap Usage;OK;HARD;4;SWAP OK - 95% free (54032 MB out of 57111 MB)
[1582279810] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.95, 0.66, 0.70
[1582279826] SERVICE ALERT: vonly-backup;Root Partition;OK;HARD;4;/ is at 52% capacity, 107G of 220G
[1582280111] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1582280111] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.44, 4.43, 4.96
[1582280180] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1582280180] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (25.6% change >= 20.0% threshold)
[1582280312] SERVICE ALERT: WINV-101;D:\ Drive Space;CRITICAL;SOFT;1;d:\ - total: 931.32 Gb - used: 846.02 Gb (91%) - free 85.30 Gb (9%)
[1582280410] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.26, 4.40, 2.02
[1582280432] SERVICE ALERT: WINV-101;D:\ Drive Space;CRITICAL;SOFT;2;d:\ - total: 931.32 Gb - used: 846.02 Gb (91%) - free 85.30 Gb (9%)
[1582280470] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.94, 5.82, 2.66
[1582280530] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 6.56, 5.48, 2.75
[1582280552] SERVICE ALERT: WINV-101;D:\ Drive Space;CRITICAL;HARD;3;d:\ - total: 931.32 Gb - used: 846.02 Gb (91%) - free 85.30 Gb (9%)
[1582281129] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 3.69, 4.63, 7.63
[1582281729] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 8.71, 8.79, 8.11
[1582281789] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 7.38, 8.30, 7.98
[1582282749] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 12.06, 12.39, 9.59
[1582282800] SERVICE ALERT: WINV-102;D:\ Drive Space;WARNING;SOFT;1;d:\ - total: 1862.64 Gb - used: 1566.59 Gb (84%) - free 296.05 Gb (16%)
[1582282809] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 7.55, 10.99, 9.28
[1582282811] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.89, 6.27, 5.03
[1582282815] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.40, 9.57, 9.91
[1582282818] SERVICE ALERT: WINV-111;D:\ Drive Space;CRITICAL;SOFT;1;d:\ - total: 931.22 Gb - used: 849.48 Gb (91%) - free 81.74 Gb (9%)
[1582282869] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 5.68, 9.82, 8.98
[1582282871] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 5.47, 5.67, 4.90
[1582282878] SERVICE ALERT: WINV-110;D:\ Drive Space;CRITICAL;SOFT;1;d:\ - total: 931.32 Gb - used: 893.73 Gb (96%) - free 37.59 Gb (4%)
[1582282920] SERVICE ALERT: WINV-102;D:\ Drive Space;WARNING;SOFT;2;d:\ - total: 1862.64 Gb - used: 1566.66 Gb (84%) - free 295.98 Gb (16%)
[1582282929] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.82, 9.69, 8.99
[1582282938] SERVICE ALERT: WINV-107;D:\ Drive Space;CRITICAL;SOFT;1;d:\ - total: 931.51 Gb - used: 924.36 Gb (99%) - free 7.15 Gb (1%)
[1582282938] SERVICE ALERT: WINV-111;D:\ Drive Space;CRITICAL;SOFT;2;d:\ - total: 931.22 Gb - used: 849.52 Gb (91%) - free 81.70 Gb (9%)
[1582282998] SERVICE ALERT: WINV-110;D:\ Drive Space;CRITICAL;SOFT;2;d:\ - total: 931.32 Gb - used: 893.73 Gb (96%) - free 37.59 Gb (4%)
[1582283040] SERVICE ALERT: WINV-102;D:\ Drive Space;WARNING;HARD;3;d:\ - total: 1862.64 Gb - used: 1566.81 Gb (84%) - free 295.83 Gb (16%)
[1582283057] SERVICE ALERT: WINV-105;D:\ Drive Space;CRITICAL;SOFT;1;d:\ - total: 931.32 Gb - used: 886.61 Gb (95%) - free 44.71 Gb (5%)
[1582283058] SERVICE ALERT: WINV-107;D:\ Drive Space;CRITICAL;SOFT;2;d:\ - total: 931.51 Gb - used: 924.42 Gb (99%) - free 7.09 Gb (1%)
[1582283058] SERVICE ALERT: WINV-111;D:\ Drive Space;CRITICAL;HARD;3;d:\ - total: 931.22 Gb - used: 849.58 Gb (91%) - free 81.64 Gb (9%)
[1582283118] SERVICE ALERT: WINV-110;D:\ Drive Space;CRITICAL;HARD;3;d:\ - total: 931.32 Gb - used: 893.73 Gb (96%) - free 37.59 Gb (4%)
[1582283177] SERVICE ALERT: WINV-105;D:\ Drive Space;CRITICAL;SOFT;2;d:\ - total: 931.32 Gb - used: 886.61 Gb (95%) - free 44.71 Gb (5%)
[1582283178] SERVICE ALERT: WINV-107;D:\ Drive Space;CRITICAL;HARD;3;d:\ - total: 931.51 Gb - used: 924.43 Gb (99%) - free 7.08 Gb (1%)
[1582283297] SERVICE ALERT: WINV-105;D:\ Drive Space;CRITICAL;HARD;3;d:\ - total: 931.32 Gb - used: 886.61 Gb (95%) - free 44.71 Gb (5%)
[1582284129] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.44, 5.47, 7.05
[1582284729] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 9.96, 11.43, 9.00
[1582284789] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 7.47, 10.50, 8.83
[1582284849] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 4.96, 9.25, 8.50
[1582284909] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582284909] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.21, 8.62, 8.34
[1582285209] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.75, 6.48, 7.52
[1582286709] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.43, 11.60, 8.67
[1582286769] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 14.54, 11.63, 8.85
[1582286829] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 8.52, 10.63, 8.69
[1582286889] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.08, 9.36, 8.38
[1582287189] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.56, 4.05, 6.35
[1582288544] SERVICE ALERT: vapps-100;Total Processes;OK;HARD;4;PROCS OK: 3757 processes with STATE = RSZDT
[1582292889] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582294068] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1582294156] SERVICE ALERT: WINV-111;CPU Load;WARNING;SOFT;1;CPU Load 86% (5 min average)
[1582294188] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;2;CPU Load 88% (5 min average)
[1582294276] SERVICE ALERT: WINV-111;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1582294308] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1582294578] SERVICE ALERT: WINV-107;D:\ Drive Space;WARNING;HARD;3;d:\ - total: 931.51 Gb - used: 818.41 Gb (88%) - free 113.10 Gb (12%)
[1582295178] SERVICE ALERT: WINV-107;D:\ Drive Space;OK;HARD;3;d:\ - total: 931.51 Gb - used: 548.94 Gb (59%) - free 382.57 Gb (41%)
[1582296152] SERVICE ALERT: WINV-101;D:\ Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 837.47 Gb (90%) - free 93.85 Gb (10%)
[1582296736] SERVICE ALERT: WINV-108;C:\ Drive Space;OK;HARD;3;c:\ - total: 67.89 Gb - used: 49.01 Gb (72%) - free 18.88 Gb (28%)
[1582296752] SERVICE ALERT: WINV-101;D:\ Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 741.65 Gb (80%) - free 189.67 Gb (20%)
[1582305044] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;1;PROCS WARNING: 4529 processes with STATE = RSZDT
[1582305105] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;2;PROCS WARNING: 4530 processes with STATE = RSZDT
[1582305165] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;3;PROCS WARNING: 4529 processes with STATE = RSZDT
[1582305226] SERVICE ALERT: vapps-100;Total Processes;WARNING;HARD;4;PROCS WARNING: 4530 processes with STATE = RSZDT
[1582310508] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1582310628] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 19% (5 min average)
[1582322654] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1582322783] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;2;No data was received from host!
[1582322903] SERVICE ALERT: WINV-104;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1582328589] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.19, 8.83, 6.43
[1582328649] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 10.56, 8.73, 6.56
[1582338318] SERVICE ALERT: WINV-110;D:\ Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 826.96 Gb (89%) - free 104.36 Gb (11%)
[1582341654] SERVICE ALERT: WINV-101;Uptime;WARNING;SOFT;1;No data was received from host!
[1582341781] SERVICE ALERT: WINV-101;Uptime;OK;SOFT;2;System Uptime - 67 day(s) 12 hour(s) 40 minute(s)
[1582352658] SERVICE ALERT: WINV-111;D:\ Drive Space;OK;HARD;3;d:\ - total: 931.22 Gb - used: 625.04 Gb (67%) - free 306.18 Gb (33%)
[1582352897] SERVICE ALERT: WINV-105;D:\ Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 752.98 Gb (81%) - free 178.34 Gb (19%)
[1582353240] SERVICE ALERT: WINV-102;D:\ Drive Space;OK;HARD;3;d:\ - total: 1862.64 Gb - used: 1429.69 Gb (77%) - free 432.94 Gb (23%)
[1582354086] SERVICE ALERT: WINV-105;D:\ Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 740.95 Gb (80%) - free 190.37 Gb (20%)
[1582354670] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes with STATE = RSZDT
[1582354702] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 812.02 Gb (87%) - free 119.30 Gb (13%)
[1582354730] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 253 processes with STATE = RSZDT
[1582354791] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1582354822] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 812.02 Gb (87%) - free 119.30 Gb (13%)
[1582354851] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582354942] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 812.02 Gb (87%) - free 119.30 Gb (13%)
[1582355261] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 90% capacity! 3.1T of 3.6T
[1582355321] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 90% capacity! 3.1T of 3.6T
[1582355381] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;3;/mnt/data01 is at 90% capacity! 3.1T of 3.6T
[1582355441] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 90% capacity! 3.1T of 3.6T
[1582355497] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 84% capacity! 2.9T of 3.6T
[1582355557] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 84% capacity! 2.9T of 3.6T
[1582355617] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 84% capacity! 2.9T of 3.6T
[1582355677] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 84% capacity! 2.9T of 3.6T