[1586589897] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.37, 10.11, 4.52
[1586589958] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 25.81, 13.57, 6.05
[1586590018] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 22.69, 14.75, 6.92
[1586590078] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 24.84, 16.85, 8.14
[1586590441] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.77, 8.79, 5.60
[1586590501] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.01, 8.67, 5.77
[1586590561] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.36, 8.02, 5.72
[1586590621] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.11, 8.49, 6.03
[1586590678] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 43.51, 36.70, 21.76
[1586590685] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586590805] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586590925] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1586591525] SERVICE FLAPPING ALERT: WINV-101;CPU Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1586591525] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 7% (5 min average)
[1586599325] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 93% (5 min average)
[1586599445] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 10% (5 min average)
[1586599978] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.11, 2.55, 18.30
[1586600221] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.81, 2.30, 4.34
[1586600878] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.07, 0.27, 7.04
[1586602565] SERVICE FLAPPING ALERT: WINV-101;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586616965] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586617085] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 18% (5 min average)
[1586623920] SERVICE ALERT: vapps-109;SSH;CRITICAL;SOFT;1;Server answer:
[1586623980] SERVICE ALERT: vapps-109;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1586635877] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586636000] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586636120] SERVICE ALERT: WINV-104;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1586649178] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.79, 23.35, 9.32
[1586649238] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 63.02, 31.37, 12.94
[1586649298] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 62.51, 36.97, 16.02
[1586649358] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 46.14, 37.35, 17.49
[1586650021] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.91, 6.86, 5.06
[1586650081] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.51, 6.52, 5.06
[1586650141] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 4.59, 5.93, 4.94
[1586650502] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.09, 6.73, 5.53
[1586650562] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.63, 6.72, 5.60
[1586650622] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 10.46, 7.61, 5.97
[1586650682] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.20, 7.68, 6.11
[1586651882] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.95, 2.27, 4.18
[1586652058] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.21, 2.67, 14.88
[1586652658] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.03, 0.43, 7.83
[1586671201] HOST ALERT: vdb-103;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1586671262] SERVICE ALERT: vdb-103;Swap Usage;OK;HARD;4;SWAP OK - 100% free (10239 MB out of 10239 MB)
[1586671264] SERVICE ALERT: vdb-103;Current Load;OK;SOFT;1;OK - load average: 0.00, 0.00, 0.00
[1586671287] SERVICE ALERT: vdb-103;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 38% capacity, 2.6T of 7.3T
[1586671320] SERVICE ALERT: vdb-103;Current Users;OK;SOFT;1;USERS OK - 1 users currently logged in
[1586671320] SERVICE ALERT: vdb-103;Total Processes;OK;SOFT;1;PROCS OK: 181 processes with STATE = RSZDT
[1586671323] SERVICE ALERT: vdb-103;SSH;OK;SOFT;1;SSH OK - OpenSSH_7.6p1 Ubuntu-4ubuntu0.2 (protocol 2.0)
[1586671326] SERVICE ALERT: vdb-103;PING;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.07 ms
[1586671381] SERVICE ALERT: vdb-103;Root Partition;OK;SOFT;1;/ is at 5% capacity, 18G of 457G
[1586671827] HOST ALERT: vdb-103;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1586671865] SERVICE ALERT: vdb-103;Current Load;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586671865] SERVICE ALERT: vdb-103;Swap Usage;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586671890] SERVICE ALERT: vdb-103;/mnt/data01 Partition;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586671890] HOST ALERT: vdb-103;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.4.103)
[1586671923] SERVICE ALERT: vdb-103;Total Processes;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586671923] SERVICE ALERT: vdb-103;Current Users;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586671923] SERVICE ALERT: vdb-103;PING;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586671923] SERVICE ALERT: vdb-103;SSH;CRITICAL;HARD;1;connect to address 192.168.4.103 and port 22: No route to host
[1586671953] HOST ALERT: vdb-103;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.4.103)
[1586671984] SERVICE ALERT: vdb-103;Root Partition;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586672016] HOST ALERT: vdb-103;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.4.103)
[1586672079] HOST ALERT: vdb-103;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.4.103)
[1586672143] HOST ALERT: vdb-103;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.40 ms
[1586672162] SERVICE ALERT: vdb-103;Swap Usage;OK;SOFT;1;SWAP OK - 100% free (10239 MB out of 10239 MB)
[1586672163] SERVICE ALERT: vdb-103;Current Load;OK;SOFT;1;OK - load average: 0.50, 0.23, 0.09
[1586672187] SERVICE ALERT: vdb-103;/mnt/data01 Partition;OK;SOFT;1;/mnt/data01 is at 38% capacity, 2.6T of 7.3T
[1586672220] SERVICE ALERT: vdb-103;Current Users;OK;SOFT;1;USERS OK - 0 users currently logged in
[1586672220] SERVICE ALERT: vdb-103;Total Processes;OK;SOFT;1;PROCS OK: 179 processes with STATE = RSZDT
[1586672223] SERVICE ALERT: vdb-103;SSH;OK;SOFT;1;SSH OK - OpenSSH_7.6p1 Ubuntu-4ubuntu0.2 (protocol 2.0)
[1586672226] SERVICE ALERT: vdb-103;PING;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.09 ms
[1586672281] SERVICE ALERT: vdb-103;Root Partition;OK;SOFT;1;/ is at 5% capacity, 18G of 457G