[1585033264] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585033384] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1035.65 Gb (56%) - free 826.74 Gb (44%)
[1585034225] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 17.99, 13.55, 10.44
[1585034286] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 18.60, 14.53, 10.97
[1585034302] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.59, 27.44, 12.77
[1585034346] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 18.21, 15.13, 11.40
[1585034362] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 49.02, 32.11, 15.31
[1585034406] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 18.84, 15.85, 11.88
[1585034422] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 33.17, 30.75, 15.90
[1585034482] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.66, 30.24, 16.64
[1585034782] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 27.27, 24.65, 17.92
[1585034800] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.58, 8.13, 5.37
[1585034860] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.24, 7.86, 5.45
[1585034920] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.49, 8.09, 5.69
[1585034980] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.29, 7.86, 5.76
[1585035082] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.51, 28.59, 21.65
[1585036882] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1585036882] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.33, 18.77, 19.81
[1585037182] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 16.65, 20.06, 20.31
[1585038723] SERVICE ALERT: vonly-replica;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1585038783] SERVICE ALERT: vonly-replica;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1585038843] SERVICE ALERT: vonly-replica;Current Users;WARNING;SOFT;3;USERS WARNING - 6 users currently logged in
[1585038903] SERVICE ALERT: vonly-replica;Current Users;WARNING;HARD;4;USERS WARNING - 6 users currently logged in
[1585038906] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.18, 27.31, 25.18
[1585039206] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 17.36, 23.17, 24.19
[1585039882] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.00, 6.22, 15.53
[1585040380] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.98, 2.24, 4.21
[1585040782] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.23, 0.74, 6.32
[1585042806] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.86, 8.82, 9.94
[1585044003] SERVICE ALERT: vonly-replica;Current Users;OK;HARD;4;USERS OK - 5 users currently logged in
[1585044303] SERVICE ALERT: vonly-replica;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1585044363] SERVICE ALERT: vonly-replica;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1585044423] SERVICE ALERT: vonly-replica;Current Users;WARNING;SOFT;3;USERS WARNING - 6 users currently logged in
[1585044483] SERVICE ALERT: vonly-replica;Current Users;WARNING;HARD;4;USERS WARNING - 6 users currently logged in
[1585045206] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.18, 10.30, 10.06
[1585045266] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.16, 10.27, 10.06
[1585045326] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.10, 10.22, 10.06
[1585045386] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.03, 10.19, 10.06
[1585045612] SERVICE ALERT: WINV-103;Uptime;WARNING;SOFT;1;No data was received from host!
[1585045732] SERVICE ALERT: WINV-103;Uptime;OK;SOFT;2;System Uptime - 41 day(s) 19 hour(s) 32 minute(s)
[1585045882] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1585045973] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;No data was received from host!
[1585046461] SERVICE ALERT: WINV-103;Uptime;WARNING;SOFT;1;No data was received from host!
[1585046573] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 63.84 Gb (94%) - free 4.43 Gb (6%)
[1585046581] SERVICE ALERT: WINV-103;Uptime;OK;SOFT;2;System Uptime - 41 day(s) 19 hour(s) 46 minute(s)
[1585046913] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 87% (5 min average)
[1585047033] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1585047182] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;No data was received from host!
[1585047310] SERVICE ALERT: WINV-103;Uptime;WARNING;SOFT;1;No data was received from host!
[1585047313] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585047430] SERVICE ALERT: WINV-103;Uptime;OK;SOFT;2;System Uptime - 41 day(s) 20 hour(s) 0 minute(s)
[1585047433] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1041.57 Gb (56%) - free 820.82 Gb (44%)
[1585047483] SERVICE ALERT: vonly-replica;Current Users;OK;HARD;4;USERS OK - 3 users currently logged in
[1585047782] SERVICE FLAPPING ALERT: WINV-103;OS Drive Space;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1585047782] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 63.84 Gb (94%) - free 4.43 Gb (6%)
[1585048391] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;No data was received from host!
[1585048759] SERVICE ALERT: WINV-103;Uptime;WARNING;SOFT;1;No data was received from host!
[1585048762] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585048879] SERVICE ALERT: WINV-103;Uptime;OK;SOFT;2;System Uptime - 41 day(s) 20 hour(s) 24 minute(s)
[1585048882] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1041.65 Gb (56%) - free 820.74 Gb (44%)
[1585048991] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 63.84 Gb (94%) - free 4.43 Gb (6%)
[1585049444] SERVICE ALERT: WINV-103;Memory Usage;WARNING;HARD;3;No data was received from host!
[1585050044] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;HARD;3;Memory usage: total:309720.24 MB - used: 286401.72 MB (92%) - free: 23318.52 MB (8%)
[1585052886] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.14, 9.37, 9.81
[1585053487] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 17.36, 13.61, 11.44
[1585053547] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 18.80, 14.71, 11.96
[1585053607] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 17.98, 15.26, 12.32
[1585053667] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 17.93, 15.78, 12.69
[1585053682] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.02, 14.61, 6.75
[1585053743] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 21.81, 15.77, 7.63
[1585053803] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 15.61, 15.23, 7.96
[1585053863] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.28, 14.81, 8.27
[1585054463] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 31.04, 30.33, 19.29
[1585054480] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.02, 7.53, 5.27
[1585054540] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.39, 7.66, 5.46
[1585054600] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.63, 8.01, 5.73
[1585054660] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.46, 8.32, 5.98
[1585059800] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;No data was received from host!
[1585060404] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 63.99 Gb (94%) - free 4.28 Gb (6%)
[1585061011] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585061013] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;No data was received from host!
[1585061138] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1047.82 Gb (56%) - free 814.57 Gb (44%)
[1585061613] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 64.01 Gb (94%) - free 4.26 Gb (6%)
[1585061860] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.34, 4.00, 4.74
[1585064560] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.26, 6.40, 5.33
[1585064620] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.87, 5.68, 5.15
[1585064680] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.70, 5.51, 5.12
[1585064740] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.41, 5.79, 5.23
[1585065340] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.20, 4.01, 4.65
[1585065467] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585065587] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1049.76 Gb (56%) - free 812.63 Gb (44%)
[1585065863] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.05, 4.88, 15.67
[1585065903] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (11034 MB out of 57111 MB)
[1585065963] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (10926 MB out of 57111 MB)
[1585066023] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 19% free (10725 MB out of 57111 MB)
[1585066083] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 19% free (10683 MB out of 57111 MB)
[1585066763] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.94, 1.00, 6.43
[1585071053] SERVICE ALERT: WINV-103;Memory Usage;WARNING;HARD;3;No data was received from host!
[1585071116] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585071236] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1051.96 Gb (56%) - free 810.43 Gb (44%)
[1585071822] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;No data was received from host!
[1585072258] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;HARD;3;Memory usage: total:320203.12 MB - used: 296178.64 MB (92%) - free: 24024.48 MB (8%)
[1585073622] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 64.01 Gb (94%) - free 4.25 Gb (6%)
[1585074831] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;No data was received from host!
[1585075431] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 64.02 Gb (94%) - free 4.25 Gb (6%)
[1585076765] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585076885] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1053.84 Gb (57%) - free 808.55 Gb (43%)
[1585080614] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585080734] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1055.21 Gb (57%) - free 807.18 Gb (43%)
[1585084817] HOST ALERT: WINV-103;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.3.103)
[1585084880] HOST ALERT: WINV-103;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.3.103)
[1585084923] SERVICE ALERT: WINV-103;CPU Load;CRITICAL;HARD;1;connect to address 192.168.3.103 and port 12489: No route to host
[1585084943] HOST ALERT: WINV-103;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.3.103)
[1585085000] SERVICE ALERT: WINV-103;Uptime;CRITICAL;HARD;1;connect to address 192.168.3.103 and port 12489: Connection refused
[1585085033] HOST ALERT: WINV-103;DOWN;SOFT;4;(Host check timed out after 30.03 seconds)
[1585085092] HOST ALERT: WINV-103;DOWN;SOFT;5;(Host check timed out after 30.03 seconds)
[1585085156] HOST ALERT: WINV-103;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1585085459] SERVICE ALERT: WINV-103;Memory Usage;OK;HARD;3;Memory usage: total:196589.32 MB - used: 61137.20 MB (31%) - free: 135452.12 MB (69%)
[1585085520] SERVICE ALERT: WINV-103;CPU Load;OK;SOFT;1;CPU Load 0% (5 min average)
[1585085600] SERVICE ALERT: WINV-103;Uptime;OK;SOFT;1;System Uptime - 0 day(s) 0 hour(s) 14 minute(s)
[1585085883] SERVICE ALERT: vonly-database;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (5621 MB out of 57111 MB)
[1585086831] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;c:\ - total: 68.27 Gb - used: 54.86 Gb (80%) - free 13.41 Gb (20%)
[1585087431] SERVICE ALERT: WINV-103;OS Drive Space;OK;HARD;3;c:\ - total: 68.27 Gb - used: 49.46 Gb (72%) - free 18.80 Gb (28%)
[1585092940] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.15, 6.50, 5.13
[1585093000] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.75, 6.09, 5.07
[1585093060] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.23, 6.22, 5.18
[1585093120] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.22, 5.69, 5.07
[1585093354] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585093463] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 11.71, 14.57, 10.58
[1585093474] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1585093523] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 11.75, 14.10, 10.68
[1585093584] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 15.01, 14.47, 11.01
[1585093644] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 23.47, 17.28, 12.21
[1585093944] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 35.46, 28.88, 18.60
[1585094194] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1585094244] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.66, 17.73, 16.60
[1585094314] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1585094544] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1585094544] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 41.13, 28.64, 21.48
[1585098831] SERVICE FLAPPING ALERT: WINV-103;OS Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585099644] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.15, 7.26, 17.72
[1585099720] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.50, 2.85, 4.45
[1585100845] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.47, 1.26, 6.08
[1585103167] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.33, 8.97, 9.70
[1585103402] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 88% capacity! 6.1T of 7.3T
[1585105645] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1585116404] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585116524] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585116644] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585117244] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)