[1586329391] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.61, 9.92, 9.98
[1586329991] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.81, 10.65, 10.07
[1586330035] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 19.82, 10.85, 4.66
[1586330051] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.60, 10.99, 10.23
[1586330095] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 21.81, 12.88, 5.73
[1586330111] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.53, 11.57, 10.48
[1586330155] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 21.08, 14.76, 6.86
[1586330171] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1586330171] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.45, 11.89, 10.66
[1586330215] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 21.70, 15.97, 7.76
[1586330515] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.56, 26.04, 14.36
[1586330578] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.11, 7.03, 5.07
[1586330638] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.40, 7.28, 5.29
[1586330698] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.87, 7.59, 5.51
[1586330758] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.83, 7.30, 5.55
[1586331661] SERVICE ALERT: vonly-database;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 3.1T of 3.6T
[1586334720] SERVICE ALERT: WINV-111;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1586334840] SERVICE ALERT: WINV-111;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1586334960] SERVICE ALERT: WINV-111;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1586335571] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1586337964] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1586338084] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1586338271] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.44, 8.95, 9.89
[1586338558] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.00, 2.39, 4.63
[1586338916] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.19, 1.35, 14.72
[1586339516] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.16, 0.67, 8.00
[1586342228] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 80% (5 min average)
[1586342356] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586342404] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586342476] SERVICE ALERT: WINV-104;CPU Load;OK;HARD;3;CPU Load 40% (5 min average)
[1586342524] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586342644] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 7% (5 min average)
[1586346066] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1586346186] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1586347316] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.70, 9.28, 4.32
[1586347376] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 22.67, 12.41, 5.71
[1586347436] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 34.17, 17.49, 7.87
[1586347496] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.26, 20.25, 9.42
[1586348458] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.12, 7.11, 5.48
[1586348518] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.14, 6.70, 5.44
[1586348578] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.32, 6.57, 5.47
[1586348638] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.87, 6.81, 5.63
[1586348644] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1586348764] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586348884] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 6% (5 min average)
[1586348938] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.19, 3.69, 4.62
[1586349185] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 4.62, 4.66, 4.90
[1586349296] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.33, 4.95, 14.99
[1586349337] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 10.45, 7.17, 5.39
[1586349397] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 11.82, 8.13, 5.83
[1586349457] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 13.05, 9.12, 6.32
[1586349517] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 13.97, 10.10, 6.83
[1586350196] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.75, 0.98, 6.12
[1586351317] SERVICE ALERT: vonly-database;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.38, 23.27, 20.47
[1586355154] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1586355274] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 10% (5 min average)
[1586355671] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 14.17, 13.13, 10.23
[1586355731] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 14.60, 13.45, 10.53
[1586355791] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 15.17, 13.83, 10.84
[1586355851] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 14.91, 14.00, 11.09
[1586366317] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 18.76, 19.09, 19.78
[1586370517] SERVICE ALERT: vonly-database;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 20.51, 20.59, 20.19
[1586372617] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 19.28, 19.65, 19.95
[1586372951] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.46, 9.54, 9.93
[1586389796] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 47.64, 24.22, 10.64
[1586389856] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 62.16, 32.02, 14.13
[1586389916] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 61.84, 37.47, 17.14
[1586389976] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 64.45, 42.89, 20.29
[1586390338] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.14, 7.53, 5.40
[1586390398] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.42, 7.97, 5.67
[1586390458] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.27, 7.61, 5.69
[1586390519] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.67, 7.53, 5.78
[1586391485] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 6.32, 6.34, 5.05
[1586391545] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 6.28, 6.33, 5.13
[1586391605] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 5.72, 6.15, 5.14
[1586391665] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 5.20, 5.93, 5.13
[1586392319] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.64, 2.97, 4.50
[1586392376] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.09, 5.68, 19.61
[1586393276] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.86, 0.96, 7.88
[1586408917] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.37, 4.42, 4.91
[1586413507] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1586413627] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1586413865] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 4.79, 4.56, 4.94
[1586414761] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 85% capacity! 2.9T of 3.6T
[1586415065] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 5.97, 5.41, 5.15
[1586415125] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 5.70, 5.43, 5.17
[1586415186] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 5.82, 5.49, 5.21
[1586415246] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 5.37, 5.41, 5.19