[1586416561] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 78% capacity, 2.7T of 3.6T
[1586416976] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 36.84, 23.53, 11.34
[1586417036] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 41.75, 27.40, 13.44
[1586417096] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 32.83, 27.69, 14.42
[1586417156] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 26.11, 26.54, 14.85
[1586417219] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.24, 7.36, 5.08
[1586417279] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.39, 8.05, 5.47
[1586417339] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.84, 7.52, 5.45
[1586417399] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.16, 7.73, 5.67
[1586417646] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 3.33, 4.54, 4.92
[1586419751] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 13.52, 11.80, 10.62
[1586419811] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.05, 11.95, 10.75
[1586419871] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 14.74, 12.56, 11.03
[1586419931] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 15.56, 13.17, 11.34
[1586422594] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1586422714] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1586426161] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 2.8T of 3.6T
[1586426222] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 2.8T of 3.6T
[1586426282] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 2.8T of 3.6T
[1586426342] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 2.8T of 3.6T
[1586427360] SERVICE ALERT: WINV-111;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1586427431] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.76, 7.49, 9.21
[1586427480] SERVICE ALERT: WINV-111;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1586428799] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.23, 2.18, 4.21
[1586428856] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.23, 3.28, 15.70
[1586429756] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.46, 1.54, 7.08
[1586430817] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.72, 6.98, 5.46
[1586430877] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.27, 6.80, 5.49
[1586430937] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.09, 6.66, 5.53
[1586430997] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.16, 6.57, 5.57
[1586433397] SERVICE ALERT: vonly-database;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.64, 23.19, 20.11
[1586438342] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 79% capacity, 2.7T of 3.6T
[1586443142] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 2.8T of 3.6T
[1586443202] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 2.8T of 3.6T
[1586443262] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 2.8T of 3.6T
[1586443323] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 2.8T of 3.6T
[1586450284] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 89% (5 min average)
[1586450404] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586450524] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1586451124] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1586452031] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 14.17, 13.52, 10.81
[1586452091] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 14.34, 13.67, 11.04
[1586452151] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 15.00, 13.97, 11.31
[1586452211] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 15.49, 14.31, 11.59
[1586461298] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 19.01, 19.68, 19.96
[1586469911] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.24, 9.50, 9.87
[1586476256] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.58, 8.15, 3.70
[1586476316] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 23.74, 11.66, 5.18
[1586476376] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.93, 15.90, 7.07
[1586476436] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 31.68, 19.07, 8.72
[1586477699] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.03, 6.81, 5.63
[1586477759] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.90, 7.04, 5.79
[1586477819] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 10.95, 7.99, 6.19
[1586477879] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.79, 8.07, 6.34
[1586479137] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.76, 5.14, 18.33
[1586479379] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.95, 2.03, 4.04
[1586480037] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.11, 0.41, 7.05
[1586483532] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1586483652] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1586491598] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.73, 4.27, 4.76
[1586494635] SERVICE ALERT: WINV-102;OS Drive Space;WARNING;SOFT;1;No data was received from host!
[1586494753] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586494755] SERVICE ALERT: WINV-102;OS Drive Space;OK;SOFT;2;c:\ - total: 931.41 Gb - used: 489.23 Gb (53%) - free 442.19 Gb (47%)
[1586494873] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1586495281] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 16% free (1607 MB out of 10239 MB)
[1586495341] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 16% free (1623 MB out of 10239 MB)
[1586495401] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 16% free (1629 MB out of 10239 MB)
[1586495461] SERVICE ALERT: vdb-103;Swap Usage;WARNING;HARD;4;SWAP WARNING - 17% free (1646 MB out of 10239 MB)
[1586496600] SERVICE ALERT: WINV-111;CPU Load;WARNING;SOFT;1;CPU Load 80% (5 min average)
[1586496720] SERVICE ALERT: WINV-111;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)