[1582791437] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 26.83, 16.51, 7.89
[1582791497] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 31.58, 19.96, 9.64
[1582791590] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 18.52, 18.69, 10.15
[1582791650] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.94, 17.68, 10.33
[1582791793] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 324 processes with STATE = RSZDT
[1582791853] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 246 processes with STATE = RSZDT
[1582792060] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.27, 5.74, 3.65
[1582792060] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 269 processes with STATE = RSZDT
[1582792060] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 48.80, 32.85, 18.72
[1582792120] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 6.56, 5.78, 3.80
[1582792120] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1582792408] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.12, 24.08, 19.98
[1582792708] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1582792708] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 44.08, 26.08, 21.03
[1582792708] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.51, 5.92, 5.09
[1582792788] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.23, 6.41, 5.33
[1582792848] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.79, 7.00, 5.60
[1582792908] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.98, 6.86, 5.64
[1582794587] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1582794647] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 252 processes with STATE = RSZDT
[1582794707] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1582794767] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1582794807] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.11, 12.16, 10.15
[1582794867] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;2;OK - load average: 8.99, 11.32, 9.98
[1582795227] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 14.20, 11.28, 10.18
[1582795287] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 14.64, 11.97, 10.49
[1582795347] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 15.00, 12.58, 10.79
[1582795407] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 16.45, 13.42, 11.20
[1582796007] SERVICE ALERT: vdb-103;Swap Usage;OK;HARD;4;SWAP OK - 21% free (418 MB out of 2047 MB)
[1582796098] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.42, 8.59, 6.18
[1582796158] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 5.34, 7.64, 6.00
[1582796218] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.69, 7.02, 5.88
[1582796278] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 3.41, 6.29, 5.70
[1582796308] SERVICE ALERT: vdb-staging;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 60% capacity, 5.1T of 9.1T
[1582796878] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.75, 3.96, 4.66
[1582797167] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582797467] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582797527] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1582797587] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1582797647] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1582798408] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582799008] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.51, 13.56, 18.71
[1582799813] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.46, 3.71, 4.77
[1582801707] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.73, 8.73, 9.83
[1582803508] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 44.76, 29.38, 22.38
[1582806808] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.54, 7.29, 17.85
[1582807708] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.20, 0.86, 7.09
[1582813547] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 249 processes with STATE = RSZDT
[1582814860] SERVICE ALERT: WINV-111;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1582814980] SERVICE ALERT: WINV-111;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1582819247] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1582819307] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1582831801] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.43 Gb (80%) - free 185.89 Gb (20%)
[1582831921] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.53 Gb (80%) - free 185.79 Gb (20%)
[1582832041] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.63 Gb (80%) - free 185.69 Gb (20%)
[1582833841] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 741.93 Gb (80%) - free 189.39 Gb (20%)
[1582852408] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 26.29, 16.90, 10.65
[1582852468] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 25.47, 18.23, 11.49
[1582852528] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 23.74, 19.03, 12.19
[1582852588] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 26.38, 20.58, 13.15
[1582852888] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 40.53, 31.38, 19.74
[1582853869] SERVICE ALERT: WINV-101;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1582853998] SERVICE ALERT: WINV-101;Memory Usage;WARNING;SOFT;2;No data was received from host!
[1582854082] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582854120] SERVICE ALERT: WINV-101;Memory Usage;OK;HARD;3;Memory usage: total:98284.48 MB - used: 42032.29 MB (43%) - free: 56252.19 MB (57%)
[1582854202] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1582854414] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.48, 6.16, 5.40
[1582854474] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.92, 5.79, 5.32
[1582854534] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.68, 5.67, 5.30
[1582854594] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.30, 5.91, 5.41
[1582855794] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.96, 4.39, 4.98
[1582857388] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.94, 17.82, 19.69
[1582858288] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.28, 1.22, 7.83
[1582871440] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.67, 4.64, 2.86
[1582871505] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 5.18, 5.29, 3.53
[1582873665] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.49, 6.84, 4.77
[1582873725] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.36, 6.29, 4.71
[1582873785] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 3.56, 5.51, 4.54
[1582876113] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.15, 8.35, 5.06
[1582876113] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 315 processes with STATE = RSZDT
[1582876173] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 8.01, 7.85, 5.11
[1582876173] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1582876233] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.66, 6.91, 4.96
[1582876233] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 249 processes with STATE = RSZDT
[1582876293] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 3.54, 6.18, 4.83
[1582876597] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 5.46, 5.30, 4.76
[1582876597] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1582876657] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 257 processes with STATE = RSZDT
[1582876717] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 254 processes with STATE = RSZDT
[1582876797] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 248 processes with STATE = RSZDT