[1585293930] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 26.05, 16.13, 7.93
[1585293990] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.51, 18.01, 9.07
[1585294050] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 25.49, 19.40, 10.13
[1585294110] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 25.01, 20.35, 11.03
[1585294410] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.94, 26.89, 16.11
[1585294447] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 15.66, 13.65, 10.95
[1585294507] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 14.41, 13.69, 11.14
[1585294567] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.78, 13.63, 11.28
[1585294627] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.89, 13.67, 11.44
[1585294721] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.92, 7.12, 5.37
[1585294781] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.96, 6.97, 5.42
[1585294841] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.65, 6.80, 5.46
[1585294901] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 10.51, 7.62, 5.81
[1585295467] SERVICE ALERT: vinfra-200;Swap Usage;OK;HARD;4;SWAP OK - 21% free (2289 MB out of 11215 MB)
[1585297083] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 5.5T of 7.3T
[1585297143] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 5.5T of 7.3T
[1585297203] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 5.5T of 7.3T
[1585297263] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 5.5T of 7.3T
[1585298310] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.27, 15.37, 19.71
[1585298767] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 18% free (2011 MB out of 11215 MB)
[1585298827] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 18% free (2015 MB out of 11215 MB)
[1585298887] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 19% free (2019 MB out of 11215 MB)
[1585298910] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1585298910] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.06, 27.77, 23.06
[1585298947] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;HARD;4;SWAP WARNING - 19% free (2023 MB out of 11215 MB)
[1585299423] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 89% capacity! 3.1T of 3.6T
[1585299723] SERVICE ALERT: vdb-101;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 74% capacity, 2.6T of 3.6T
[1585301347] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 1% free (1 MB out of 11215 MB)
[1585302275] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1585302395] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1585303410] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.33, 12.62, 19.75
[1585303710] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 61.58, 28.95, 23.62
[1585306110] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.31, 14.40, 18.48
[1585306410] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 11.63, 19.21, 20.27
[1585306710] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.42, 16.08, 18.72
[1585307010] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 56.22, 32.47, 23.91
[1585308394] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585308514] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 14% (5 min average)
[1585312710] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585312920] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;1;/ is at 81% capacity! 351G of 457G
[1585312980] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;2;/ is at 82% capacity! 352G of 457G
[1585313040] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;3;/ is at 82% capacity! 353G of 457G
[1585313100] SERVICE ALERT: vdb-103;Root Partition;WARNING;HARD;4;/ is at 82% capacity! 355G of 457G
[1585313201] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.93, 4.04, 4.89
[1585313282] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 9.66, 8.72, 5.27
[1585313342] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 7.53, 8.21, 5.31
[1585313402] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 8.62, 8.41, 5.56
[1585313462] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 8.51, 8.42, 5.74
[1585313610] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.16, 5.83, 18.69
[1585314510] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.04, 0.40, 7.19
[1585316100] SERVICE ALERT: vdb-103;Root Partition;CRITICAL;HARD;4;/ is at 91% capacity! 391G of 457G
[1585317293] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585317414] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585317534] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1585318597] SERVICE ALERT: vinfra-200;MySQL Connection Time;WARNING;SOFT;1;WARNING - 4.17 seconds to connect as nagios
[1585318657] SERVICE ALERT: vinfra-200;MySQL Connection Time;OK;SOFT;2;OK - 0.01 seconds to connect as nagios
[1585318801] SERVICE ALERT: vdb-103;Root Partition;OK;HARD;4;/ is at 22% capacity, 93G of 457G
[1585320662] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 3.64, 4.54, 4.93
[1585321262] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 5.03, 5.16, 5.07
[1585321322] SERVICE ALERT: vdb-103;Current Load;OK;SOFT;2;OK - load average: 4.51, 4.95, 5.00
[1585321382] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 5.23, 5.08, 5.04
[1585321442] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 5.40, 5.15, 5.07
[1585321502] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 5.77, 5.29, 5.12
[1585321562] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 4.85, 5.11, 5.07
[1585322162] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 5.08, 4.74, 4.90
[1585331324] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1585331444] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 26% (5 min average)
[1585333743] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;No data was received from host!
[1585333828] SERVICE ALERT: WINV-101;OS Drive Space;WARNING;SOFT;1;No data was received from host!
[1585333829] SERVICE ALERT: WINV-101;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1585333871] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585333951] SERVICE ALERT: WINV-101;Memory Usage;OK;SOFT;2;Memory usage: total:98284.48 MB - used: 48206.98 MB (49%) - free: 50077.50 MB (51%)
[1585333951] SERVICE ALERT: WINV-101;OS Drive Space;OK;SOFT;2;c:\ - total: 68.27 Gb - used: 43.58 Gb (64%) - free 24.68 Gb (36%)
[1585333993] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585334593] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1585337462] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 6.18, 6.31, 5.08
[1585337522] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 6.21, 6.29, 5.15
[1585337527] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.84, 9.55, 9.99
[1585337564] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 80% (5 min average)
[1585337582] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 5.76, 6.12, 5.16
[1585337642] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 6.32, 6.22, 5.26
[1585337684] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 22% (5 min average)
[1585338793] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585338842] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 5.12, 4.90, 4.97
[1585338913] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585339033] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 10% (5 min average)
[1585343527] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.85, 10.64, 10.12
[1585343587] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.55, 10.58, 10.13
[1585343647] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.11, 10.47, 10.12
[1585343707] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.22, 10.41, 10.12
[1585344307] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.46, 9.85, 10.00
[1585346707] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.47, 10.39, 10.02
[1585346767] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.83, 10.49, 10.08
[1585346827] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.87, 10.54, 10.12
[1585346887] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.59, 10.52, 10.14
[1585352011] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 14.13, 11.10, 7.20
[1585352071] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 12.69, 11.22, 7.49
[1585352131] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 13.04, 11.64, 7.87
[1585352191] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.21, 12.01, 8.22
[1585353091] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 6.95, 5.38, 7.66
[1585353391] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.68, 19.48, 12.95
[1585353451] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 45.86, 25.96, 15.59
[1585353511] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 43.57, 28.61, 17.14
[1585353571] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.57, 30.05, 18.36
[1585354001] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.96, 6.90, 5.42
[1585354061] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.89, 6.46, 5.36
[1585354121] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.95, 6.75, 5.54
[1585354182] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.52, 6.67, 5.59
[1585358804] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585358924] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 18% (5 min average)
[1585358971] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.30, 8.25, 16.78
[1585359582] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.87, 3.33, 4.80
[1585359871] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.99, 1.35, 7.32
[1585366387] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.90, 9.90, 10.00
[1585377187] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.83, 10.79, 10.14
[1585377247] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.84, 10.80, 10.18
[1585377307] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.55, 10.72, 10.19
[1585377367] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.37, 10.62, 10.19
[1585377421] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 9.49, 10.33, 6.53
[1585377481] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 9.39, 10.16, 6.71
[1585377541] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 9.35, 9.99, 6.88
[1585377602] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 9.26, 9.82, 7.01
[1585377634] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585377754] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 33% (5 min average)