[1581497042] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.78, 6.11, 5.27
[1581497102] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.85, 6.09, 5.31
[1581497162] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.24, 6.12, 5.37
[1581497222] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.80, 6.04, 5.39
[1581499166] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581499226] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581499286] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1581499346] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1581499646] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581499706] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581499766] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1581499826] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1581500126] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581500186] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581500246] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1581500606] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581500666] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581500726] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1581501686] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581501746] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581501806] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1581501866] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1581502166] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 249 processes with STATE = RSZDT
[1581503366] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581503426] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581503486] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1581503546] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1581503846] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1581503846] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1581504146] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581504206] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581504266] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1581504326] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1581504926] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1581505226] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581505286] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1581507146] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581507206] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581507266] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1581507422] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.14, 4.08, 4.78
[1581507626] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1581507686] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1581508046] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581508106] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 252 processes with STATE = RSZDT
[1581508166] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1581508226] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 252 processes with STATE = RSZDT
[1581508633] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 8.11, 12.08, 18.87
[1581508933] SERVICE ALERT: vdb-staging;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 29.19, 21.97, 21.16
[1581509822] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.87, 5.52, 5.04
[1581509882] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.60, 5.49, 5.06
[1581509942] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.41, 5.43, 5.06
[1581510002] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.39, 5.42, 5.08
[1581510902] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.62, 4.60, 4.85
[1581511226] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1581513302] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.91, 5.61, 5.16
[1581513362] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.84, 5.64, 5.20
[1581513422] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.62, 5.61, 5.22
[1581513482] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1581513482] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.19, 5.74, 5.28
[1581513926] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1581515126] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581515186] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1581515846] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1581515906] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1581515966] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1581517082] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STOPPED; Service appears to have stopped flapping (4.7% change < 5.0% threshold)
[1581520742] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.46, 7.41, 5.67
[1581520802] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.03, 6.91, 5.61
[1581520862] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.16, 6.51, 5.55
[1581520922] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.99, 6.19, 5.50
[1581522722] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.59, 4.68, 4.96
[1581538682] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.07, 4.59, 4.96
[1581540782] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.32, 5.07, 5.02
[1581540842] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.07, 5.04, 5.01
[1581540902] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 4.89, 5.00, 5.00
[1581541262] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 4.82, 5.01, 5.01
[1581541322] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.25, 5.10, 5.04
[1581541382] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.00, 5.06, 5.03
[1581541442] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 4.94, 5.02, 5.01
[1581542342] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 5.14, 5.01, 5.00
[1581542642] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.10, 5.01, 5.01
[1581542702] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;2;OK - load average: 4.89, 4.96, 4.99
[1581543362] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.10, 5.05, 5.02
[1581543422] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;2;OK - load average: 4.97, 5.01, 5.00
[1581543482] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.13, 5.05, 5.02
[1581543542] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;2;OK - load average: 4.96, 5.01, 5.00
[1581543602] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.12, 5.04, 5.01
[1581543662] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;2;OK - load average: 4.79, 4.98, 4.99
[1581543722] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.10, 5.03, 5.01
[1581543782] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.22, 5.06, 5.02
[1581543842] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 4.81, 4.97, 4.99
[1581545402] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.40, 5.18, 5.04
[1581545462] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.23, 5.19, 5.05
[1581545522] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.28, 5.22, 5.07
[1581545582] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 4.87, 5.12, 5.04
[1581546482] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1581546482] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.27, 4.71, 4.94
[1581551582] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1581553934] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 17.91, 19.25, 19.91
[1581554534] SERVICE ALERT: vdb-staging;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 20.51, 20.65, 20.23
[1581555782] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 8.02, 6.54, 5.37
[1581555842] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 7.14, 6.55, 5.45
[1581555902] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.26, 6.41, 5.47
[1581555962] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.18, 6.34, 5.51
[1581560234] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 19.63, 18.93, 19.59
[1581561662] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.29, 4.52, 4.96
[1581562334] SERVICE ALERT: vdb-staging;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 21.62, 21.07, 20.29
[1581564134] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 16.19, 18.39, 19.50
[1581576734] SERVICE ALERT: vdb-staging;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 20.11, 21.01, 20.27
[1581577334] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 14.62, 17.26, 18.95