[1583568178] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1583569373] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.78, 6.59, 2.74
[1583569433] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 13.52, 7.53, 3.31
[1583569793] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.60, 20.99, 10.37
[1583569853] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 29.18, 22.18, 11.44
[1583569913] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.16, 23.17, 12.44
[1583569973] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.47, 24.48, 13.56
[1583570626] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.35, 12.03, 10.24
[1583570686] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.31, 12.09, 10.37
[1583570746] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.45, 12.13, 10.49
[1583570806] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.89, 12.32, 10.66
[1583571046] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.34, 7.36, 5.43
[1583571106] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.62, 7.03, 5.44
[1583571166] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.95, 6.87, 5.47
[1583571226] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.53, 6.18, 5.32
[1583576222] SERVICE ALERT: WINV-102;Data Drive Space;WARNING;SOFT;1;d:\ - total: 1862.64 Gb - used: 1490.30 Gb (80%) - free 372.34 Gb (20%)
[1583576342] SERVICE ALERT: WINV-102;Data Drive Space;WARNING;SOFT;2;d:\ - total: 1862.64 Gb - used: 1490.35 Gb (80%) - free 372.29 Gb (20%)
[1583576462] SERVICE ALERT: WINV-102;Data Drive Space;WARNING;HARD;3;d:\ - total: 1862.64 Gb - used: 1490.39 Gb (80%) - free 372.25 Gb (20%)
[1583576563] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583576683] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1583581006] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.27, 9.09, 9.78
[1583581789] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583581909] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583582026] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.57, 4.37, 4.86
[1583582029] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 8% (5 min average)
[1583582326] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.63, 5.05, 5.04
[1583582386] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.38, 5.10, 5.05
[1583582446] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.63, 5.49, 5.19
[1583582506] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.07, 4.71, 4.94
[1583582506] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.06, 10.31, 10.07
[1583582566] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;2;OK - load average: 8.56, 9.83, 9.92
[1583583062] SERVICE ALERT: WINV-102;Data Drive Space;OK;HARD;3;d:\ - total: 1862.64 Gb - used: 1489.13 Gb (80%) - free 373.51 Gb (20%)
[1583583173] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.28, 3.32, 16.84
[1583584073] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.23, 0.58, 6.71
[1583590183] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.81 Gb (80%) - free 185.51 Gb (20%)
[1583590303] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.97 Gb (80%) - free 185.35 Gb (20%)
[1583590423] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 746.16 Gb (80%) - free 185.16 Gb (20%)
[1583613205] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583613325] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1583618580] SERVICE ALERT: WINV-104;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.11 Gb (80%) - free 186.21 Gb (20%)
[1583618700] SERVICE ALERT: WINV-104;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.11 Gb (80%) - free 186.21 Gb (20%)
[1583618820] SERVICE ALERT: WINV-104;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.11 Gb (80%) - free 186.21 Gb (20%)
[1583623973] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.72, 15.42, 8.62
[1583624033] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 32.63, 19.85, 10.59
[1583624093] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 25.17, 20.32, 11.35
[1583624153] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.66, 22.80, 12.71
[1583624326] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.83, 11.88, 10.48
[1583624386] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.34, 11.57, 10.46
[1583624447] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.08, 11.53, 10.51
[1583624453] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.13, 17.26, 13.35
[1583624507] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.13, 11.15, 10.44
[1583625707] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.60, 9.67, 9.90
[1583626553] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.63, 6.16, 7.47
[1583628953] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 24.23, 15.02, 7.60
[1583629007] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 13.70, 12.11, 10.33
[1583629013] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.90, 17.70, 8.98
[1583629068] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.11, 12.25, 10.49
[1583629073] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 31.13, 20.72, 10.58
[1583629128] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.00, 12.37, 10.64
[1583629133] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.77, 23.36, 12.11
[1583629188] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.37, 12.56, 10.82
[1583630807] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.05, 5.16, 5.03
[1583630867] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.69, 5.19, 5.04
[1583630927] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.65, 5.49, 5.15
[1583630987] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.65, 5.76, 5.25
[1583633387] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.47, 4.51, 4.96
[1583633623] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 743.15 Gb (80%) - free 188.17 Gb (20%)
[1583634778] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1583634833] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.20, 7.59, 17.58
[1583634898] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 31% (5 min average)
[1583635733] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.08, 0.60, 6.86
[1583635788] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.50, 8.64, 9.85
[1583639220] SERVICE ALERT: WINV-104;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 745.00 Gb (80%) - free 186.32 Gb (20%)
[1583647623] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1583647743] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 47% (5 min average)
[1583647864] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583647987] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1583648107] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 67% (5 min average)
[1583648462] SERVICE ALERT: WINV-102;Data Drive Space;WARNING;SOFT;1;d:\ - total: 1862.64 Gb - used: 1490.12 Gb (80%) - free 372.51 Gb (20%)
[1583648582] SERVICE ALERT: WINV-102;Data Drive Space;WARNING;SOFT;2;d:\ - total: 1862.64 Gb - used: 1490.44 Gb (80%) - free 372.20 Gb (20%)
[1583648702] SERVICE ALERT: WINV-102;Data Drive Space;WARNING;HARD;3;d:\ - total: 1862.64 Gb - used: 1490.44 Gb (80%) - free 372.20 Gb (20%)
[1583650429] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583650549] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583650669] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)