[1585552776] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.57, 9.80, 4.80
[1585552836] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 23.78, 12.92, 6.18
[1585552896] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 22.55, 14.60, 7.18
[1585552956] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 27.58, 17.28, 8.56
[1585553556] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 39.53, 29.30, 18.67
[1585553684] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.37, 6.38, 5.19
[1585553744] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.60, 6.18, 5.19
[1585553804] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.61, 7.20, 5.60
[1585553864] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.86, 6.80, 5.56
[1585556289] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.28, 27.45, 25.71
[1585558865] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1585560189] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 20.66, 22.46, 24.04
[1585561356] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.56, 13.05, 19.39
[1585564594] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1585564714] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;2;CPU Load 92% (5 min average)
[1585564835] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585565435] SERVICE ALERT: WINV-104;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1585566157] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 49.55, 27.29, 21.11
[1585566457] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.56, 16.86, 18.81
[1585568257] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 44.23, 25.36, 18.40
[1585572087] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1585572157] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.13, 7.14, 15.92
[1585572164] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.20, 2.83, 4.54
[1585572207] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1585573057] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.54, 2.21, 6.88
[1585577264] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 11.03, 6.87, 5.16
[1585577324] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 11.18, 7.87, 5.62
[1585577384] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 15.84, 9.63, 6.37
[1585577444] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.60, 8.70, 6.25
[1585578004] SERVICE ALERT: vdb-103;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 71% capacity, 4.9T of 7.3T
[1585578157] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 10.33, 10.32, 7.57
[1585578217] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 17.36, 12.43, 8.47
[1585578277] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 13.03, 11.90, 8.53
[1585578338] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 15.69, 13.28, 9.24
[1585578634] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585578754] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1585578874] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1585578938] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.15, 4.93, 6.92
[1585579244] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.92, 3.19, 4.41
[1585579474] SERVICE ALERT: WINV-107;CPU Load;OK;HARD;3;CPU Load 20% (5 min average)
[1585583727] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1585583847] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1585584567] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585584687] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1585586400] SERVICE ALERT: vdb-103;Swap Usage;WARNING;HARD;4;SWAP WARNING - 18% free (1833 MB out of 10239 MB)
[1585586700] SERVICE ALERT: vdb-103;Swap Usage;OK;HARD;4;SWAP OK - 23% free (2301 MB out of 10239 MB)
[1585595465] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 90% capacity! 6.2T of 7.3T
[1585595607] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585595727] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585595847] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1585607144] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.65, 7.00, 5.05
[1585607204] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 6.29, 6.52, 5.00
[1585607264] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.53, 6.30, 5.02
[1585607324] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.41, 5.84, 4.95
[1585607384] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.27, 5.97, 5.06
[1585607438] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 11.23, 10.30, 7.26
[1585607444] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.14, 6.57, 5.33
[1585607498] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 12.11, 10.65, 7.57
[1585607504] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.66, 6.26, 5.31
[1585607558] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 10.69, 10.53, 7.72
[1585607564] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.74, 6.65, 5.51
[1585607618] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.61, 10.93, 8.04
[1585609064] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.72, 2.67, 4.24
[1585609118] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.40, 3.34, 6.99
[1585612419] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 29.88, 13.42, 5.85
[1585612479] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 29.17, 15.74, 7.10
[1585612539] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 31.28, 18.62, 8.62
[1585612599] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 28.57, 19.87, 9.66
[1585612899] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1585612899] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.65, 30.24, 17.05
[1585613564] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.99, 6.18, 5.16
[1585613624] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.26, 5.96, 5.15
[1585613684] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.81, 6.89, 5.51
[1585613744] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.81, 7.57, 5.84
[1585617399] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.13, 14.29, 18.54
[1585618013] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585618133] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)
[1585620044] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.74, 3.03, 4.56
[1585620053] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1585620173] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1585620399] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.63, 2.48, 7.91
[1585623399] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)