[1584428434] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1584428554] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 76% (5 min average)
[1584429794] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 24.70, 8.76, 3.73
[1584429854] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 18.68, 10.04, 4.50
[1584429914] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 17.07, 11.21, 5.25
[1584429974] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.64, 10.92, 5.52
[1584430868] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.15, 6.95, 5.01
[1584430928] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 10.25, 7.56, 5.34
[1584430988] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.86, 7.10, 5.32
[1584431048] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.57, 7.36, 5.53
[1584431174] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.52, 27.21, 20.02
[1584437075] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 93% (5 min average)
[1584437195] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 24% (5 min average)
[1584437401] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.15, 11.46, 10.33
[1584437461] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.42, 11.70, 10.48
[1584437481] SERVICE ALERT: vonly-database;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1584437521] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 14.17, 12.26, 10.74
[1584437550] SERVICE ALERT: vonly-database;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.06 ms
[1584437581] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.49, 12.45, 10.91
[1584438121] SERVICE ALERT: vinfra-200;Swap Usage;OK;HARD;4;SWAP OK - 90% free (10005 MB out of 11215 MB)
[1584439681] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 5.34, 8.10, 9.88
[1584439874] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.58, 15.12, 18.64
[1584440221] SERVICE ALERT: vdb-103;Swap Usage;OK;HARD;4;SWAP OK - 100% free (10239 MB out of 10239 MB)
[1584443474] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.17, 24.93, 19.08
[1584444001] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 5.5T of 7.3T
[1584444061] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 5.5T of 7.3T
[1584444121] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 5.5T of 7.3T
[1584444181] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 5.5T of 7.3T
[1584445274] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.53, 11.78, 17.36
[1584445749] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.15, 3.66, 4.80
[1584446474] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.39, 1.78, 7.05
[1584448020] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 5.81, 6.30, 5.32
[1584448080] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 5.91, 6.21, 5.35
[1584448140] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 6.42, 6.32, 5.44
[1584448200] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 6.53, 6.37, 5.52
[1584452400] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 3.38, 4.05, 4.73
[1584453394] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1584453514] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;2;CPU Load 87% (5 min average)
[1584453634] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584454234] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 7% (5 min average)
[1584472681] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 5.5T of 7.3T
[1584487874] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 11.82, 10.24, 6.35
[1584487935] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 23.31, 13.97, 7.88
[1584487995] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 23.77, 15.85, 8.91
[1584488055] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.16, 20.97, 11.15
[1584488349] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.81, 6.73, 5.42
[1584488355] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.58, 21.44, 14.92
[1584488409] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.66, 6.42, 5.40
[1584488469] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.95, 6.66, 5.53
[1584488529] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.37, 6.67, 5.62
[1584488955] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 46.54, 38.92, 25.13
[1584494356] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.83, 3.59, 18.44
[1584494529] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.14, 3.67, 4.61
[1584495256] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.65, 1.01, 7.54
[1584497435] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584497555] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1584497675] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;1;c:\ - total: 68.27 Gb - used: 54.98 Gb (81%) - free 13.28 Gb (19%)
[1584497795] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;2;c:\ - total: 68.27 Gb - used: 55.09 Gb (81%) - free 13.17 Gb (19%)
[1584497915] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;c:\ - total: 68.27 Gb - used: 55.19 Gb (81%) - free 13.07 Gb (19%)
[1584499355] SERVICE ALERT: WINV-103;Memory Usage;WARNING;SOFT;1;Memory usage: total:209629.53 MB - used: 168498.12 MB (80%) - free: 41131.42 MB (20%)
[1584499475] SERVICE ALERT: WINV-103;Memory Usage;WARNING;SOFT;2;Memory usage: total:209629.53 MB - used: 169877.59 MB (81%) - free: 39751.94 MB (19%)
[1584499595] SERVICE ALERT: WINV-103;Memory Usage;WARNING;HARD;3;Memory usage: total:209629.53 MB - used: 171817.59 MB (82%) - free: 37811.94 MB (18%)
[1584501395] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;HARD;3;Memory usage: total:209629.53 MB - used: 191178.76 MB (91%) - free: 18450.78 MB (9%)
[1584511934] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 89% (5 min average)
[1584512054] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;2;CPU Load 89% (5 min average)
[1584512174] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584512774] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 61% (5 min average)