[1585898147] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.07, 7.68, 3.26
[1585898207] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 23.40, 10.76, 4.61
[1585898267] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 20.52, 12.18, 5.48
[1585898327] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 19.07, 13.21, 6.25
[1585898710] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 31.51, 28.97, 25.64
[1585899891] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.80, 5.67, 5.06
[1585899951] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.28, 5.89, 5.17
[1585900011] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.11, 5.70, 5.15
[1585900071] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.19, 5.33, 5.05
[1585902827] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.60, 25.30, 21.75
[1585907711] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 23.85, 23.53, 24.79
[1585911741] SERVICE ALERT: WINV-102;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585912000] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585912001] SERVICE ALERT: WINV-102;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585912004] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585912053] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585912130] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1585912131] SERVICE ALERT: WINV-102;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1585912134] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1585912183] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1585912235] SERVICE ALERT: WINV-112;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1585912260] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585912261] SERVICE ALERT: WINV-102;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585912264] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585912313] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585912355] SERVICE ALERT: WINV-112;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1585913460] SERVICE ALERT: WINV-102;Memory Usage;OK;HARD;3;Memory usage: total:298295.47 MB - used: 30000.61 MB (10%) - free: 268294.87 MB (90%)
[1585913461] SERVICE ALERT: WINV-102;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 5 minute(s)
[1585913464] SERVICE ALERT: WINV-102;OS Drive Space;OK;HARD;3;c:\ - total: 931.41 Gb - used: 490.28 Gb (53%) - free 441.13 Gb (47%)
[1585913513] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1585913541] SERVICE ALERT: WINV-102;Data Drive Space;WARNING;HARD;3;d:\ - total: 1862.64 Gb - used: 1506.16 Gb (81%) - free 356.48 Gb (19%)
[1585913571] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.09, 2.79, 4.67
[1585914227] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.09, 1.45, 15.37
[1585915127] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.54, 0.68, 6.18
[1585915500] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 5.06, 5.13, 5.01
[1585915560] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 5.15, 5.17, 5.02
[1585915620] SERVICE ALERT: vdb-101;Current Load;OK;SOFT;3;OK - load average: 4.16, 4.88, 4.93
[1585915941] SERVICE ALERT: WINV-102;Data Drive Space;OK;HARD;3;d:\ - total: 1862.64 Gb - used: 1308.52 Gb (70%) - free 554.12 Gb (30%)
[1585916881] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 5.49, 5.18, 5.04
[1585916941] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.19, 5.06
[1585917001] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 5.43, 5.25, 5.09
[1585917061] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 5.37, 5.24, 5.09
[1585918861] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 4.80, 4.74, 4.94
[1585919161] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 5.17, 5.13, 5.07
[1585919221] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 5.10, 5.12, 5.07
[1585919281] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 5.20, 5.12, 5.08
[1585919341] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 5.18, 5.12, 5.08
[1585919641] SERVICE FLAPPING ALERT: vdb-101;Current Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1585919641] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 4.51, 4.80, 4.95
[1585920011] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.24, 27.31, 25.44
[1585920841] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 5.55, 5.32, 5.08
[1585920901] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 5.50, 5.34, 5.10
[1585920961] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 5.79, 5.44, 5.15
[1585921021] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 5.64, 5.49, 5.19
[1585922521] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 4.79, 4.85, 5.00
[1585923722] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 5.33, 5.14, 5.05
[1585923782] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 5.62, 5.25, 5.10
[1585923842] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 5.31, 5.24, 5.10
[1585923902] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 4.85, 5.12, 5.07
[1585924727] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.58, 7.16, 3.32
[1585924787] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 10.52, 7.35, 3.63
[1585925102] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 4.30, 4.54, 4.85
[1585925147] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 12.36, 15.38, 9.22
[1585925207] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 15.56, 16.09, 9.87
[1585925267] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 18.49, 17.02, 10.59
[1585925271] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.22, 8.49, 5.59
[1585925327] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 28.65, 19.85, 11.94
[1585925331] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.12, 8.44, 5.75
[1585925391] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.92, 9.03, 6.14
[1585925451] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.16, 8.87, 6.26
[1585926002] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 5.15, 5.14, 5.03
[1585926062] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 5.03, 5.12, 5.03
[1585926122] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 5.60, 5.27, 5.08
[1585926182] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 5.16, 5.20, 5.07
[1585926527] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.92, 20.17, 15.80
[1585926827] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.87, 18.87, 16.43
[1585927228] SERVICE ALERT: WINV-104;Uptime;WARNING;SOFT;1;No data was received from host!
[1585927238] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 93% (5 min average)
[1585927348] SERVICE ALERT: WINV-104;Uptime;OK;SOFT;2;System Uptime - 189 day(s) 21 hour(s) 49 minute(s)
[1585927358] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;2;CPU Load 88% (5 min average)
[1585927478] SERVICE ALERT: WINV-104;CPU Load;OK;HARD;3;CPU Load 64% (5 min average)
[1585927982] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 4.83, 4.86, 4.96
[1585928927] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1585928927] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 6.12, 5.08, 7.51
[1585930551] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.13, 2.58, 4.44
[1585932182] SERVICE FLAPPING ALERT: vdb-101;Current Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1585932827] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1585933481] SERVICE ALERT: WINV-101;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1585933602] SERVICE ALERT: WINV-101;Memory Usage;OK;SOFT;2;Memory usage: total:98284.48 MB - used: 48106.48 MB (49%) - free: 50178.00 MB (51%)
[1585934017] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;No data was received from host!
[1585934137] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1585934257] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 12% (5 min average)
[1585936657] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585936777] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1585936897] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585938097] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 17% (5 min average)
[1585947001] SERVICE ALERT: vdb-101;Swap Usage;OK;HARD;4;SWAP OK - 99% free (21028 MB out of 21455 MB)
[1585953227] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.25, 11.41, 7.71
[1585953287] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 15.91, 12.02, 8.14
[1585953347] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 12.62, 11.85, 8.33
[1585953407] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.27, 13.67, 9.20
[1585953652] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 11.67, 8.50, 6.13
[1585953712] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 10.07, 8.70, 6.35
[1585953772] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.34, 7.89, 6.22
[1585953832] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.71, 7.16, 6.08
[1585955032] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.57, 2.54, 4.10
[1585955207] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.41, 1.82, 6.20
[1585957608] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 48.21, 22.39, 9.86
[1585957668] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 31.09, 22.55, 10.73
[1585957728] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 21.06, 21.42, 11.09
[1585957788] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 20.53, 21.39, 11.74
[1585958688] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1585958688] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 29.14, 25.21, 19.31
[1585958988] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.52, 19.27, 18.27
[1585960732] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.82, 6.08, 5.26
[1585960792] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.63, 5.96, 5.27
[1585960852] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.54, 6.56, 5.51
[1585960912] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.05, 5.92, 5.36
[1585964688] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585967512] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.39, 3.01, 4.67
[1585967988] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.42, 1.69, 6.15
[1585971097] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585971218] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585971338] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)