[1583655420] SERVICE ALERT: WINV-104;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.18 Gb (80%) - free 186.14 Gb (20%)
[1583655533] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 19.59, 11.99, 5.88
[1583655540] SERVICE ALERT: WINV-104;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.21 Gb (80%) - free 186.11 Gb (20%)
[1583655593] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 19.86, 13.37, 6.74
[1583655653] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 23.69, 15.67, 7.95
[1583655660] SERVICE ALERT: WINV-104;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.26 Gb (80%) - free 186.06 Gb (20%)
[1583655713] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 28.13, 18.47, 9.39
[1583656013] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.53, 25.47, 14.89
[1583656188] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.51, 11.51, 10.16
[1583656248] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.10, 11.63, 10.29
[1583656308] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.47, 11.22, 10.23
[1583656368] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.43, 11.03, 10.22
[1583657087] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.18, 6.11, 5.18
[1583657147] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.36, 5.67, 5.08
[1583657207] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 4.25, 5.30, 4.99
[1583657267] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.14, 5.73, 5.16
[1583657327] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.54, 5.37, 5.08
[1583657387] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.57, 5.94, 5.30
[1583657447] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.04, 5.60, 5.22
[1583657707] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583657827] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 14% (5 min average)
[1583663813] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.42, 13.66, 19.18
[1583664468] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.31, 8.95, 9.85
[1583666514] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.02, 24.28, 20.79
[1583667114] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.78, 18.14, 19.44
[1583667714] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1583667714] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.24, 22.72, 20.59
[1583669747] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.52, 3.28, 4.53
[1583670114] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.99, 8.36, 17.29
[1583671014] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.34, 0.90, 6.92
[1583673223] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.19 Gb (80%) - free 186.12 Gb (20%)
[1583673343] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.36 Gb (80%) - free 185.96 Gb (20%)
[1583673463] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.50 Gb (80%) - free 185.82 Gb (20%)
[1583676114] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583690269] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583690389] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1583698629] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;No data was received from host!
[1583698753] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583698882] SERVICE ALERT: WINV-104;CPU Load;WARNING;HARD;3;No data was received from host!
[1583699482] SERVICE ALERT: WINV-104;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1583708863] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 739.24 Gb (79%) - free 192.08 Gb (21%)
[1583710314] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 28.85, 17.47, 9.26
[1583710375] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.92, 21.34, 11.09
[1583710435] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 32.12, 22.76, 12.22
[1583710495] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.38, 27.13, 14.44
[1583710795] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 23.48, 21.08, 14.98
[1583710847] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.13, 6.06, 5.01
[1583710907] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.45, 6.62, 5.26
[1583710967] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 10.05, 7.05, 5.47
[1583710968] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.04, 12.02, 10.23
[1583711027] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.09, 6.72, 5.46
[1583711028] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.06, 12.32, 10.44
[1583711088] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.09, 12.52, 10.63
[1583711148] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.89, 12.58, 10.77
[1583711995] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.03, 23.48, 20.33
[1583713911] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583714040] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;2;No data was received from host!
[1583714160] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 9% (5 min average)
[1583715548] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1583715668] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1583717327] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.63, 2.85, 4.32
[1583717395] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.37, 7.58, 19.05
[1583717448] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.47, 7.56, 9.31
[1583718295] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.31, 0.81, 7.54
[1583729160] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583729280] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 94% (5 min average)
[1583729400] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1583738395] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 21.36, 13.38, 6.76
[1583738455] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 24.96, 15.84, 8.02
[1583738515] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 21.84, 16.65, 8.79
[1583738575] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 29.60, 19.63, 10.30
[1583738875] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.37, 26.48, 15.65
[1583739348] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.66, 11.16, 10.13
[1583739408] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.54, 11.36, 10.27
[1583739468] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.33, 11.92, 10.53
[1583739527] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.40, 6.28, 5.11
[1583739528] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.97, 12.33, 10.76
[1583739587] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.72, 6.56, 5.28
[1583739647] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.15, 6.76, 5.44
[1583739707] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.95, 7.27, 5.67