[1582100575] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 36.85, 16.23, 6.63
[1582100635] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.22, 19.70, 8.44
[1582100695] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 34.50, 22.70, 10.20
[1582100755] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 29.83, 23.50, 11.25
[1582101355] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.63, 27.75, 18.73
[1582101741] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.02, 6.39, 5.02
[1582101801] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.93, 6.22, 5.05
[1582101861] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.13, 6.03, 5.06
[1582101921] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.98, 6.07, 5.13
[1582101955] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.54, 18.16, 18.77
[1582103043] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 5.53, 6.43, 3.90
[1582103103] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 3.12, 5.58, 3.77
[1582104321] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.41, 4.64, 4.91
[1582104363] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.58, 3.53, 2.93
[1582104423] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.49, 5.98, 3.81
[1582104483] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 10.32, 6.56, 4.15
[1582104543] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 5.75, 5.97, 4.11
[1582104621] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 3.47, 5.01, 5.13
[1582104681] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.60, 5.22, 5.19
[1582104741] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.93, 5.42, 5.26
[1582104801] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.37, 5.64, 5.33
[1582104955] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1582104955] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 24.96, 24.45, 20.83
[1582107955] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1582108143] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 14.00, 6.26, 3.84
[1582108203] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.15, 7.25, 4.34
[1582108263] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 7.45, 6.84, 4.38
[1582108323] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.40, 6.10, 4.28
[1582108623] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 2.40, 3.82, 3.78
[1582109823] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 4.97, 7.07, 5.11
[1582109883] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 3.48, 6.24, 4.94
[1582109943] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 2.72, 5.51, 4.77
[1582111401] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.05, 4.29, 4.91
[1582111555] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.14, 12.16, 19.56
[1582111701] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.15, 5.37, 5.14
[1582111761] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.41, 4.91, 4.99
[1582113955] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.63, 26.28, 21.06
[1582114221] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.66, 5.90, 5.11
[1582114281] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.73, 5.56, 5.04
[1582114341] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.74, 5.37, 5.01
[1582114401] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.28, 5.77, 5.17
[1582114701] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.54, 3.82, 4.55
[1582114855] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.53, 12.82, 19.61
[1582115155] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1582115155] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 137.64, 66.27, 37.91
[1582115301] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 13.34, 9.49, 6.51
[1582115361] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.91, 9.29, 6.63
[1582115421] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 10.51, 9.51, 6.87
[1582115481] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1582115481] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 12.92, 10.22, 7.27
[1582115912] SERVICE ALERT: vapps-100;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 91% capacity! 471G of 551G
[1582115972] SERVICE ALERT: vapps-100;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 91% capacity! 471G of 551G
[1582116032] SERVICE ALERT: vapps-100;/mnt/data01 Partition;WARNING;SOFT;3;/mnt/data01 is at 91% capacity! 471G of 551G
[1582116092] SERVICE ALERT: vapps-100;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 91% capacity! 471G of 551G
[1582116392] SERVICE ALERT: vapps-100;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 78% capacity, 407G of 551G
[1582116681] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.38, 1.71, 4.22
[1582118181] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.38, 6.68, 5.53
[1582118241] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 3.95, 5.79, 5.29
[1582118301] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 2.23, 4.98, 5.04
[1582118361] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.45, 4.52, 4.88
[1582119055] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.76, 10.65, 19.54
[1582120255] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.28, 1.58, 7.05
[1582121661] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582125055] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582129503] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.75, 9.74, 6.57
[1582129563] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.12, 8.89, 6.48
[1582129623] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.62, 8.12, 6.36
[1582129683] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.11, 7.52, 6.26
[1582131183] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.55, 4.47, 4.86
[1582134783] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 5.32, 5.32, 5.04
[1582134843] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 4.58, 5.11, 4.98
[1582154755] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 8.48, 22.14, 12.03
[1582154815] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 7.70, 19.26, 11.66
[1582154875] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 5.48, 16.49, 11.18
[1582154935] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.98, 14.78, 10.93
[1582155535] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 26.88, 30.03, 18.62
[1582155835] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 4.65, 15.03, 15.36
[1582157035] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1582157035] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.12, 2.39, 7.81
[1582161835] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582175935] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 27.79, 26.81, 13.42
[1582175995] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 11.93, 22.40, 12.74
[1582176055] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 7.73, 19.32, 12.29
[1582176115] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.11, 16.45, 11.74
[1582176715] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 50.54, 18.37, 12.29
[1582177015] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.40, 11.19, 11.35
[1582177615] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1582177615] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 6.94, 4.64, 7.38
[1582183015] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)