[1586244114] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.11, 12.69, 7.06
[1586244174] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 17.28, 13.73, 7.78
[1586244234] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 24.88, 16.24, 9.00
[1586244294] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.75, 19.39, 10.53
[1586244311] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.07, 29.42, 25.50
[1586244537] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.53, 8.27, 5.78
[1586244597] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.11, 7.93, 5.81
[1586244657] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.84, 8.57, 6.17
[1586244717] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.23, 8.65, 6.35
[1586247594] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.57, 12.44, 17.81
[1586249994] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 15.22, 21.56, 20.42
[1586250294] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.89, 10.84, 16.12
[1586251494] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1586251494] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.27, 4.52, 7.73
[1586251617] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.51, 3.42, 4.85
[1586253011] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 23.86, 23.59, 24.57
[1586255463] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 20.16, 9.33, 3.86
[1586255524] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 16.79, 10.41, 4.57
[1586255584] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 17.35, 11.74, 5.40
[1586255644] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 20.42, 13.45, 6.37
[1586256294] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1586257902] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586257905] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586257908] SERVICE ALERT: WINV-103;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586257911] SERVICE ALERT: WINV-105;Data Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586257963] SERVICE ALERT: WINV-101;Data Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586257966] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586257969] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586257972] SERVICE ALERT: WINV-105;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258024] SERVICE ALERT: WINV-101;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258027] SERVICE ALERT: WINV-102;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258030] SERVICE ALERT: WINV-104;Data Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258032] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258033] SERVICE ALERT: WINV-105;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258035] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258038] SERVICE ALERT: WINV-103;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258041] SERVICE ALERT: WINV-105;Data Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258085] SERVICE ALERT: WINV-101;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258088] SERVICE ALERT: WINV-103;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258091] SERVICE ALERT: WINV-104;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258093] SERVICE ALERT: WINV-101;Data Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258094] SERVICE ALERT: WINV-105;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258096] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258099] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258102] SERVICE ALERT: WINV-105;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258146] SERVICE ALERT: WINV-101;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258149] SERVICE ALERT: WINV-103;Data Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258152] SERVICE ALERT: WINV-104;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258154] SERVICE ALERT: WINV-101;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258157] SERVICE ALERT: WINV-102;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258160] SERVICE ALERT: WINV-104;Data Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258162] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258163] SERVICE ALERT: WINV-105;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258165] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258168] SERVICE ALERT: WINV-103;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258171] SERVICE ALERT: WINV-105;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258206] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258209] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258212] SERVICE ALERT: WINV-104;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258215] SERVICE ALERT: WINV-101;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258218] SERVICE ALERT: WINV-103;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258221] SERVICE ALERT: WINV-104;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258223] SERVICE ALERT: WINV-101;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258224] SERVICE ALERT: WINV-105;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258226] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258229] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258232] SERVICE ALERT: WINV-105;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258267] SERVICE ALERT: WINV-102;Data Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258270] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258273] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586258276] SERVICE ALERT: WINV-101;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258279] SERVICE ALERT: WINV-103;Data Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258282] SERVICE ALERT: WINV-104;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258284] SERVICE ALERT: WINV-101;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258287] SERVICE ALERT: WINV-102;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258290] SERVICE ALERT: WINV-104;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258293] SERVICE ALERT: WINV-105;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258336] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258339] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258342] SERVICE ALERT: WINV-104;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258345] SERVICE ALERT: WINV-101;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258348] SERVICE ALERT: WINV-103;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258351] SERVICE ALERT: WINV-104;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258354] SERVICE ALERT: WINV-105;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258397] SERVICE ALERT: WINV-102;Data Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258400] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258403] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586258406] SERVICE ALERT: WINV-101;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258409] SERVICE ALERT: WINV-103;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258412] SERVICE ALERT: WINV-104;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258466] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258469] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258472] SERVICE ALERT: WINV-104;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258527] SERVICE ALERT: WINV-102;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258530] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258533] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586258764] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 19% (5 min average)
[1586258765] SERVICE ALERT: WINV-102;Memory Usage;OK;HARD;3;Memory usage: total:298295.47 MB - used: 33208.63 MB (11%) - free: 265086.84 MB (89%)
[1586258768] SERVICE ALERT: WINV-103;Uptime;OK;HARD;3;System Uptime - 13 day(s) 14 hour(s) 7 minute(s)
[1586258823] SERVICE ALERT: WINV-101;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 783.01 Gb (84%) - free 148.31 Gb (16%)
[1586258826] SERVICE ALERT: WINV-102;OS Drive Space;OK;HARD;3;c:\ - total: 931.41 Gb - used: 489.72 Gb (53%) - free 441.69 Gb (47%)
[1586258829] SERVICE ALERT: WINV-104;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1586258884] SERVICE ALERT: WINV-101;Memory Usage;OK;HARD;3;Memory usage: total:98284.48 MB - used: 40391.63 MB (41%) - free: 57892.85 MB (59%)
[1586258887] SERVICE ALERT: WINV-102;Uptime;OK;HARD;3;System Uptime - 1 day(s) 13 hour(s) 31 minute(s)
[1586258890] SERVICE ALERT: WINV-104;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 787.76 Gb (85%) - free 143.56 Gb (15%)
[1586258945] SERVICE ALERT: WINV-101;OS Drive Space;OK;HARD;3;c:\ - total: 68.27 Gb - used: 44.18 Gb (65%) - free 24.08 Gb (35%)
[1586258948] SERVICE ALERT: WINV-103;CPU Load;OK;HARD;3;CPU Load 16% (5 min average)
[1586258951] SERVICE ALERT: WINV-104;Memory Usage;OK;HARD;3;Memory usage: total:129142.02 MB - used: 42270.28 MB (33%) - free: 86871.73 MB (67%)
[1586258957] SERVICE ALERT: WINV-105;Uptime;OK;HARD;3;System Uptime - 357 day(s) 20 hour(s) 1 minute(s)
[1586259006] SERVICE ALERT: WINV-101;Uptime;OK;HARD;3;System Uptime - 33 day(s) 5 hour(s) 8 minute(s)
[1586259009] SERVICE ALERT: WINV-103;Data Drive Space;OK;HARD;3;d:\ - total: 1862.39 Gb - used: 947.91 Gb (51%) - free 914.48 Gb (49%)
[1586259012] SERVICE ALERT: WINV-104;OS Drive Space;OK;HARD;3;c:\ - total: 68.27 Gb - used: 48.11 Gb (70%) - free 20.16 Gb (30%)
[1586259066] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1586259069] SERVICE ALERT: WINV-103;Memory Usage;OK;HARD;3;Memory usage: total:196589.32 MB - used: 64087.71 MB (33%) - free: 132501.61 MB (67%)
[1586259072] SERVICE ALERT: WINV-104;Uptime;OK;HARD;3;System Uptime - 193 day(s) 17 hour(s) 57 minute(s)
[1586259127] SERVICE ALERT: WINV-102;Data Drive Space;OK;HARD;3;d:\ - total: 1862.64 Gb - used: 1286.40 Gb (69%) - free 576.24 Gb (31%)
[1586259130] SERVICE ALERT: WINV-103;OS Drive Space;OK;HARD;3;c:\ - total: 68.27 Gb - used: 48.75 Gb (71%) - free 19.52 Gb (29%)
[1586259133] SERVICE ALERT: WINV-105;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1586259371] SERVICE ALERT: WINV-105;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 602.32 Gb (65%) - free 329.00 Gb (35%)
[1586259432] SERVICE ALERT: WINV-105;Memory Usage;OK;HARD;3;Memory usage: total:98284.48 MB - used: 40406.64 MB (41%) - free: 57877.84 MB (59%)
[1586259493] SERVICE ALERT: WINV-105;OS Drive Space;OK;HARD;3;c:\ - total: 231.90 Gb - used: 98.19 Gb (42%) - free 133.71 Gb (58%)
[1586259601] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 18% free (3773 MB out of 21455 MB)
[1586259661] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 15% free (3181 MB out of 21455 MB)
[1586259722] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 13% free (2695 MB out of 21455 MB)
[1586259782] SERVICE ALERT: vdb-101;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (2303 MB out of 21455 MB)
[1586260082] SERVICE ALERT: vdb-101;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 9% free (1875 MB out of 21455 MB)
[1586260382] SERVICE ALERT: vdb-101;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (2166 MB out of 21455 MB)
[1586260982] SERVICE FLAPPING ALERT: vdb-101;Swap Usage;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1586260982] SERVICE ALERT: vdb-101;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (2056 MB out of 21455 MB)
[1586261282] SERVICE ALERT: vdb-101;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (2161 MB out of 21455 MB)
[1586263494] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 31.98, 10.74, 4.20
[1586263554] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 47.52, 19.95, 7.82
[1586263614] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 46.10, 24.46, 10.13
[1586263674] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 70.11, 35.42, 14.81
[1586264517] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.11, 6.75, 5.33
[1586264578] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.85, 6.41, 5.30
[1586264638] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.03, 6.02, 5.23
[1586264698] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.41, 6.27, 5.35
[1586265781] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 2.8T of 3.6T
[1586265841] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 2.8T of 3.6T
[1586265901] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 2.8T of 3.6T
[1586265961] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 2.8T of 3.6T
[1586266982] SERVICE FLAPPING ALERT: vdb-101;Swap Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1586267398] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.81, 4.55, 4.92
[1586267698] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.14, 6.21, 5.56
[1586267758] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.57, 6.67, 5.75
[1586267818] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.39, 6.53, 5.76
[1586267878] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.32, 5.94, 5.61
[1586267882] SERVICE ALERT: vdb-101;Swap Usage;OK;HARD;4;SWAP OK - 22% free (4547 MB out of 21455 MB)
[1586269920] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;1;/media/data is at 91% capacity! 6.2T of 7.3T
[1586269980] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;2;/media/data is at 91% capacity! 6.2T of 7.3T
[1586270040] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;3;/media/data is at 91% capacity! 6.2T of 7.3T
[1586270100] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1586271144] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (2227 MB out of 11215 MB)
[1586271204] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (2150 MB out of 11215 MB)
[1586271265] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 19% free (2034 MB out of 11215 MB)
[1586271325] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;HARD;4;SWAP WARNING - 18% free (1956 MB out of 11215 MB)
[1586272674] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.57, 4.50, 17.92
[1586273278] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.68, 2.66, 4.08
[1586273574] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.96, 1.44, 7.33
[1586273582] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (4286 MB out of 21455 MB)
[1586273642] SERVICE ALERT: vdb-101;Swap Usage;OK;SOFT;2;SWAP OK - 21% free (4322 MB out of 21455 MB)
[1586274002] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (4276 MB out of 21455 MB)
[1586274062] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (4288 MB out of 21455 MB)
[1586274123] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 20% free (4193 MB out of 21455 MB)
[1586274183] SERVICE ALERT: vdb-101;Swap Usage;WARNING;HARD;4;SWAP WARNING - 20% free (4229 MB out of 21455 MB)
[1586277000] SERVICE ALERT: vonly-backup;/media/data Partition;OK;HARD;4;/media/data is at 86% capacity, 5.9T of 7.3T
[1586277244] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 3.26, 3.59, 4.66
[1586277445] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.74, 6.16, 3.25
[1586277505] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 10.67, 7.06, 3.74
[1586277565] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 11.87, 8.08, 4.30
[1586277625] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 12.60, 8.98, 4.85
[1586283061] SERVICE ALERT: vonly-database;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 3.1T of 3.6T
[1586283483] SERVICE ALERT: vdb-101;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (2100 MB out of 21455 MB)
[1586284444] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 6.14, 5.50, 5.06
[1586284504] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 5.16, 5.37, 5.05
[1586284564] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 6.51, 5.72, 5.19
[1586284624] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 6.92, 6.00, 5.32
[1586298724] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 2.19, 3.04, 4.78
[1586298925] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 6.81, 7.10, 7.74
[1586299324] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 8.86, 6.67, 5.40
[1586299384] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 9.16, 7.08, 5.62
[1586299444] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 7.27, 6.94, 5.66
[1586299504] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 7.71, 7.15, 5.82
[1586301183] SERVICE ALERT: vdb-101;Swap Usage;OK;HARD;4;SWAP OK - 98% free (20906 MB out of 21455 MB)
[1586301625] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.32, 4.53, 4.90
[1586302204] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 0.00, 0.55, 3.75
[1586303274] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.01, 8.31, 3.56
[1586303334] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 28.97, 12.97, 5.44
[1586303394] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 43.36, 19.86, 8.28
[1586303454] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 57.20, 28.12, 11.84
[1586303704] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 11.50, 10.09, 6.03
[1586303764] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 10.87, 10.21, 6.33
[1586303825] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 10.68, 10.34, 6.62
[1586303885] SERVICE FLAPPING ALERT: vdb-101;Current Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1586303885] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 10.53, 10.35, 6.86
[1586305554] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.62, 5.43, 16.49
[1586306454] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.12, 0.59, 6.55
[1586308385] SERVICE FLAPPING ALERT: vdb-101;Current Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1586323854] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 13.97, 11.34, 5.67
[1586323914] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 25.75, 15.46, 7.47
[1586323975] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 22.45, 16.54, 8.34
[1586324035] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 22.64, 17.63, 9.22
[1586324674] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1586324794] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 16% (5 min average)
[1586324914] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1586325034] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1586325835] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.12, 1.10, 5.95
[1586326261] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 89% capacity! 3.1T of 3.6T
[1586328611] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.86, 9.72, 10.00
[1586328911] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.11, 9.96, 10.01
[1586328971] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.24, 10.01, 10.03
[1586329031] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.04, 10.00, 10.02
[1586329091] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.12, 10.02, 10.02