[1580717709] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.78, 5.72, 5.01
[1580717769] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.49, 5.58, 5.01
[1580717829] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.50, 5.56, 5.03
[1580717889] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.76, 5.60, 5.08
[1580720289] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.12, 4.56, 4.90
[1580723289] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.07, 5.92, 5.33
[1580723349] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.33, 5.75, 5.30
[1580723409] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.86, 5.84, 5.36
[1580723469] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.07, 5.90, 5.41
[1580725269] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.09, 4.55, 4.85
[1580726469] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.29, 5.53, 5.13
[1580726529] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.84, 5.86, 5.27
[1580726589] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 4.99, 5.52, 5.19
[1580726649] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.13, 5.49, 5.20
[1580730202] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 256 processes with STATE = RSZDT
[1580730262] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1580730322] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 241 processes with STATE = RSZDT
[1580730682] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 261 processes with STATE = RSZDT
[1580730742] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 259 processes with STATE = RSZDT
[1580730803] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 260 processes with STATE = RSZDT
[1580730863] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 260 processes with STATE = RSZDT
[1580731463] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 242 processes with STATE = RSZDT
[1580764165] SERVICE ALERT: vdb-101;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (99 MB out of 975 MB)
[1580764765] SERVICE ALERT: vdb-101;Swap Usage;OK;HARD;4;SWAP OK - 18% free (174 MB out of 975 MB)
[1580768349] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.49, 4.71, 4.90
[1580769249] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.15, 5.10, 5.02
[1580769309] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.04, 5.07, 5.01
[1580769369] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 4.97, 5.10, 5.03
[1580769429] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 4.97, 5.08, 5.03
[1580770029] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.95, 4.93, 4.99
[1580770141] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.79, 4.77, 4.99
[1580770329] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 4.76, 5.05, 5.03
[1580770389] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 4.25, 4.87, 4.97
[1580770749] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.15, 5.10, 5.04
[1580770809] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.21, 5.18, 5.07
[1580770869] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.18, 5.12, 5.06
[1580770929] SERVICE FLAPPING ALERT: vonly-database;Current Load;STARTED; Service appears to have started flapping (22.9% change >= 20.0% threshold)
[1580770929] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.08, 5.10, 5.06
[1580771229] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.99, 4.85, 4.95
[1580771829] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.17, 5.37, 5.17
[1580771889] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.40, 5.65, 5.28
[1580771949] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.14, 5.70, 5.32
[1580772009] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.09, 5.49, 5.27
[1580772609] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.93, 4.78, 4.97
[1580776509] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 4.98, 5.08, 5.02
[1580776569] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.22, 5.11, 5.04
[1580776629] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.43, 5.20, 5.07
[1580776689] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 4.94, 5.06, 5.03
[1580776989] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.91, 4.82, 4.92
[1580777341] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 7.20, 6.47, 5.56
[1580777401] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 7.20, 6.61, 5.66
[1580777461] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 7.69, 6.87, 5.81
[1580777521] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 8.51, 7.23, 6.01
[1580782689] SERVICE FLAPPING ALERT: vonly-database;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580794765] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 15% free (137 MB out of 975 MB)
[1580794825] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 15% free (137 MB out of 975 MB)
[1580794885] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 15% free (137 MB out of 975 MB)
[1580794945] SERVICE ALERT: vdb-101;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (71 MB out of 975 MB)
[1580796489] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.50, 5.23, 5.02
[1580796549] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.13, 5.19, 5.02
[1580796609] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.23, 5.20, 5.03
[1580796669] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.65, 5.06, 5.00