[1584170829] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.00, 13.99, 6.94
[1584170889] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 19.61, 15.06, 7.77
[1584170914] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 85% (5 min average)
[1584170949] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 26.10, 17.56, 9.07
[1584171009] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 23.73, 18.53, 9.95
[1584171034] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 78% (5 min average)
[1584171609] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 26.84, 26.23, 18.24
[1584172027] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.24, 7.42, 5.73
[1584172087] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.68, 7.38, 5.82
[1584172147] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.95, 7.27, 5.89
[1584172207] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 10.30, 8.11, 6.26
[1584173300] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.96, 6.75, 5.41
[1584173361] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 7.70, 6.94, 5.56
[1584173421] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 8.66, 7.37, 5.79
[1584173482] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 10.80, 8.20, 6.18
[1584173969] SERVICE ALERT: vonly-database;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1584174033] SERVICE ALERT: vonly-database;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.05 ms
[1584175282] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 2.50, 3.15, 4.70
[1584175585] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 7.83, 5.55, 5.25
[1584175645] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 8.00, 6.05, 5.45
[1584175705] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.67, 5.74, 5.38
[1584175761] SERVICE ALERT: vonly-database;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1584175765] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.55, 5.70, 5.39
[1584175821] SERVICE ALERT: vonly-database;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1584176110] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.49, 15.70, 19.03
[1584176965] SERVICE FLAPPING ALERT: vonly-database;Current Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1584176965] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 1.56, 3.54, 4.60
[1584177310] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 26.39, 25.26, 21.81
[1584177995] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584178115] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 13% (5 min average)
[1584178380] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.17, 11.71, 10.55
[1584178440] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.36, 11.85, 10.67
[1584178500] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.27, 11.95, 10.78
[1584178560] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.87, 11.87, 10.83
[1584180360] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.55, 8.71, 9.93
[1584181510] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.66, 14.87, 19.18
[1584181765] SERVICE FLAPPING ALERT: vonly-database;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584183611] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.82, 29.88, 23.09
[1584184087] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 83% (5 min average)
[1584184207] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1584185111] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.79, 11.01, 18.73
[1584185407] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.42, 4.02, 4.83
[1584186360] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.69 Gb (80%) - free 185.63 Gb (20%)
[1584186480] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.84 Gb (80%) - free 185.48 Gb (20%)
[1584186600] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.99 Gb (80%) - free 185.33 Gb (20%)
[1584186911] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1584186911] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.10, 1.45, 5.85
[1584191111] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1584192120] SERVICE ALERT: vonly-backup;/media/data Partition;OK;HARD;4;/media/data is at 90% capacity, 6.2T of 7.3T
[1584210992] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584211112] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584211232] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584211832] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 14% (5 min average)
[1584228611] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.34, 28.41, 14.81
[1584228671] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 34.10, 28.38, 15.65
[1584228731] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 19.53, 25.69, 15.55
[1584228791] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.88, 22.18, 14.98
[1584229508] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.34, 6.20, 5.14
[1584229568] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.13, 6.43, 5.28
[1584229628] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.05, 6.21, 5.28
[1584229688] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 3.73, 5.48, 5.08
[1584229988] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.66, 4.83, 4.93
[1584230288] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.45, 5.40, 5.09
[1584230348] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.80, 5.33, 5.08
[1584230408] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.25, 5.23, 5.06
[1584230468] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.63, 4.99, 4.99
[1584230591] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.00, 23.21, 15.64
[1584230768] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.08, 5.68, 5.35
[1584230828] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.29, 6.68, 5.72
[1584230888] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.72, 6.26, 5.64
[1584230948] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.07, 6.69, 5.84
[1584235391] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.99, 11.71, 19.42
[1584235748] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.81, 2.25, 4.01
[1584236291] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.20, 0.84, 7.57
[1584245527] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 85% (5 min average)
[1584245647] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 71% (5 min average)
[1584254221] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;HARD;4;SWAP WARNING - 13% free (119 MB out of 975 MB)