[1584775458] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.93, 9.40, 4.25
[1584775518] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.75, 13.32, 5.91
[1584775578] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 22.92, 14.51, 6.78
[1584775638] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 24.07, 16.55, 7.97
[1584775938] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 40.51, 30.68, 16.35
[1584776316] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.00, 7.04, 5.10
[1584776376] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.41, 6.69, 5.10
[1584776436] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.30, 6.65, 5.19
[1584776496] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.39, 6.76, 5.31
[1584777422] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 16.42, 13.10, 10.04
[1584777482] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 17.23, 13.97, 10.54
[1584777542] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 16.40, 14.42, 10.91
[1584777602] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 16.31, 14.84, 11.28
[1584780603] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.93, 6.91, 9.56
[1584789642] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1584789696] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.56, 2.21, 4.19
[1584789738] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.97, 5.48, 18.27
[1584789762] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1584790354] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1584790474] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 69% (5 min average)
[1584790638] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.32, 0.55, 7.15
[1584792994] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1584793114] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 48% (5 min average)
[1584796930] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1584797059] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;2;No data was received from host!
[1584797179] SERVICE ALERT: WINV-106;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 585.53 Gb (63%) - free 345.79 Gb (37%)
[1584818682] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584818802] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584818922] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584819522] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 9% (5 min average)
[1584822984] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584823104] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1584824064] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.12, 5.79, 5.16
[1584824124] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.72, 5.76, 5.19
[1584824184] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 4.60, 5.44, 5.12
[1584824244] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 4.77, 5.32, 5.09
[1584824544] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.00, 4.56, 4.84
[1584833238] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 40.42, 21.83, 11.12
[1584833298] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 32.57, 23.19, 12.28
[1584833359] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 32.63, 25.04, 13.61
[1584833419] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.88, 28.58, 15.56
[1584833496] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.43, 6.54, 5.13
[1584833556] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 2.79, 5.57, 4.88
[1584833719] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.37, 13.61, 12.90
[1584834319] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.01, 2.94, 7.38
[1584834919] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 38.45, 26.77, 15.16
[1584834979] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 36.22, 28.10, 16.33
[1584835039] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 36.28, 29.94, 17.72
[1584835099] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1584835099] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.50, 29.94, 18.48
[1584835416] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.53, 6.72, 5.19
[1584835476] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.43, 6.24, 5.12
[1584835503] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.41, 12.14, 10.32
[1584835536] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.97, 6.89, 5.40
[1584835563] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.96, 12.04, 10.40
[1584835596] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.12, 7.21, 5.62
[1584835623] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.15, 12.10, 10.53
[1584835683] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.57, 12.22, 10.67
[1584837302] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1584839102] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 90% capacity! 6.2T of 7.3T
[1584839402] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1584840396] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.23, 3.30, 4.40
[1584840483] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.57, 8.84, 9.84
[1584840499] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.27, 8.14, 18.42
[1584841399] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.21, 1.31, 7.41
[1584846499] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584847344] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.75, 5.52, 5.05
[1584847404] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.56, 5.53, 5.08
[1584847464] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.98, 5.65, 5.15
[1584847524] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.79, 5.66, 5.19
[1584848124] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.62, 4.24, 4.76
[1584851523] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (11321 MB out of 57111 MB)
[1584851583] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (11337 MB out of 57111 MB)
[1584851643] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 20% free (11313 MB out of 57111 MB)
[1584851703] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 20% free (11287 MB out of 57111 MB)
[1584854922] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584855042] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)