[1585984488] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.40, 9.99, 4.50
[1585984548] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 17.50, 11.16, 5.23
[1585984608] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 21.71, 13.45, 6.39
[1585984669] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.67, 18.63, 8.61
[1585985212] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.08, 7.58, 5.58
[1585985272] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.95, 7.82, 5.78
[1585985332] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.50, 7.55, 5.81
[1585985392] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 10.48, 8.44, 6.23
[1585989169] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.65, 10.97, 19.34
[1585989469] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 54.35, 28.15, 23.38
[1585995769] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.98, 12.02, 19.31
[1585996069] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.86, 28.28, 23.87
[1585996354] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585996474] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1585996594] SERVICE ALERT: WINV-107;CPU Load;OK;HARD;3;CPU Load 7% (5 min average)
[1586002492] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.64, 3.96, 4.88
[1586002969] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.29, 5.36, 16.67
[1586003869] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.31, 0.83, 6.73
[1586004938] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1586005058] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586005178] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 97% (5 min average)
[1586005778] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 20% (5 min average)
[1586006570] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 19.85, 12.74, 7.65
[1586006630] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 13.40, 12.30, 7.83
[1586006690] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 15.02, 13.00, 8.35
[1586006750] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 20.16, 14.72, 9.23
[1586007592] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.66, 6.05, 5.31
[1586007650] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1586007650] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 26.53, 26.47, 19.70
[1586007652] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.20, 5.74, 5.24
[1586007712] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.26, 5.80, 5.30
[1586007772] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.03, 5.36, 5.18
[1586008250] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.15, 19.94, 19.67
[1586009857] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 7.23, 6.82, 5.52
[1586009917] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 7.20, 6.88, 5.63
[1586009977] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 7.26, 6.97, 5.74
[1586010037] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 7.26, 7.02, 5.83
[1586010172] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.23, 3.04, 4.69
[1586010350] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.73, 1.74, 6.04
[1586011237] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.28, 3.95, 4.77
[1586014250] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1586023565] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 90% capacity! 6.2T of 7.3T
[1586026778] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1586026898] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1586027450] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.33, 9.42, 4.25
[1586027510] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.57, 15.24, 6.56
[1586027570] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 34.01, 18.43, 8.20
[1586027630] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.98, 22.39, 10.19
[1586033930] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.34, 7.40, 15.28
[1586034830] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.36, 1.07, 6.22
[1586040230] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 21.01, 15.00, 9.49
[1586040290] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 17.95, 15.29, 9.94
[1586040350] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 16.79, 15.96, 10.54
[1586040410] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.76, 15.59, 10.76
[1586040472] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.70, 5.76, 5.12
[1586040532] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.12, 6.51, 5.42
[1586040592] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.69, 6.96, 5.66
[1586040652] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.64, 6.95, 5.75
[1586041252] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.91, 3.85, 4.83
[1586041910] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.41, 1.63, 5.94
[1586044010] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.43, 11.37, 5.10
[1586044070] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 30.25, 14.72, 6.66
[1586044130] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 34.86, 18.45, 8.43
[1586044191] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.27, 22.32, 10.38
[1586044791] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1586044791] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.67, 24.89, 18.24
[1586045091] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.64, 27.60, 21.25
[1586046291] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 15.60, 18.30, 19.72
[1586046353] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.08, 5.19, 5.09
[1586046413] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.52, 5.67, 5.28
[1586046473] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.44, 5.62, 5.28
[1586046533] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.40, 5.29, 5.19
[1586047191] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 17.02, 21.47, 20.79
[1586047491] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.73, 18.67, 19.90
[1586047791] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.23, 22.40, 21.25
[1586048392] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.79, 17.21, 19.43
[1586049592] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 20.24, 21.82, 20.79
[1586049665] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1586050192] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 16.69, 17.54, 19.15
[1586050492] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.42, 23.69, 21.55
[1586050792] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.48, 17.28, 19.52
[1586051933] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.36, 4.22, 4.90
[1586052892] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.66, 1.39, 6.10
[1586054165] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 90% capacity! 6.2T of 7.3T
[1586056792] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1586057627] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;No data was received from host!
[1586057749] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586057870] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1586058470] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 6% (5 min average)
[1586061011] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 24.42, 24.60, 24.99
[1586065511] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 24.95, 25.11, 25.01
[1586068211] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 25.68, 24.80, 24.95
[1586068511] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.26, 27.77, 26.15
[1586069270] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1586069390] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586069510] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)