[1584947661] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.33, 16.78, 7.88
[1584947722] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 29.59, 18.82, 9.13
[1584947782] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.13, 20.97, 10.50
[1584947842] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 40.04, 25.64, 12.76
[1584948758] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.56, 7.12, 5.59
[1584948819] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.01, 6.85, 5.59
[1584948879] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.70, 7.45, 5.86
[1584948939] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.62, 7.96, 6.14
[1584950044] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 13.52, 11.76, 10.10
[1584950104] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.49, 11.25, 10.03
[1584950164] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.31, 11.59, 10.22
[1584950224] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.57, 11.54, 10.29
[1584954724] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.66, 8.44, 9.55
[1584955324] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.68, 11.45, 10.42
[1584955384] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.63, 11.04, 10.34
[1584955445] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.55, 11.26, 10.46
[1584955505] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.56, 11.54, 10.60
[1584958342] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.41, 15.92, 19.99
[1584959705] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.36, 7.25, 9.25
[1584959842] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.86, 25.67, 19.31
[1584960034] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584960154] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)
[1584961505] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.13, 10.77, 10.07
[1584961565] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.98, 10.65, 10.08
[1584961625] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.16, 10.58, 10.09
[1584961685] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.31, 10.58, 10.12
[1584962842] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.00, 6.48, 18.11
[1584963340] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.20, 1.98, 4.18
[1584963742] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1584963742] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.19, 0.65, 7.14
[1584965585] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.96, 9.52, 9.92
[1584966185] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.11, 10.55, 10.06
[1584966245] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.93, 10.48, 10.07
[1584966305] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.54, 10.72, 10.17
[1584966365] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.00, 10.72, 10.21
[1584966665] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1584966665] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.95, 8.38, 9.34
[1584967042] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.39, 8.02, 3.94
[1584967102] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 15.80, 9.15, 4.57
[1584967162] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 28.03, 13.82, 6.46
[1584967222] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.24, 16.90, 7.97
[1584967840] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.39, 7.16, 5.11
[1584967900] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.19, 6.98, 5.17
[1584967960] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.53, 7.58, 5.49
[1584968020] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.75, 7.86, 5.71
[1584969220] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.09, 3.59, 4.97
[1584969274] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1584969394] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 64% (5 min average)
[1584969622] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.08, 3.83, 17.14
[1584970522] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.09, 0.75, 6.90
[1584971668] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.10, 6.05, 5.21
[1584971728] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.17, 6.07, 5.27
[1584971788] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.66, 5.98, 5.29
[1584971848] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.99, 6.00, 5.34
[1584972149] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 2.55, 4.41, 4.87
[1584972365] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584975149] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.82, 6.37, 5.20
[1584975209] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.53, 6.35, 5.27
[1584975269] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 7.18, 6.57, 5.42
[1584975329] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 7.80, 6.83, 5.58
[1584975622] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584983103] SERVICE ALERT: vonly-database;Swap Usage;OK;HARD;4;SWAP OK - 21% free (11449 MB out of 57111 MB)
[1584985224] SERVICE ALERT: vonly-database;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1584985284] SERVICE ALERT: vonly-database;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1584991529] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.26, 3.93, 4.77
[1584994473] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 86% (5 min average)
[1584994593] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584994713] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 10% (5 min average)
[1584994830] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.71, 6.50, 5.32
[1584994890] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.93, 6.58, 5.42
[1584994950] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.44, 6.50, 5.47
[1584995010] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 7.11, 6.69, 5.60
[1584997410] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 2.10, 3.17, 4.71
[1585006714] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 83% (5 min average)
[1585006822] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.93, 13.40, 8.14
[1585006834] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1585006882] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 21.33, 15.00, 9.01
[1585006942] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 23.45, 16.81, 10.01
[1585006954] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.51 Gb - used: 745.65 Gb (80%) - free 185.86 Gb (20%)
[1585007002] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 26.73, 18.55, 11.01
[1585007020] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.07, 6.74, 5.45
[1585007074] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.51 Gb - used: 745.92 Gb (80%) - free 185.59 Gb (20%)
[1585007080] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.86, 6.31, 5.39
[1585007140] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.18, 6.66, 5.56
[1585007194] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.51 Gb - used: 746.17 Gb (80%) - free 185.34 Gb (20%)
[1585007200] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.13, 6.87, 5.71
[1585007465] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.96, 12.19, 10.24
[1585007525] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.25, 12.46, 10.46
[1585007585] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.35, 12.60, 10.63
[1585007602] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 39.86, 33.09, 21.41
[1585007645] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.32, 12.68, 10.78
[1585012145] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.35, 9.64, 9.96
[1585012715] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;1;c:\ - total: 68.27 Gb - used: 54.92 Gb (80%) - free 13.35 Gb (20%)
[1585012745] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 9.83, 10.02, 10.02
[1585012805] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.37, 10.15, 10.07
[1585012835] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;2;c:\ - total: 68.27 Gb - used: 54.99 Gb (81%) - free 13.28 Gb (19%)
[1585012865] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.46, 10.23, 10.10
[1585012900] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.61, 4.52, 4.87
[1585012925] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.02, 10.21, 10.11
[1585012955] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;c:\ - total: 68.27 Gb - used: 55.08 Gb (81%) - free 13.18 Gb (19%)
[1585013225] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1585013225] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.14, 8.59, 9.46
[1585013302] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.23, 4.87, 14.94
[1585014202] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.42, 1.37, 6.33
[1585015595] SERVICE ALERT: WINV-103;Memory Usage;WARNING;SOFT;1;Memory usage: total:234403.37 MB - used: 189469.29 MB (81%) - free: 44934.08 MB (19%)
[1585015715] SERVICE ALERT: WINV-103;Memory Usage;WARNING;SOFT;2;Memory usage: total:234403.37 MB - used: 191038.84 MB (82%) - free: 43364.53 MB (18%)
[1585015835] SERVICE ALERT: WINV-103;Memory Usage;WARNING;HARD;3;Memory usage: total:234403.37 MB - used: 192641.79 MB (82%) - free: 41761.58 MB (18%)
[1585017635] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;HARD;3;Memory usage: total:234403.37 MB - used: 213897.93 MB (91%) - free: 20505.45 MB (9%)
[1585018925] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585020755] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 61.53 Gb (90%) - free 6.74 Gb (10%)
[1585024363] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585024364] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;No data was received from host!
[1585024486] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1028.06 Gb (55%) - free 834.33 Gb (45%)
[1585024964] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;HARD;3;c:\ - total: 68.27 Gb - used: 63.44 Gb (93%) - free 4.82 Gb (7%)
[1585031815] SERVICE ALERT: WINV-103;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585031935] SERVICE ALERT: WINV-103;Data Drive Space;OK;SOFT;2;d:\ - total: 1862.39 Gb - used: 1034.84 Gb (56%) - free 827.55 Gb (44%)
[1585032763] SERVICE ALERT: WINV-103;Uptime;WARNING;SOFT;1;No data was received from host!
[1585032883] SERVICE ALERT: WINV-103;Uptime;OK;SOFT;2;System Uptime - 41 day(s) 15 hour(s) 58 minute(s)