[1580374798] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.91, 6.35, 5.14
[1580374858] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 7.08, 6.49, 5.26
[1580374918] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.37, 6.42, 5.32
[1580374978] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.59, 6.49, 5.41
[1580383078] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.14, 4.32, 4.83
[1580385778] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.49, 5.62, 5.05
[1580385838] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 4.93, 5.44, 5.02
[1580385898] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.11, 5.37, 5.02
[1580385958] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.42, 5.41, 5.05
[1580386258] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.83, 4.66, 4.85
[1580435758] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.02, 5.52, 5.13
[1580435818] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 4.58, 5.32, 5.08
[1580435878] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 4.34, 5.11, 5.02
[1580435938] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.04, 4.86, 4.94
[1580436838] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.58, 5.63, 5.20
[1580436898] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.38, 5.56, 5.20
[1580436958] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.75, 5.63, 5.24
[1580437019] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.74, 5.65, 5.28
[1580438519] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.16, 4.22, 4.75