[1582877611] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 21.50, 13.22, 5.58
[1582877616] SERVICE ALERT: WINV-101;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.10 Gb (80%) - free 186.22 Gb (20%)
[1582877671] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 13.49, 12.43, 5.79
[1582877732] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 11.52, 12.05, 6.08
[1582877736] SERVICE ALERT: WINV-101;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.10 Gb (80%) - free 186.22 Gb (20%)
[1582877792] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.71, 12.39, 6.56
[1582877856] SERVICE ALERT: WINV-101;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.10 Gb (80%) - free 186.21 Gb (20%)
[1582879292] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.10, 22.81, 16.12
[1582879597] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.97, 5.91, 4.10
[1582879657] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 3.90, 5.12, 3.94
[1582879711] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.16, 5.50, 5.04
[1582879711] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.35, 11.43, 10.37
[1582879771] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.49, 5.31, 5.00
[1582879771] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.41, 11.43, 10.44
[1582879831] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.23, 11.62, 10.57
[1582879891] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.84, 11.90, 10.73
[1582880131] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.79, 6.14, 5.42
[1582880191] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.67, 6.11, 5.46
[1582880251] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.80, 6.17, 5.52
[1582880311] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.88, 6.50, 5.69
[1582885240] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1582885360] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 64% (5 min average)
[1582885592] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.62, 12.48, 18.20
[1582887092] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.16, 27.00, 21.82
[1582887811] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.58, 4.62, 4.95
[1582888111] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.87, 5.68, 5.31
[1582888171] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.54, 5.59, 5.31
[1582888231] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.32, 5.68, 5.37
[1582888291] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 3.67, 5.13, 5.19
[1582889011] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 81% capacity! 5.5T of 7.3T
[1582889071] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 81% capacity! 5.5T of 7.3T
[1582889131] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;3;/mnt/data01 is at 81% capacity! 5.5T of 7.3T
[1582889191] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 81% capacity! 5.5T of 7.3T
[1582889191] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.55, 4.65, 4.94
[1582889491] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.06, 7.41, 9.09
[1582890691] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.50, 5.58, 5.14
[1582890751] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.18, 5.05, 4.98
[1582891111] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.38, 5.70, 5.24
[1582891171] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.15, 5.87, 5.32
[1582891231] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.45, 5.40, 5.19
[1582891291] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1582891291] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 3.83, 5.07, 5.09
[1582891592] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.49, 4.19, 4.79
[1582891712] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1582891832] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1582891952] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 87% (5 min average)
[1582892072] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;2;CPU Load 86% (5 min average)
[1582892192] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.62, 4.38, 18.59
[1582892192] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1582892697] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582892757] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1582892792] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 13% (5 min average)
[1582893093] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.68, 1.28, 7.63
[1582897292] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582903111] SERVICE ALERT: vdb-staging;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 31.13, 14.47, 7.59
[1582903171] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 14.64, 12.70, 7.41
[1582903231] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 9.99, 11.69, 7.39
[1582903291] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 5.23, 9.97, 7.06
[1582904191] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 0.69, 2.62, 4.98
[1582913709] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.58 Gb (80%) - free 185.74 Gb (20%)
[1582913829] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.67 Gb (80%) - free 185.65 Gb (20%)
[1582913949] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.78 Gb (80%) - free 185.54 Gb (20%)
[1582914549] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 742.96 Gb (80%) - free 188.35 Gb (20%)
[1582919792] SERVICE ALERT: vdb-103;/mnt/data01 Partition;CRITICAL;HARD;4;/mnt/data01 is at 91% capacity! 6.2T of 7.3T
[1582932393] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.52, 11.84, 6.49
[1582932453] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 24.33, 14.36, 7.65
[1582932514] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 17.71, 14.44, 8.11
[1582932574] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 22.33, 16.22, 9.11
[1582933474] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.17, 3.72, 7.45
[1582937075] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.20, 11.34, 5.63
[1582937135] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 19.79, 13.15, 6.61
[1582937195] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 27.08, 16.31, 8.10
[1582937255] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.32, 20.03, 9.87
[1582937555] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1582937555] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 28.53, 24.59, 14.58
[1582937855] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 49.63, 39.79, 23.66
[1582938692] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.98, 5.43, 5.03
[1582938753] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.80, 5.57, 5.10
[1582938813] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.92, 6.12, 5.31
[1582938873] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.23, 6.75, 5.59
[1582942173] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.98, 4.09, 4.75
[1582942355] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.48, 14.53, 19.43
[1582943555] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.53, 1.03, 6.57
[1582948355] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1582954891] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 1.93, 8.33, 5.85
[1582954951] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 1.51, 7.02, 5.55
[1582955011] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 1.30, 5.96, 5.28
[1582955071] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 1.23, 5.09, 5.02
[1582956881] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1582957001] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 16% (5 min average)
[1582958071] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 1.04, 3.02, 4.68
[1582958921] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1582959041] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 36% (5 min average)
[1582960171] SERVICE ALERT: vdb-staging;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 33.15, 21.62, 10.10
[1582960231] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 16.62, 19.56, 10.15
[1582960291] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 6.19, 16.02, 9.52
[1582960351] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 2.43, 13.14, 8.94
[1582961036] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes with STATE = RSZDT
[1582961036] SERVICE FLAPPING ALERT: vdb-staging;Current Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1582961036] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 0.02, 1.75, 4.69
[1582961096] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1582961353] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 6.22, 7.49, 4.81
[1582961413] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 4.55, 6.79, 4.74
[1582961473] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 3.89, 6.20, 4.67
[1582961533] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.19, 5.85, 4.64
[1582961794] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582961917] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1582961953] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 11.81, 8.62, 6.45
[1582962013] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 12.04, 9.28, 6.81
[1582962074] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 10% (5 min average)
[1582962074] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 12.71, 9.96, 7.20
[1582962133] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 6.77, 6.08, 5.20
[1582962134] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 12.32, 10.35, 7.51
[1582962193] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 5.97, 5.96, 5.21
[1582962253] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.62, 5.85, 5.22
[1582962313] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.31, 5.72, 5.21
[1582962613] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 16.24, 8.85, 6.40