[1586070110] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1586071192] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.04, 20.85, 9.76
[1586071252] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 31.61, 23.56, 11.42
[1586071312] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.22, 24.57, 12.54
[1586071372] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 26.78, 24.50, 13.26
[1586071672] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.80, 26.50, 17.38
[1586072033] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.11, 8.09, 6.01
[1586072093] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 10.20, 8.76, 6.38
[1586072153] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.90, 8.36, 6.39
[1586072213] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 10.31, 9.01, 6.73
[1586073730] SERVICE ALERT: WINV-105;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586073730] SERVICE ALERT: WINV-105;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586073732] SERVICE ALERT: WINV-105;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586073779] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586073860] SERVICE ALERT: WINV-105;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586073860] SERVICE ALERT: WINV-105;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586073862] SERVICE ALERT: WINV-105;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586073878] SERVICE ALERT: WINV-104;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586073888] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586073909] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586073981] SERVICE ALERT: WINV-105;Data Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586073990] SERVICE ALERT: WINV-105;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586073990] SERVICE ALERT: WINV-105;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586073992] SERVICE ALERT: WINV-105;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074008] SERVICE ALERT: WINV-104;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074018] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074039] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074111] SERVICE ALERT: WINV-105;Data Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074132] SERVICE ALERT: WINV-101;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074138] SERVICE ALERT: WINV-104;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074142] SERVICE ALERT: WINV-101;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074148] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074228] SERVICE ALERT: WINV-104;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074228] SERVICE ALERT: WINV-104;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074241] SERVICE ALERT: WINV-105;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074262] SERVICE ALERT: WINV-101;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074271] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074272] SERVICE ALERT: WINV-101;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074320] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074330] SERVICE ALERT: WINV-103;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074358] SERVICE ALERT: WINV-104;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074358] SERVICE ALERT: WINV-104;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074364] SERVICE ALERT: WINV-104;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074366] SERVICE ALERT: WINV-101;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074392] SERVICE ALERT: WINV-101;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074401] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074402] SERVICE ALERT: WINV-101;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074411] SERVICE ALERT: WINV-103;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074441] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074450] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074460] SERVICE ALERT: WINV-103;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074466] SERVICE ALERT: WINV-103;Data Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074482] SERVICE ALERT: WINV-101;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586074488] SERVICE ALERT: WINV-104;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074488] SERVICE ALERT: WINV-104;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074531] SERVICE ALERT: WINV-103;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074541] SERVICE ALERT: WINV-103;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074571] SERVICE ALERT: WINV-103;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074580] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074590] SERVICE ALERT: WINV-103;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586074596] SERVICE ALERT: WINV-103;Data Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074612] SERVICE ALERT: WINV-101;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1586074671] SERVICE ALERT: WINV-103;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1586081779] SERVICE ALERT: vinfra-200;MySQL Connection Time;WARNING;SOFT;1;WARNING - 3.03 seconds to connect as nagios
[1586081839] SERVICE ALERT: vinfra-200;MySQL Connection Time;OK;SOFT;2;OK - 0.01 seconds to connect as nagios
[1586087796] SERVICE ALERT: WINV-107;Data Drive Space;CRITICAL;HARD;3;d:\ - total: 931.51 Gb - used: 839.06 Gb (90%) - free 92.45 Gb (10%)
[1586088114] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.68, 2.43, 4.27
[1586088172] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.38, 1.95, 15.69
[1586089073] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.13, 1.79, 6.92
[1586125373] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.13, 8.04, 3.60
[1586125433] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 20.65, 11.03, 4.92
[1586125493] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 23.77, 13.45, 6.12
[1586125553] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 28.41, 16.89, 7.79
[1586126153] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.10, 29.76, 18.89
[1586126214] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.67, 8.47, 6.04
[1586126274] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.85, 8.69, 6.26
[1586126334] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.75, 8.85, 6.47
[1586126394] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.05, 8.20, 6.39
[1586128853] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.01, 15.13, 19.42
[1586134194] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.45, 4.55, 5.00
[1586134494] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.25, 5.10, 5.09
[1586134554] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.75, 4.88, 5.02
[1586134614] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.66, 4.90, 5.02
[1586134674] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.75, 4.55, 4.88
[1586137553] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.40, 1.51, 6.39
[1586137921] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 2.8T of 3.6T
[1586137981] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 2.8T of 3.6T
[1586138041] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 2.8T of 3.6T
[1586138101] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 2.8T of 3.6T
[1586145311] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 23.94, 24.18, 24.79
[1586156101] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 73% capacity, 2.5T of 3.6T