[1583824990] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583825110] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1583825155] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.98, 9.83, 4.86
[1583825215] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 17.19, 11.03, 5.57
[1583825275] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 20.77, 13.11, 6.62
[1583825335] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 21.66, 14.81, 7.62
[1583825830] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1583825950] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 33% (5 min average)
[1583826070] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583826190] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1583826535] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 40.24, 31.50, 22.36
[1583826587] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.28, 7.66, 5.81
[1583826647] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.48, 7.07, 5.72
[1583826707] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.67, 6.79, 5.71
[1583826767] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.29, 7.14, 5.89
[1583828035] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.12, 14.57, 19.69
[1583828335] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1583828335] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 48.86, 32.85, 25.81
[1583828388] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 13.41, 11.47, 10.20
[1583828448] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 14.18, 12.04, 10.48
[1583828508] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 14.28, 12.49, 10.73
[1583828568] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 9.69, 11.49, 10.49
[1583830435] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.51, 14.98, 19.57
[1583830668] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 10.68, 9.40, 9.85
[1583830735] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.27, 20.99, 20.10
[1583830968] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.36, 10.40, 10.19
[1583831028] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.52, 10.21, 10.14
[1583831035] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.20, 19.25, 19.58
[1583831088] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 9.35, 10.01, 10.07
[1583831148] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.97, 9.76, 9.98
[1583831448] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 13.09, 10.79, 10.26
[1583831508] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 14.03, 11.41, 10.51
[1583831568] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 14.11, 11.96, 10.76
[1583831628] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.99, 12.07, 10.87
[1583831635] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 29.01, 23.45, 20.04
[1583832235] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.95, 14.85, 18.24
[1583832535] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 45.48, 24.80, 20.67
[1583833135] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.10, 15.24, 19.01
[1583833435] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.31, 32.06, 25.75
[1583834635] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.65, 16.70, 19.79
[1583835835] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 46.28, 26.20, 20.48
[1583836135] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.69, 17.34, 18.40
[1583836428] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.70, 8.36, 10.00
[1583838167] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.17, 4.00, 4.99
[1583838467] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.48, 5.76, 5.42
[1583838527] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.08, 5.39, 5.31
[1583838535] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.67, 23.85, 17.44
[1583838587] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.06, 5.33, 5.30
[1583838647] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.41, 5.05, 5.20
[1583838835] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.19, 17.95, 17.14
[1583838947] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.20, 4.42, 4.92
[1583839248] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.46, 5.09, 5.04
[1583839308] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.35, 5.66, 5.24
[1583839368] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.10, 5.32, 5.14
[1583839428] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (24.2% change >= 20.0% threshold)
[1583839428] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.22, 5.77, 5.30
[1583839435] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 31.93, 23.68, 19.39
[1583840028] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.59, 4.54, 4.92
[1583840035] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.34, 13.41, 17.10
[1583840328] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.03, 5.78, 5.32
[1583840335] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 24.27, 22.21, 20.19
[1583840388] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.69, 6.14, 5.46
[1583840448] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.41, 6.28, 5.57
[1583840508] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 3.57, 5.51, 5.35
[1583840635] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.25, 17.33, 19.24
[1583840808] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.68, 4.18, 4.77
[1583840935] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.43, 21.67, 19.93
[1583841108] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.92, 5.70, 5.22
[1583841168] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.01, 5.57, 5.20
[1583841228] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.68, 10.49, 10.15
[1583841228] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.27, 5.21, 5.10
[1583841288] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.72, 10.26, 10.09
[1583841288] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.93, 5.29, 5.14
[1583841348] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 9.73, 10.15, 10.07
[1583841408] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 9.44, 9.97, 10.01
[1583841535] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.67, 12.39, 17.34
[1583841589] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.79, 4.16, 4.65
[1583842009] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.86, 9.94, 9.99
[1583842081] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1583842201] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1583843035] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.78, 2.55, 7.52
[1583843521] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1583843642] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583843762] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1583844362] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1583847289] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583847536] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1583882635] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 24.82, 13.08, 7.01
[1583882695] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 33.13, 17.14, 8.76
[1583882755] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 38.65, 21.65, 10.84
[1583882815] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 44.54, 25.97, 12.97
[1583883115] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.64, 15.63, 12.58
[1583883715] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.12, 2.99, 7.29
[1583884315] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.05, 7.52, 6.56
[1583884375] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 16.27, 9.19, 7.20
[1583884435] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;3;OK - load average: 14.15, 9.93, 7.58
[1583884495] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.65, 11.71, 8.33
[1583884555] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 19.95, 13.34, 9.10
[1583884615] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 28.59, 16.75, 10.52
[1583884676] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1583884676] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 28.82, 19.02, 11.70
[1583884976] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 21.37, 25.18, 16.76
[1583885389] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.07, 6.75, 5.39
[1583885449] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.41, 6.75, 5.47
[1583885509] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.56, 6.68, 5.53
[1583885569] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.47, 7.14, 5.77
[1583889176] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.77, 11.95, 18.57
[1583889476] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.76, 24.09, 21.46
[1583890076] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.52, 17.18, 19.63
[1583890669] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.64, 3.21, 4.39
[1583891276] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.16, 0.97, 7.12
[1583896076] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583909443] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583909563] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 61% (5 min average)