[1581149582] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.62, 4.56, 4.83
[1581154787] SERVICE ALERT: vdb-102;PING;CRITICAL;SOFT;1;(Return code of 255 for service 'PING' on host 'vdb-102' was out of bounds)
[1581154788] SERVICE ALERT: vdb-102;/mnt/data01 Partition;CRITICAL;SOFT;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1581154790] HOST ALERT: vdb-102;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.4.102)
[1581154791] SERVICE ALERT: vdb-102;Current Users;CRITICAL;HARD;1;CHECK_NRPE: Socket timeout after 30 seconds.
[1581154850] SERVICE ALERT: vdb-102;/mnt/data01 Partition;CRITICAL;HARD;4;(Return code of 255 for service '/mnt/data01 Partition' on host 'vdb-102' was out of bounds)
[1581154850] SERVICE ALERT: vdb-102;PING;CRITICAL;HARD;4;(Return code of 255 for service 'PING' on host 'vdb-102' was out of bounds)
[1581154853] HOST ALERT: vdb-102;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.4.102)
[1581154865] SERVICE ALERT: vdb-102;SSH;CRITICAL;HARD;1;connect to address 192.168.4.102 and port 22: No route to host
[1581154891] SERVICE ALERT: vdb-102;Current Load;CRITICAL;HARD;1;(Return code of 255 for service 'Current Load' on host 'vdb-102' was out of bounds)
[1581154909] SERVICE ALERT: vdb-102;Root Partition;CRITICAL;HARD;1;(Return code of 255 for service 'Root Partition' on host 'vdb-102' was out of bounds)
[1581154916] HOST ALERT: vdb-102;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.4.102)
[1581154976] SERVICE ALERT: vdb-102;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'vdb-102' was out of bounds)
[1581154979] HOST ALERT: vdb-102;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.4.102)
[1581155042] HOST ALERT: vdb-102;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.4.102)
[1581155097] SERVICE ALERT: vonly-database;/media/data Partition;CRITICAL;SOFT;1;(Return code of 255 for service '/media/data Partition' on host 'vonly-database' was out of bounds)
[1581155100] HOST ALERT: vonly-database;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.3.60)
[1581155105] HOST ALERT: vdb-102;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.4.102)
[1581155112] SERVICE ALERT: vonly-database;Current Load;CRITICAL;HARD;1;(Return code of 255 for service 'Current Load' on host 'vonly-database' was out of bounds)
[1581155148] SERVICE ALERT: vonly-database;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'vonly-database' was out of bounds)
[1581155160] SERVICE ALERT: vonly-database;/media/data Partition;CRITICAL;HARD;4;(Return code of 255 for service '/media/data Partition' on host 'vonly-database' was out of bounds)
[1581155163] HOST ALERT: vonly-database;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.3.60)
[1581155168] HOST ALERT: vdb-102;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.4.102)
[1581155223] SERVICE ALERT: vonly-database;Root Partition;CRITICAL;HARD;1;(Return code of 255 for service 'Root Partition' on host 'vonly-database' was out of bounds)
[1581155226] HOST ALERT: vonly-database;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.3.60)
[1581155229] SERVICE ALERT: vonly-database;PING;CRITICAL;HARD;1;(Return code of 255 for service 'PING' on host 'vonly-database' was out of bounds)
[1581155231] HOST ALERT: vdb-102;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.4.102)
[1581155262] SERVICE ALERT: vonly-database;SSH;CRITICAL;HARD;1;connect to address 192.168.3.60 and port 22: No route to host
[1581155273] SERVICE ALERT: vdb-102;Total Processes;OK;HARD;4;PROCS OK: 202 processes with STATE = RSZDT
[1581155277] HOST ALERT: vdb-102;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.32 ms
[1581155289] HOST ALERT: vonly-database;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.3.60)
[1581155295] SERVICE ALERT: vonly-database;Swap Usage;CRITICAL;HARD;1;(Return code of 255 for service 'Swap Usage' on host 'vonly-database' was out of bounds)
[1581155301] SERVICE ALERT: vonly-database;Total Processes;CRITICAL;HARD;1;(Return code of 255 for service 'Total Processes' on host 'vonly-database' was out of bounds)
[1581155352] HOST ALERT: vonly-database;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.3.60)
[1581155361] SERVICE ALERT: vdb-102;Current Users;OK;SOFT;1;USERS OK - 2 users currently logged in
[1581155378] SERVICE ALERT: vdb-102;Swap Usage;OK;HARD;4;SWAP OK - 98% free (55806 MB out of 57111 MB)
[1581155415] HOST ALERT: vonly-database;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.3.60)
[1581155448] SERVICE ALERT: vdb-102;/mnt/data01 Partition;OK;HARD;4;/media/data is at 84% capacity, 2.9T of 3.6T
[1581155451] SERVICE ALERT: vdb-102;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1581155462] SERVICE ALERT: vdb-102;SSH;OK;SOFT;1;SSH OK - OpenSSH_6.6.1p1 Ubuntu-2ubuntu2.8 (protocol 2.0)
[1581155478] HOST ALERT: vonly-database;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.3.60)
[1581155488] SERVICE ALERT: vdb-102;Current Load;OK;SOFT;1;OK - load average: 0.00, 0.01, 0.05
[1581155506] SERVICE ALERT: vdb-102;Root Partition;OK;SOFT;1;/ is at 67% capacity, 140G of 220G
[1581155541] HOST ALERT: vonly-database;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.3.60)
[1581155604] HOST ALERT: vonly-database;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.3.60)
[1581155667] HOST ALERT: vonly-database;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.3.60)
[1581157811] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 0.08, 0.17, 0.21
[1581157815] HOST ALERT: vonly-database;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.26 ms
[1581157847] SERVICE ALERT: vonly-database;Current Users;OK;HARD;4;USERS OK - 3 users currently logged in
[1581157857] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 84% capacity, 2.9T of 3.6T
[1581157920] SERVICE ALERT: vonly-database;Root Partition;OK;HARD;4;/ is at 67% capacity, 140G of 220G
[1581157931] SERVICE ALERT: vonly-database;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1581157959] SERVICE ALERT: vonly-database;SSH;OK;HARD;4;SSH OK - OpenSSH_6.6.1p1 Ubuntu-2ubuntu2.8 (protocol 2.0)
[1581157992] SERVICE ALERT: vonly-database;Swap Usage;OK;HARD;4;SWAP OK - 98% free (55829 MB out of 57111 MB)
[1581158000] SERVICE ALERT: vonly-database;Total Processes;OK;HARD;4;PROCS OK: 203 processes with STATE = RSZDT
[1581158078] SERVICE ALERT: vdb-102;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 1% free (0 MB out of 975 MB)
[1581158138] SERVICE ALERT: vdb-102;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 1% free (0 MB out of 975 MB)
[1581158198] SERVICE ALERT: vdb-102;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 1% free (0 MB out of 975 MB)
[1581158258] SERVICE ALERT: vdb-102;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 1% free (0 MB out of 975 MB)
[1581161711] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 7.78, 7.51, 5.43
[1581161771] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 7.34, 7.39, 5.51
[1581161831] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.94, 7.29, 5.60
[1581161891] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 7.01, 7.27, 5.70
[1581163382] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 7.72, 6.58, 5.32
[1581163442] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 7.84, 6.84, 5.49
[1581163502] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 8.02, 7.04, 5.64
[1581163562] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.93, 6.90, 5.68
[1581163691] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.09, 3.68, 4.96
[1581163991] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 4.70, 4.56, 5.01
[1581164051] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 3.95, 4.37, 4.92
[1581164411] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.15, 5.36, 5.19
[1581164471] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.59, 5.36, 5.20
[1581164531] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.14, 5.28, 5.19
[1581164591] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 3.86, 4.93, 5.07
[1581164891] SERVICE FLAPPING ALERT: vonly-database;Current Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1581164891] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.47, 4.07, 4.65
[1581165176] SERVICE ALERT: vdb-102;Total Processes;CRITICAL;SOFT;1;(Return code of 255 for service 'Total Processes' on host 'vdb-102' was out of bounds)
[1581165179] HOST ALERT: vdb-102;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.4.202)
[1581165239] SERVICE ALERT: vdb-102;Total Processes;CRITICAL;HARD;4;(Return code of 255 for service 'Total Processes' on host 'vdb-102' was out of bounds)
[1581165242] HOST ALERT: vdb-102;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.4.202)
[1581165264] SERVICE ALERT: vdb-102;Current Users;CRITICAL;HARD;1;(Return code of 255 for service 'Current Users' on host 'vdb-102' was out of bounds)
[1581165305] HOST ALERT: vdb-102;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.4.202)
[1581165352] SERVICE ALERT: vdb-102;PING;CRITICAL;HARD;1;(Return code of 255 for service 'PING' on host 'vdb-102' was out of bounds)
[1581165352] SERVICE ALERT: vdb-102;/mnt/data01 Partition;CRITICAL;HARD;1;(Return code of 255 for service '/mnt/data01 Partition' on host 'vdb-102' was out of bounds)
[1581165365] SERVICE ALERT: vdb-102;SSH;CRITICAL;HARD;1;connect to address 192.168.4.202 and port 22: No route to host
[1581165365] HOST ALERT: vdb-102;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.4.202)
[1581165391] SERVICE ALERT: vdb-102;Current Load;CRITICAL;HARD;1;(Return code of 255 for service 'Current Load' on host 'vdb-102' was out of bounds)
[1581165409] SERVICE ALERT: vdb-102;Root Partition;CRITICAL;HARD;1;(Return code of 255 for service 'Root Partition' on host 'vdb-102' was out of bounds)
[1581165428] HOST ALERT: vdb-102;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.4.202)
[1581165491] HOST ALERT: vdb-102;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.4.202)
[1581165554] HOST ALERT: vdb-102;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.4.202)
[1581165617] HOST ALERT: vdb-102;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.4.202)
[1581165680] HOST ALERT: vdb-102;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.4.202)
[1581165743] HOST ALERT: vdb-102;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.4.202)
[1581170591] SERVICE FLAPPING ALERT: vonly-database;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581175322] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.80, 5.26, 3.58
[1581175382] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.95, 5.53, 3.79
[1581175442] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.55, 4.99, 3.71
[1581175562] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.69, 4.71, 4.95
[1581181347] HOST ALERT: vdb-102;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.30 ms
[1581181562] SERVICE ALERT: vdb-102;SSH;OK;SOFT;1;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.4 (protocol 2.0)
[1581187091] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 18.49, 14.60, 8.12
[1581187151] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 19.55, 15.55, 8.85
[1581187211] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 21.39, 16.82, 9.70
[1581187271] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 23.62, 18.28, 10.65
[1581187571] SERVICE ALERT: vonly-database;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 41.87, 30.48, 17.79
[1581214871] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 16.95, 17.37, 19.18
[1581221162] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.79, 7.06, 5.60
[1581221222] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.39, 6.90, 5.64
[1581221282] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.90, 6.68, 5.64
[1581221342] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.28, 6.69, 5.71
[1581222371] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 1.70, 3.04, 4.64
[1581227942] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.36, 3.65, 4.75