[1584343333] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 20.38, 15.04, 7.94
[1584343393] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 23.99, 16.66, 8.92
[1584343453] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 24.16, 18.08, 9.90
[1584343514] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 29.37, 21.05, 11.45
[1584343814] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.19, 29.94, 17.78
[1584344108] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.96, 6.48, 5.17
[1584344168] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.93, 6.93, 5.39
[1584344228] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.73, 7.11, 5.56
[1584344288] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.21, 6.54, 5.46
[1584345848] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584345968] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1584348092] SERVICE ALERT: vonly-database;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1584348106] SERVICE ALERT: vonly-database;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1584348152] SERVICE ALERT: vonly-database;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1584348170] SERVICE ALERT: vonly-database;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.05 ms
[1584351721] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.55, 10.92, 10.22
[1584351781] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.40, 10.79, 10.22
[1584351841] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 9.95, 10.62, 10.20
[1584351901] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.04, 10.54, 10.20
[1584352801] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.09, 9.40, 9.95
[1584354914] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.65, 14.81, 19.31
[1584354994] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1584355114] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1584355814] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 24.97, 24.26, 20.93
[1584357488] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.50, 3.86, 4.86
[1584357914] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.72, 10.40, 18.08
[1584358021] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 5.5T of 7.3T
[1584358081] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 5.5T of 7.3T
[1584358141] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 5.5T of 7.3T
[1584358201] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 5.5T of 7.3T
[1584358814] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1584358814] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.26, 1.17, 7.85
[1584363222] SERVICE ALERT: vonly-database;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1584363242] SERVICE ALERT: vonly-database;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1584363284] SERVICE ALERT: vonly-database;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1584363306] SERVICE ALERT: vonly-database;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.05 ms
[1584363654] SERVICE ALERT: vonly-database;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1584363715] SERVICE ALERT: vonly-database;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1584363914] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584364501] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 5.5T of 7.3T
[1584401714] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 31.14, 25.29, 13.96
[1584401774] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 15.55, 21.86, 13.49
[1584401834] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 9.73, 19.02, 13.03
[1584401894] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.59, 16.95, 12.69
[1584403088] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.62, 5.54, 5.20
[1584403148] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 3.60, 5.03, 5.05
[1584403208] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.78, 5.24, 5.13
[1584403268] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.43, 5.08, 5.08
[1584403994] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 29.21, 26.38, 19.96
[1584407294] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.89, 13.15, 18.41
[1584408068] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.62, 3.93, 4.68
[1584409394] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 5.23, 5.95, 7.83
[1584418692] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584418813] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584418934] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584419534] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 10% (5 min average)
[1584421353] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584421474] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584421594] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584422794] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 66% (5 min average)