[1581066302] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581067502] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.94, 5.83, 4.13
[1581067562] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 8.89, 6.19, 4.37
[1581067622] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 5.51, 5.70, 4.32
[1581067982] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.68, 7.36, 5.18
[1581068042] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.61, 6.92, 5.16
[1581068102] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.65, 6.20, 5.02
[1581068162] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.49, 5.58, 4.88
[1581068199] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1581068260] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1581068320] SERVICE ALERT: vdb-101;/mnt/data01 Partition;OK;SOFT;3;/mnt/data01 is at 90% capacity, 3.1T of 3.6T
[1581068680] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1581068740] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1581068800] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;3;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1581068860] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1581072362] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.74, 5.79, 4.10
[1581072422] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 8.44, 5.92, 4.26
[1581072482] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 5.01, 5.39, 4.17
[1581074642] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 6.67, 6.04, 4.35
[1581074702] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 4.66, 5.60, 4.31
[1581075962] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.58, 6.40, 4.88
[1581076022] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.98, 8.08, 5.56
[1581076082] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 7.84, 7.56, 5.54
[1581076142] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.15, 6.83, 5.41
[1581076442] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.98, 4.50, 4.75
[1581077342] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.70, 6.06, 5.21
[1581077402] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.21, 6.22, 5.32
[1581077462] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 7.56, 6.50, 5.47
[1581077522] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 6.34, 6.37, 5.50
[1581078122] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1581078122] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 12.61, 9.07, 7.23
[1581078422] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.00, 6.70, 6.73
[1581084122] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581087422] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 14.09, 8.36, 6.62
[1581087722] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.81, 6.20, 6.17
[1581088922] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 14.78, 8.50, 6.53
[1581089222] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1581089222] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.32, 6.61, 6.27
[1581093722] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.93, 4.05, 4.89
[1581094022] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 6.94, 6.89, 5.91
[1581094082] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 5.36, 6.44, 5.81
[1581094142] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 7.91, 6.85, 5.98
[1581094202] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 7.17, 6.91, 6.07
[1581096122] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.64, 4.64, 4.95
[1581096722] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.12, 5.10, 5.04
[1581096782] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.12, 5.10, 5.04
[1581096842] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.12, 5.10, 5.05
[1581096902] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.02, 5.08, 5.04
[1581096902] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.02, 4.48, 4.92
[1581097202] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.13, 9.15, 6.87
[1581097262] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.86, 8.28, 6.72
[1581097322] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.66, 7.39, 6.51
[1581097382] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.05, 6.71, 6.33
[1581098282] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.83, 4.06, 4.85
[1581098882] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 4.07, 5.58, 5.32
[1581098942] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 3.94, 5.29, 5.23
[1581099002] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.04, 5.05, 5.15
[1581099062] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.36, 4.96, 5.11
[1581099362] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.29, 4.56, 4.89
[1581099662] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.60, 7.67, 5.98
[1581099722] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.59, 7.06, 5.88
[1581099782] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.85, 6.48, 5.75
[1581099842] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.28, 6.02, 5.64
[1581101042] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.30, 3.79, 4.72
[1581101942] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.27, 5.40, 4.47
[1581102002] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 9.15, 6.01, 4.74
[1581102062] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 5.56, 5.58, 4.68
[1581103202] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.76, 4.65, 4.86
[1581105002] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.03, 5.06, 5.01
[1581105062] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.24, 5.11, 5.03
[1581105122] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 4.93, 5.04, 5.00
[1581105182] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.16, 5.10, 5.03
[1581105242] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.41, 5.18, 5.06
[1581105302] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.16, 5.16, 5.06
[1581105362] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.23, 5.20, 5.08
[1581105722] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1581105962] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.87, 4.94, 5.00
[1581106262] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.02, 5.40, 5.16
[1581106322] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.81, 5.47, 5.20
[1581106382] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.95, 5.58, 5.25
[1581106442] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1581106442] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.06, 5.66, 5.30
[1581112142] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1581114242] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.85, 4.91, 5.00
[1581117242] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 7.80, 6.46, 5.44
[1581117302] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 7.75, 6.68, 5.59
[1581117362] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 7.89, 6.91, 5.73
[1581117422] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 7.94, 7.15, 5.89
[1581119522] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.72, 4.76, 4.98
[1581122222] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 7.74, 5.67, 5.08
[1581122282] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 8.33, 6.21, 5.30
[1581122342] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 8.05, 6.51, 5.46
[1581122402] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 7.68, 6.68, 5.59
[1581142795] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 84% capacity, 2.9T of 3.6T
[1581148202] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.85, 4.93, 5.00
[1581148502] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.48, 5.17, 5.07
[1581148562] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.29, 5.16, 5.07
[1581148622] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.05, 5.13, 5.06
[1581148682] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 4.96, 5.10, 5.06