[1584082805] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 89% (5 min average)
[1584082925] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584083045] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 14% (5 min average)
[1584084368] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 20.38, 15.14, 8.10
[1584084428] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 22.89, 16.78, 9.10
[1584084488] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.85, 19.75, 10.60
[1584084548] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.94, 18.28, 10.66
[1584084848] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.44, 25.31, 15.96
[1584085147] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.19, 7.52, 5.55
[1584085207] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.61, 7.51, 5.68
[1584085267] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.83, 7.76, 5.89
[1584085327] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.29, 7.04, 5.76
[1584088594] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 87% (5 min average)
[1584088714] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1584089648] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.53, 14.25, 18.19
[1584089948] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1584089948] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.62, 21.48, 20.29
[1584090248] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.31, 17.78, 18.99
[1584091748] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 22.24, 21.71, 20.55
[1584093000] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.85, 10.88, 10.24
[1584093060] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.16, 10.94, 10.30
[1584093120] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.70, 10.81, 10.29
[1584093180] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.94, 11.15, 10.44
[1584094080] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.24, 7.92, 9.46
[1584095648] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.92, 13.85, 19.33
[1584096034] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584096154] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1584096634] SERVICE ALERT: WINV-103;OS Drive Space;OK;HARD;3;c:\ - total: 68.27 Gb - used: 54.55 Gb (80%) - free 13.71 Gb (20%)
[1584097748] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.57, 28.79, 21.58
[1584098074] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584098194] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 71% (5 min average)
[1584099848] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.96, 9.97, 17.23
[1584100327] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.83, 2.84, 4.36
[1584100749] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.32, 1.48, 7.75
[1584101912] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584102032] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584102152] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584102752] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1584105849] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584116531] SERVICE ALERT: vonly-database;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1584116591] SERVICE ALERT: vonly-database;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1584117311] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 7.85, 6.63, 5.38
[1584117376] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 8.86, 7.21, 5.67
[1584117436] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 7.92, 7.23, 5.77
[1584117497] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 7.51, 7.25, 5.87
[1584118697] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 2.36, 3.53, 4.74
[1584118955] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584119075] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 66% (5 min average)
[1584121632] SERVICE ALERT: vonly-database;PING;CRITICAL;SOFT;1;CRITICAL - Plugin timed out after 10 seconds
[1584121697] SERVICE ALERT: vonly-database;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.08 ms
[1584141634] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;1;c:\ - total: 68.27 Gb - used: 54.66 Gb (80%) - free 13.60 Gb (20%)
[1584141754] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;2;c:\ - total: 68.27 Gb - used: 54.70 Gb (80%) - free 13.57 Gb (20%)
[1584141874] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;c:\ - total: 68.27 Gb - used: 54.62 Gb (80%) - free 13.64 Gb (20%)
[1584142149] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 28.92, 17.98, 9.84
[1584142209] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 25.27, 19.04, 10.73
[1584142269] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 36.41, 23.48, 12.79
[1584142327] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.33, 7.12, 5.06
[1584142329] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 44.64, 28.11, 15.04
[1584142387] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.23, 6.19, 4.87
[1584142629] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.64, 12.87, 12.20
[1584144247] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.23, 6.17, 5.05
[1584144307] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.67, 5.90, 5.03
[1584144367] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.72, 5.70, 5.01
[1584144427] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.47, 5.62, 5.03
[1584144429] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 43.90, 31.80, 20.48
[1584148329] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.38, 11.28, 17.90
[1584148927] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.16, 3.78, 4.67
[1584150129] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.75, 5.81, 7.73
[1584155760] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 737.01 Gb (79%) - free 194.31 Gb (21%)
[1584159247] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584159367] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 28% (5 min average)
[1584163352] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584163472] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584163592] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584164192] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 63% (5 min average)