[1580286565] SERVICE ALERT: vdb-101;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 5% free (47 MB out of 975 MB)
[1580287016] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.48, 6.31, 5.37
[1580287076] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.62, 6.38, 5.45
[1580287136] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.85, 6.20, 5.45
[1580287196] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.86, 6.13, 5.48
[1580296796] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 2.78, 3.54, 4.95
[1580297660] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 287 processes with STATE = RSZDT
[1580297721] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 287 processes with STATE = RSZDT
[1580297781] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 284 processes with STATE = RSZDT
[1580297841] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 234 processes with STATE = RSZDT
[1580298441] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 305 processes with STATE = RSZDT
[1580298501] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 305 processes with STATE = RSZDT
[1580298561] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 305 processes with STATE = RSZDT
[1580298621] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 292 processes with STATE = RSZDT
[1580298921] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 235 processes with STATE = RSZDT
[1580300162] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 6.80, 6.33, 3.77
[1580300222] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 5.95, 6.18, 3.88
[1580300282] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.84, 5.82, 3.90
[1580300942] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 6.52, 6.49, 5.13
[1580301002] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.14, 6.37, 5.17
[1580301062] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.99, 6.27, 5.21
[1580301122] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.75, 6.14, 5.23
[1580302022] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.06, 4.36, 4.92
[1580302622] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 5.77, 5.69, 5.32
[1580302682] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.39, 5.86, 5.40
[1580302742] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.91, 5.82, 5.42
[1580302802] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 6.96, 6.14, 5.56
[1580304902] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1580304902] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 10.02, 7.26, 6.54
[1580305202] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 6.85, 6.89, 6.55
[1580306102] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 2.97, 3.52, 4.88
[1580309121] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 293 processes with STATE = RSZDT
[1580309181] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 289 processes with STATE = RSZDT
[1580309241] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 269 processes with STATE = RSZDT
[1580309301] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 268 processes with STATE = RSZDT
[1580309602] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 235 processes with STATE = RSZDT
[1580311202] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)