[1582186015] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 53.95, 16.88, 7.45
[1582186075] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 28.63, 16.28, 7.83
[1582186135] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 20.40, 16.13, 8.31
[1582186195] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.31, 16.01, 8.76
[1582186495] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.69, 30.84, 17.20
[1582188261] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.58, 5.55, 5.10
[1582188321] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.12, 5.15, 4.99
[1582188381] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.08, 5.45, 5.12
[1582188441] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 3.88, 5.03, 4.99
[1582188801] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.82, 6.08, 5.55
[1582188861] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 11.68, 7.80, 6.18
[1582188921] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 11.13, 8.21, 6.41
[1582188981] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 12.12, 9.11, 6.84
[1582189582] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.66, 2.43, 4.27
[1582189795] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.02, 3.85, 15.50
[1582191650] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 14.31, 13.69, 10.69
[1582191710] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.57, 13.61, 10.86
[1582191770] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.08, 13.19, 10.89
[1582191830] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.58, 12.85, 10.92
[1582192495] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.52, 3.32, 7.10
[1582193030] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 1.35, 4.35, 7.75
[1582198684] SERVICE ALERT: WINV-108;C:\ Drive Space;UNKNOWN;SOFT;1;Server port must be an integer
[1582198705] SERVICE ALERT: WINV-108;NSClient++ Version;UNKNOWN;SOFT;1;Server port must be an integer
[1582198725] SERVICE ALERT: WINV-108;CPU Load;UNKNOWN;SOFT;1;Server port must be an integer
[1582198766] SERVICE ALERT: WINV-108;Explorer;UNKNOWN;SOFT;1;Server port must be an integer
[1582198804] SERVICE ALERT: WINV-108;C:\ Drive Space;UNKNOWN;SOFT;2;NSClient - ERROR: Invalid password.
[1582198807] SERVICE ALERT: WINV-108;Memory Usage;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid password.
[1582198815] SERVICE ALERT: WINV-108;C:\ Drive Space;UNKNOWN;HARD;3;NSClient - ERROR: Invalid password.
[1582198825] SERVICE ALERT: WINV-108;NSClient++ Version;UNKNOWN;SOFT;2;NSClient - ERROR: Invalid password.
[1582198845] SERVICE ALERT: WINV-108;CPU Load;UNKNOWN;SOFT;2;NSClient - ERROR: Invalid password.
[1582198886] SERVICE ALERT: WINV-108;Explorer;UNKNOWN;SOFT;2;NSClient - ERROR: Invalid password.
[1582198889] SERVICE ALERT: WINV-108;Uptime;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid password.
[1582198927] SERVICE ALERT: WINV-108;Memory Usage;UNKNOWN;SOFT;2;NSClient - ERROR: No performance data from command: check_memory
[1582198930] SERVICE ALERT: WINV-108;W3SVC;UNKNOWN;SOFT;1;Unknown command(s): check_service
[1582198945] SERVICE ALERT: WINV-108;NSClient++ Version;OK;HARD;3;NSClient++ 0.5.2.35 2018-01-28
[1582198965] SERVICE ALERT: WINV-108;CPU Load;UNKNOWN;HARD;3;NSClient - ERROR: No performance data from command: check_cpu
[1582199006] SERVICE ALERT: WINV-108;Explorer;UNKNOWN;HARD;3;Unknown command(s): check_process
[1582199009] SERVICE ALERT: WINV-108;Uptime;UNKNOWN;SOFT;2;NSClient - ERROR: No performance data from command: check_uptime
[1582199047] SERVICE ALERT: WINV-108;Memory Usage;UNKNOWN;HARD;3;NSClient - ERROR: No performance data from command: check_memory
[1582199050] SERVICE ALERT: WINV-108;W3SVC;UNKNOWN;SOFT;2;Unknown command(s): check_service
[1582199129] SERVICE ALERT: WINV-108;Uptime;UNKNOWN;HARD;3;NSClient - ERROR: No performance data from command: check_uptime
[1582199170] SERVICE ALERT: WINV-108;W3SVC;UNKNOWN;HARD;3;Unknown command(s): check_service
[1582199545] SERVICE ALERT: WINV-108;NSClient++ Version;UNKNOWN;SOFT;1;NSClient - ERROR: Invalid password.
[1582199665] SERVICE ALERT: WINV-108;NSClient++ Version;UNKNOWN;SOFT;2;NSClient - ERROR: Invalid password.
[1582199729] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.22, 5.10, 2.91
[1582199785] SERVICE ALERT: WINV-108;NSClient++ Version;UNKNOWN;HARD;3;NSClient - ERROR: Invalid password.
[1582199789] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.76, 6.38, 3.50
[1582199849] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 6.23, 5.77, 3.47
[1582200128] SERVICE ALERT: WINV-108;C:\ Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1582200175] SERVICE ALERT: WINV-108;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1582200216] SERVICE ALERT: WINV-108;Explorer;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1582200257] SERVICE ALERT: WINV-108;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1582200339] SERVICE ALERT: WINV-108;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1582200380] SERVICE ALERT: WINV-108;W3SVC;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1582200395] SERVICE ALERT: WINV-108;NSClient++ Version;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1582203229] SERVICE ALERT: vapps-100;Current Users;WARNING;SOFT;1;USERS WARNING - 6 users currently logged in
[1582203289] SERVICE ALERT: vapps-100;Current Users;WARNING;SOFT;2;USERS WARNING - 6 users currently logged in
[1582203349] SERVICE ALERT: vapps-100;Current Users;WARNING;SOFT;3;USERS WARNING - 6 users currently logged in
[1582203409] SERVICE ALERT: vapps-100;Current Users;WARNING;HARD;4;USERS WARNING - 6 users currently logged in
[1582206339] SERVICE ALERT: WINV-108;Uptime;OK;HARD;3;System Uptime - 449 day(s) 21 hour(s) 54 minute(s)
[1582206380] SERVICE ALERT: WINV-108;W3SVC;UNKNOWN;HARD;3;Failed to open service W3SVC: 1060: The specified service does not exist as an installed service.
[1582206395] SERVICE ALERT: WINV-108;NSClient++ Version;OK;HARD;3;NSClient++ 0.4.4.23 2016-04-05
[1582206728] SERVICE ALERT: WINV-108;C:\ Drive Space;WARNING;HARD;3;c:\ - total: 67.89 Gb - used: 54.96 Gb (81%) - free 12.93 Gb (19%)
[1582206751] SERVICE ALERT: WINV-108;CPU Load;OK;HARD;3;CPU Load 12% (5 min average)
[1582206857] SERVICE ALERT: WINV-108;Memory Usage;OK;HARD;3;Memory usage: total:160950.28 MB - used: 41646.32 MB (26%) - free: 119303.95 MB (74%)
[1582207309] SERVICE ALERT: vapps-100;Current Users;OK;HARD;4;USERS OK - 5 users currently logged in
[1582210709] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.85, 3.73, 1.84
[1582210769] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.41, 5.79, 2.67
[1582210829] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 8.51, 5.89, 2.91
[1582210889] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.61, 4.96, 2.78
[1582210995] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 21.46, 14.57, 11.59
[1582211055] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 22.30, 16.05, 12.28
[1582211115] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 21.53, 17.02, 12.86
[1582211175] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 20.67, 17.58, 13.31
[1582212075] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 5.30, 6.56, 9.84
[1582215389] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.04, 3.20, 1.53
[1582215449] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.70, 5.27, 2.35
[1582215509] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 10.50, 6.11, 2.83
[1582215569] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.78, 5.28, 2.75
[1582227788] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.75, 9.34, 6.92
[1582227848] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 21.42, 11.68, 7.87
[1582227908] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 12.74, 11.07, 7.91
[1582227968] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.60, 10.15, 7.79
[1582228268] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 5.19, 7.26, 7.22
[1582228868] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.83, 11.55, 8.75
[1582228928] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 9.86, 10.39, 8.52
[1582228988] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 8.02, 9.75, 8.42
[1582228991] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.24, 5.60, 5.04
[1582229048] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.85, 9.64, 8.46
[1582229051] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 3.90, 5.32, 4.98
[1582229111] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.67, 5.72, 5.15
[1582229171] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 3.56, 5.16, 4.99
[1582229648] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1582229648] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 3.55, 5.26, 6.88
[1582229948] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 7.77, 10.29, 8.93
[1582230008] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 4.88, 9.01, 8.57
[1582230068] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 4.94, 8.32, 8.36
[1582230128] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.98, 8.05, 8.27
[1582231028] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 3.40, 6.17, 7.36
[1582234628] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.51, 11.43, 8.21
[1582234688] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 10.32, 10.39, 8.06
[1582234748] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;3;OK - load average: 7.13, 9.49, 7.89
[1582234808] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 8.18, 9.54, 8.02
[1582234868] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 6.17, 8.79, 7.86
[1582235828] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582237628] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 9.69, 11.35, 8.84
[1582237688] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 6.83, 10.28, 8.63
[1582237748] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 6.35, 9.59, 8.50
[1582237808] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 4.91, 8.63, 8.24
[1582238108] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 7.60, 7.52, 7.83
[1582239608] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 7.80, 10.94, 9.06
[1582239668] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 5.88, 9.85, 8.80
[1582239728] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 4.99, 8.88, 8.53
[1582239788] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 4.90, 8.16, 8.31
[1582240688] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582240688] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 3.74, 5.83, 7.30
[1582240931] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.54, 5.62, 5.07
[1582240991] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.80, 5.42, 5.03
[1582241051] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 3.92, 4.94, 4.89
[1582245788] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582246334] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;1;PROCS WARNING: 2013 processes with STATE = RSZDT
[1582246394] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;2;PROCS WARNING: 2013 processes with STATE = RSZDT
[1582246455] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;3;PROCS WARNING: 2013 processes with STATE = RSZDT
[1582246515] SERVICE ALERT: vapps-100;Total Processes;WARNING;HARD;4;PROCS WARNING: 2012 processes with STATE = RSZDT
[1582246575] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 15.64, 13.70, 10.39
[1582246635] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 15.21, 13.91, 10.67
[1582246695] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 14.62, 13.94, 10.88
[1582246755] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 14.23, 13.97, 11.09
[1582248311] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.02, 5.33, 5.05
[1582248371] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.08, 5.19, 5.02
[1582248431] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.91, 5.23, 5.05
[1582248491] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.13, 4.93, 4.96
[1582248791] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.48, 5.00, 5.01
[1582248851] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.31, 4.84, 4.95
[1582248911] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.49, 5.10, 5.03
[1582248971] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 11.01, 6.59, 5.54
[1582249031] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.34, 6.17, 5.46
[1582249091] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.46, 5.99, 5.45
[1582251491] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.34, 4.42, 4.89
[1582251788] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.64, 12.65, 8.55
[1582251848] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 13.76, 12.19, 8.67
[1582251909] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 9.14, 11.12, 8.52
[1582251969] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.47, 10.46, 8.45
[1582252269] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 5.74, 7.29, 7.61
[1582252691] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.79, 5.27, 5.04
[1582252751] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.19, 5.82, 5.24
[1582252811] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.14, 6.44, 5.48
[1582252869] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 12.31, 12.41, 9.68
[1582252871] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.72, 6.04, 5.41
[1582252929] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 8.77, 11.39, 9.50
[1582252989] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 6.80, 10.39, 9.27
[1582253049] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.87, 9.46, 9.02
[1582254371] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (20.9% change >= 20.0% threshold)
[1582254371] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.88, 4.47, 4.87
[1582254549] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1582254549] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 6.00, 6.22, 7.57
[1582254555] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 5.15, 7.52, 9.55
[1582254671] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.36, 5.23, 5.03
[1582254731] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.91, 5.00, 4.96
[1582255749] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.77, 12.75, 9.25
[1582255809] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 10.73, 11.84, 9.16
[1582255869] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 8.68, 11.00, 9.03
[1582255929] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.39, 10.45, 8.96
[1582256891] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.44, 5.55, 5.20
[1582256951] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 3.98, 5.11, 5.07
[1582257011] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 3.08, 4.66, 4.91
[1582257129] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.63, 5.57, 7.10
[1582257729] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 11.53, 11.26, 8.96
[1582257789] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 10.37, 11.00, 9.02
[1582257849] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 8.40, 10.22, 8.87
[1582257909] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.37, 9.94, 8.86
[1582258571] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1582259109] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.73, 5.19, 7.28
[1582259709] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 9.22, 10.47, 8.80
[1582259769] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 9.39, 10.28, 8.84
[1582259829] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 7.36, 9.54, 8.67
[1582259889] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.40, 9.54, 8.73
[1582263189] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.11, 4.84, 7.08
[1582263489] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 12.39, 11.34, 9.27
[1582263549] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 9.15, 10.63, 9.15
[1582263609] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 8.25, 10.02, 9.03
[1582263669] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.48, 9.42, 8.88
[1582264569] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.57, 6.75, 7.68
[1582265469] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.12, 14.47, 10.20
[1582265529] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 17.18, 14.96, 10.63
[1582265589] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 19.14, 15.95, 11.25
[1582265649] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 21.78, 17.14, 11.94
[1582266549] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.89, 29.09, 21.50
[1582267571] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.20, 5.42, 5.01
[1582267631] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.55, 5.79, 5.16
[1582267691] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.63, 5.34, 5.04
[1582267751] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.21, 4.78, 4.86
[1582268649] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.32, 6.48, 16.04
[1582269549] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.06, 0.45, 6.19
[1582269899] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (25.1% change >= 20.0% threshold)
[1582270243] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1582270255] SERVICE ALERT: vonly-backup;Total Processes;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270275] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270312] SERVICE ALERT: vonly-backup;PING;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270315] SERVICE ALERT: vonly-backup;Total Processes;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270335] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270344] SERVICE ALERT: vonly-backup;Root Partition;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270372] SERVICE ALERT: vonly-backup;PING;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270375] SERVICE ALERT: vonly-backup;Total Processes;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270395] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270404] SERVICE ALERT: vonly-backup;Root Partition;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270411] SERVICE ALERT: vonly-backup;Swap Usage;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270432] SERVICE ALERT: vonly-backup;PING;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270435] SERVICE ALERT: vonly-backup;Total Processes;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270455] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270464] SERVICE ALERT: vonly-backup;Root Partition;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270469] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270471] SERVICE ALERT: vonly-backup;Swap Usage;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270473] SERVICE ALERT: vonly-backup;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270492] SERVICE ALERT: vonly-backup;PING;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270524] SERVICE ALERT: vonly-backup;Root Partition;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270529] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270531] SERVICE ALERT: vonly-backup;Swap Usage;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270533] SERVICE ALERT: vonly-backup;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270589] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270591] SERVICE ALERT: vonly-backup;Swap Usage;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270593] SERVICE ALERT: vonly-backup;Current Users;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270649] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582270653] SERVICE ALERT: vonly-backup;Current Users;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582271392] SERVICE ALERT: vonly-backup;PING;UNKNOWN;HARD;4;NRPE: Command 'check_ping' not defined
[1582271424] SERVICE ALERT: vonly-backup;Root Partition;UNKNOWN;HARD;4;NRPE: Command 'check_disk' not defined
[1582271426] SERVICE ALERT: vonly-backup;Current Users;OK;HARD;4;USERS OK - 1 users currently logged in
[1582271491] SERVICE ALERT: vonly-backup;Swap Usage;UNKNOWN;HARD;4;NRPE: Command 'check_swap' not defined
[1582271515] SERVICE ALERT: vonly-backup;/media/data Partition;UNKNOWN;HARD;4;NRPE: Command 'check_volume' not defined
[1582271547] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average per CPU: 0.02, 0.04, 0.04