[1585638274] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1585638394] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 40% (5 min average)
[1585639599] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 23.17, 14.16, 6.79
[1585639659] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 21.38, 15.51, 7.73
[1585639719] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 17.80, 15.77, 8.32
[1585639779] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 22.95, 17.67, 9.44
[1585640444] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.51, 7.72, 5.61
[1585640504] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.47, 8.05, 5.87
[1585640564] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.80, 7.61, 5.85
[1585640624] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.60, 7.44, 5.91
[1585640980] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 22.71, 26.39, 21.41
[1585642389] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.70, 26.88, 25.13
[1585644877] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 9.58, 8.00, 5.06
[1585644937] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 8.00, 7.86, 5.20
[1585644997] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 7.82, 7.81, 5.35
[1585645057] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 7.65, 7.81, 5.51
[1585648389] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 22.94, 23.09, 24.67
[1585649080] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.83, 8.69, 19.46
[1585649680] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 14.12, 20.41, 20.55
[1585650365] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1585650603] SERVICE ALERT: vdb-103;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 36.59, 21.65, 9.25
[1585650663] SERVICE ALERT: vdb-103;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.82, 24.16, 10.89
[1585650723] SERVICE ALERT: vdb-103;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 34.60, 25.98, 12.35
[1585650783] SERVICE ALERT: vdb-103;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.42, 28.15, 13.96
[1585651780] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.92, 10.16, 17.44
[1585652080] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1585652080] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 16.92, 22.48, 21.12
[1585652924] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.35, 4.41, 4.80
[1585653524] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.84, 5.33, 5.08
[1585653584] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.40, 5.53, 5.17
[1585653644] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 3.19, 4.89, 4.97
[1585653704] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.59, 5.12, 5.03
[1585653764] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 3.14, 4.50, 4.83
[1585653824] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.53, 5.11, 5.01
[1585653884] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 5.10, 4.93, 4.95
[1585654083] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 0.00, 3.52, 17.09
[1585655080] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.00, 12.49, 17.99
[1585655144] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.91, 6.09, 5.42
[1585655204] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.60, 6.41, 5.57
[1585655264] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.67, 6.09, 5.51
[1585655283] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 0.00, 0.05, 4.66
[1585655324] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.40, 5.66, 5.40
[1585655380] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 26.53, 27.37, 23.00
[1585656824] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.61, 4.56, 4.85
[1585656880] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.44, 11.12, 16.55
[1585657180] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 19.76, 23.68, 20.88
[1585657480] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.15, 17.95, 19.76
[1585657657] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 1.93, 2.73, 4.70
[1585657780] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.89, 22.25, 20.26
[1585658380] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.15, 14.51, 19.06
[1585658680] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 17.17, 22.60, 21.60
[1585659580] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.63, 12.01, 17.11
[1585660480] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.33, 0.79, 6.66
[1585665580] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1585669647] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1585669767] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585669887] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585670487] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 51% (5 min average)
[1585676165] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 90% capacity! 6.2T of 7.3T
[1585682624] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 24.40, 12.15, 5.73
[1585682684] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 25.46, 14.69, 7.01
[1585682744] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 14.07, 13.71, 7.17
[1585682804] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.97, 11.45, 6.80
[1585683087] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585683207] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1585683704] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.58, 3.52, 4.76
[1585687005] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.92, 8.93, 5.16
[1585687067] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.00, 7.92, 5.05
[1585687127] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 12.00, 9.70, 5.87
[1585687187] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.32, 8.19, 5.58
[1585688087] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1585688087] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.45, 2.55, 4.08
[1585693187] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1585694914] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1585695034] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 20% (5 min average)
[1585698280] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 23.32, 13.08, 6.04
[1585698289] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 24.04, 12.44, 6.03
[1585698341] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 30.61, 16.94, 7.81
[1585698349] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 13.51, 11.53, 6.12
[1585698401] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 32.68, 20.19, 9.51
[1585698409] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 15.19, 12.54, 6.81
[1585698430] SERVICE ALERT: vinfra-200;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585698461] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 21.47, 19.90, 10.11
[1585698469] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 12.20, 12.16, 7.05
[1585698490] SERVICE ALERT: vinfra-200;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.6 (protocol 2.0)
[1585698495] SERVICE ALERT: vinfra-200;MySQL Uptime;CRITICAL;SOFT;1;CRITICAL - database is up since 3 minutes
[1585698547] SERVICE ALERT: vinfra-200;Swap Usage;OK;HARD;4;SWAP OK - 83% free (9307 MB out of 11215 MB)
[1585698555] SERVICE ALERT: vinfra-200;MySQL Uptime;CRITICAL;SOFT;2;CRITICAL - database is up since 4 minutes
[1585698615] SERVICE ALERT: vinfra-200;MySQL Uptime;WARNING;SOFT;3;WARNING - database is up since 5 minutes
[1585698675] SERVICE ALERT: vinfra-200;MySQL Uptime;WARNING;HARD;4;WARNING - database is up since 6 minutes
[1585698761] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 45.44, 29.62, 16.48
[1585698975] SERVICE ALERT: vinfra-200;MySQL Uptime;OK;HARD;4;OK - database is up since 11 minutes
[1585703861] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.58, 12.82, 18.42
[1585705669] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.83, 3.25, 4.48
[1585705961] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.10, 1.60, 7.92
[1585706093] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585706213] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1585713489] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.15, 9.45, 9.82