[1580809869] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.88, 5.67, 5.15
[1580809929] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.83, 5.71, 5.20
[1580809989] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.07, 5.78, 5.25
[1580810049] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.00, 5.80, 5.29
[1580813949] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.47, 4.30, 4.83
[1580814249] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.84, 5.13, 5.04
[1580814309] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.90, 5.26, 5.09
[1580814369] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.05, 5.12, 5.05
[1580814429] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.72, 5.29, 5.11
[1580815629] SERVICE FLAPPING ALERT: vonly-database;Current Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1580815629] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.20, 4.69, 4.93
[1580817729] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.59, 5.50, 5.17
[1580817789] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.54, 5.48, 5.18
[1580817849] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.51, 5.52, 5.21
[1580817909] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.45, 5.49, 5.22
[1580821809] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.54, 4.32, 4.82
[1580824809] SERVICE FLAPPING ALERT: vonly-database;Current Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1580832909] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 4.84, 5.11, 5.04
[1580832969] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 4.74, 5.08, 5.04
[1580833029] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.07, 5.07, 5.04
[1580833089] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 4.72, 5.02, 5.02
[1580834589] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.80, 4.85, 4.95
[1580835189] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.57, 5.24, 5.07
[1580835249] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.71, 5.37, 5.13
[1580835309] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.46, 5.35, 5.14
[1580835369] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.03, 5.55, 5.22
[1580838669] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.09, 4.60, 4.88
[1580839869] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.39, 5.14, 5.03
[1580839929] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 4.45, 4.91, 4.96
[1580839989] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.49, 5.08, 5.01
[1580840049] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 5.06, 5.01, 4.99
[1580840109] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.18, 5.09, 5.02
[1580840170] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.11, 5.05, 5.01
[1580840230] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;3;OK - load average: 4.76, 4.99, 4.99
[1580842090] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.62, 5.19, 5.01
[1580842150] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.98, 5.35, 5.08
[1580842210] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.00, 5.45, 5.13
[1580842270] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.02, 5.28, 5.09
[1580842570] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.26, 4.80, 4.96
[1580842870] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.14, 5.16, 5.08
[1580842930] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.33, 5.17, 5.08
[1580842990] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.18, 5.17, 5.09
[1580843050] SERVICE FLAPPING ALERT: vonly-database;Current Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1580843050] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.50, 5.20, 5.10
[1580846350] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.77, 4.85, 4.99
[1580846650] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.60, 5.18, 5.09
[1580846710] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 4.71, 5.00, 5.03
[1580846770] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 4.55, 4.94, 5.02
[1580846830] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.49, 4.81, 4.96
[1580847130] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.58, 5.18, 5.08
[1580847190] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.59, 5.23, 5.10
[1580847250] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.68, 5.33, 5.14
[1580847310] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 4.85, 5.17, 5.10
[1580850910] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 5.10, 4.96, 5.00
[1580851510] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.41, 5.10, 5.03
[1580851570] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 4.49, 4.91, 4.97
[1580852230] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.16, 5.14, 5.03
[1580852290] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.07, 5.13, 5.03
[1580852350] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.19, 5.20, 5.06
[1580852410] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.16, 5.19, 5.07
[1580853610] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.29, 4.74, 4.94
[1580854510] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.03, 5.07, 5.01
[1580854570] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 4.90, 5.04, 5.00
[1580854630] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 4.99, 5.07, 5.02
[1580854690] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 4.81, 5.02, 5.01
[1580854750] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.22, 5.11, 5.04
[1580854810] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.01, 5.08, 5.04
[1580856721] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.62, 4.45, 4.91
[1580859610] SERVICE FLAPPING ALERT: vonly-database;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1580864221] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.80, 6.57, 5.42
[1580864281] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.70, 6.57, 5.49
[1580864341] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 7.31, 6.79, 5.64
[1580864402] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 7.15, 6.83, 5.72
[1580866502] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.73, 4.56, 4.88
[1580871602] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 7.92, 7.03, 5.47
[1580871662] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.50, 7.59, 5.76
[1580871722] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 8.98, 7.80, 5.95
[1580871782] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 9.47, 8.04, 6.15
[1580871910] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.67, 4.83, 4.99
[1580872510] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.41, 5.09, 5.02
[1580872570] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 4.92, 5.05, 5.02
[1580872630] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.70, 5.23, 5.08
[1580872690] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.24, 5.20, 5.08
[1580873590] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 5.51, 4.93, 4.99
[1580874790] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.53, 5.28, 5.08
[1580874850] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.98, 5.47, 5.16
[1580874910] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.23, 5.36, 5.14
[1580874970] SERVICE FLAPPING ALERT: vonly-database;Current Load;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1580874970] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 4.54, 5.18, 5.09
[1580877370] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.11, 4.77, 5.00
[1580877670] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.35, 5.38, 5.22
[1580877730] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.06, 5.24, 5.18
[1580877790] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.20, 5.26, 5.19
[1580877850] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.56, 5.36, 5.23
[1580879282] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.48, 4.01, 4.73
[1580883550] SERVICE FLAPPING ALERT: vonly-database;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1580888650] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.73, 4.63, 4.92