[1580547360] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.69, 6.54, 5.46
[1580547420] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.11, 6.41, 5.48
[1580547480] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.33, 6.40, 5.54
[1580547540] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.67, 6.48, 5.62
[1580559540] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.24, 4.43, 4.85
[1580563140] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 4.94, 5.09, 5.01
[1580563200] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.00, 5.10, 5.02
[1580563260] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 4.58, 4.97, 4.98
[1580563620] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.32, 5.13
[1580563680] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.50, 5.37, 5.16
[1580563740] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.60, 5.43, 5.20
[1580563800] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.38, 5.40, 5.20
[1580568600] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.63, 4.64, 4.90
[1580572200] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.51, 5.26, 5.08
[1580572260] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.31, 5.25, 5.09
[1580572320] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.03, 5.18, 5.08
[1580572380] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 4.72, 5.08, 5.05
[1580572680] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.29, 4.62, 4.86
[1580575080] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.98, 5.47, 5.13
[1580575140] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 4.83, 5.26, 5.08
[1580575200] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 4.17, 5.00, 5.00
[1580580960] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.10, 5.67, 5.16
[1580581020] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.26, 5.82, 5.25
[1580581080] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.04, 5.81, 5.28
[1580581140] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.82, 5.77, 5.30
[1580583808] SERVICE ALERT: vdb-102;Current Load;WARNING;SOFT;1;WARNING - load average: 5.29, 5.20, 5.03
[1580583868] SERVICE ALERT: vdb-102;Current Load;WARNING;SOFT;2;WARNING - load average: 5.46, 5.25, 5.06
[1580583928] SERVICE ALERT: vdb-102;Current Load;WARNING;SOFT;3;WARNING - load average: 5.45, 5.29, 5.09
[1580583988] SERVICE ALERT: vdb-102;Current Load;WARNING;HARD;4;WARNING - load average: 5.58, 5.35, 5.12
[1580585188] SERVICE ALERT: vdb-102;Current Load;OK;HARD;4;OK - load average: 4.65, 4.72, 4.96
[1580594040] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.97, 4.37, 4.88
[1580605440] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.91, 6.86, 5.49
[1580605500] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.82, 6.84, 5.57
[1580605560] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.25, 6.67, 5.59
[1580605620] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.22, 6.35, 5.55
[1580613120] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.64, 3.85, 4.80
[1580613720] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 18.30, 13.02, 8.31
[1580613780] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 18.46, 14.03, 8.96
[1580613840] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 18.57, 14.88, 9.57
[1580613900] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.82, 13.56, 9.46
[1580615100] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 2.65, 2.88, 4.54