[1581945422] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.88, 4.72, 3.07
[1581945482] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 6.85, 4.78, 3.21
[1581947523] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 9.93, 8.49, 5.92
[1581947583] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.80, 8.99, 6.25
[1581947643] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.91, 9.38, 6.56
[1581947703] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.19, 9.47, 6.77
[1581962103] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.35, 4.61, 4.91
[1581988203] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.87, 5.97, 5.07
[1581988263] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 7.17, 6.20, 5.21
[1581988323] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 7.67, 6.52, 5.38
[1581988383] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 7.06, 6.57, 5.48
[1581991083] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.61, 4.79, 4.97
[1581991383] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.45, 5.13, 5.06
[1581991443] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.24, 5.14, 5.07
[1581991503] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 4.71, 5.01, 5.03
[1581991563] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.63, 4.92, 5.00