[1580890982] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 7.73, 6.91, 5.14
[1580891042] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 8.15, 7.16, 5.33
[1580891102] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 8.66, 7.47, 5.56
[1580891162] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 8.18, 7.54, 5.70
[1580898380] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 91% capacity! 6.2T of 7.3T
[1580898440] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 91% capacity! 6.2T of 7.3T
[1580898500] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 91% capacity! 6.2T of 7.3T
[1580898560] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1580902718] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1580902778] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1580902802] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.31, 4.17, 2.40
[1580902838] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;SOFT;3;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1580902862] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 5.86, 4.08, 2.48
[1580902898] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1580906360] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 90% capacity, 6.2T of 7.3T
[1580919422] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.98, 4.46, 2.44
[1580919482] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 5.35, 4.13, 2.45
[1580920262] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.62, 4.59, 4.89
[1580944562] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 7.78, 6.67, 5.38
[1580944622] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.66, 6.60, 5.44
[1580944682] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.36, 6.50, 5.48
[1580944742] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.98, 6.37, 5.50
[1580945642] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.95, 4.50, 4.96
[1580949542] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 9.64, 6.86, 5.39
[1580949602] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.86, 7.39, 5.67
[1580949662] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 9.55, 7.80, 5.92
[1580949722] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 8.56, 7.86, 6.06
[1580960335] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 91% capacity! 3.1T of 3.6T
[1580960395] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 91% capacity! 3.1T of 3.6T
[1580960455] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 91% capacity! 3.1T of 3.6T
[1580960515] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 91% capacity! 3.1T of 3.6T
[1580970242] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.73, 3.98, 2.45
[1580970302] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 6.72, 4.39, 2.70
[1580970715] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 82% capacity, 2.8T of 3.6T