[1582358542] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 705.29 Gb (76%) - free 226.03 Gb (24%)
[1582360127] SERVICE ALERT: vapps-100;Total Processes;OK;HARD;4;PROCS OK: 201 processes with STATE = RSZDT
[1582360943] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 49.88, 29.45, 14.95
[1582361003] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 42.39, 31.01, 16.40
[1582361063] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 37.75, 31.90, 17.63
[1582361123] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.06, 31.59, 18.41
[1582361377] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 71% capacity, 2.4T of 3.6T
[1582361831] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.04, 7.23, 5.77
[1582361891] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.85, 6.59, 5.64
[1582361951] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.52, 6.13, 5.54
[1582362011] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.22, 5.73, 5.44
[1582362323] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.43, 12.73, 18.25
[1582362540] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582362624] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 35.18, 20.29, 19.44
[1582362660] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 20% (5 min average)
[1582363211] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.38, 4.41, 4.82
[1582363811] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.04, 5.17, 5.03
[1582363871] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.46, 5.73, 5.24
[1582363931] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.10, 5.36, 5.14
[1582363991] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.15, 5.80, 5.30
[1582365315] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.53, 11.37, 10.14
[1582365376] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.19, 11.07, 10.12
[1582365408] SERVICE ALERT: vonly-database;Current Load;CRITICAL;SOFT;1;(Return code of 255 for service 'Current Load' on host 'vonly-database' was out of bounds)
[1582365419] SERVICE ALERT: vonly-database;PING;CRITICAL;SOFT;1;(Return code of 255 for service 'PING' on host 'vonly-database' was out of bounds)
[1582365427] HOST ALERT: vonly-database;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1582365436] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.31, 11.25, 10.24
[1582365448] HOST ALERT: vonly-database;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.28 ms
[1582365468] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 0.24, 0.10, 0.03
[1582365483] SERVICE ALERT: vonly-database;PING;OK;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.03 ms
[1582365496] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 9.43, 10.72, 10.12
[1582365496] SERVICE ALERT: WINV-112;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1582365616] SERVICE ALERT: WINV-112;CPU Load;OK;SOFT;2;CPU Load 57% (5 min average)
[1582365796] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.86, 9.67, 9.78
[1582367724] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.56, 10.93, 19.47
[1582368024] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1582368024] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 35.28, 21.72, 21.24
[1582368324] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.27, 15.45, 19.09
[1582368791] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.84, 4.05, 4.88
[1582368948] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582369068] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1582369188] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1582369391] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 3.38, 4.78, 5.05
[1582369451] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.92, 5.26, 5.21
[1582369511] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 3.84, 4.89, 5.09
[1582369571] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.75, 4.66, 5.00
[1582369703] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1582369788] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1582369823] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1582370751] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582370811] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1582370871] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582370931] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1582370991] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 249 processes with STATE = RSZDT
[1582371052] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582371112] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 248 processes with STATE = RSZDT
[1582372072] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1582372132] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes with STATE = RSZDT
[1582372192] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1582372252] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 254 processes with STATE = RSZDT
[1582372824] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.07, 5.33, 7.72
[1582372852] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582373124] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 24.28, 13.40, 10.12
[1582373152] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1582373184] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 19.00, 14.21, 10.62
[1582373212] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1582373244] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 12.72, 13.31, 10.54
[1582373272] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes with STATE = RSZDT
[1582373304] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1582373304] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.24, 12.18, 10.32
[1582373332] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 252 processes with STATE = RSZDT
[1582373392] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1582373452] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1582373640] SERVICE ALERT: vinfra-200;VAULT;UNKNOWN;SOFT;1;UNKNOWN - Vault is not initialized.
[1582373650] SERVICE ALERT: vinfra-200;VAULT;UNKNOWN;SOFT;2;UNKNOWN - Vault is not initialized.
[1582373710] SERVICE ALERT: vinfra-200;VAULT;UNKNOWN;SOFT;3;UNKNOWN - Vault is not initialized.
[1582373752] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1582373752] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582373770] SERVICE ALERT: vinfra-200;VAULT;UNKNOWN;HARD;4;UNKNOWN - Vault is not initialized.
[1582374025] SERVICE ALERT: vinfra-200;VAULT;OK;HARD;4;OK - Vault is initialized, unsealed, and active
[1582374052] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582374112] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1582374172] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1582374232] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 252 processes with STATE = RSZDT
[1582375271] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.92, 6.31, 5.10
[1582375331] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.11, 6.00, 5.07
[1582375391] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 4.28, 5.50, 4.95
[1582376639] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1044 processes with STATE = RSZDT
[1582376699] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1045 processes with STATE = RSZDT
[1582376760] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;3;PROCS WARNING: 1045 processes with STATE = RSZDT
[1582376820] SERVICE ALERT: vapps-100;Total Processes;WARNING;HARD;4;PROCS WARNING: 1045 processes with STATE = RSZDT
[1582377204] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.23, 15.33, 11.15
[1582377504] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.53, 13.74, 11.66
[1582378644] SERVICE ALERT: vdb-staging;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 81% capacity! 6.9T of 9.1T
[1582378704] SERVICE ALERT: vdb-staging;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 81% capacity! 6.9T of 9.1T
[1582378764] SERVICE ALERT: vdb-staging;/mnt/data01 Partition;WARNING;SOFT;3;/mnt/data01 is at 81% capacity! 6.9T of 9.1T
[1582378824] SERVICE ALERT: vdb-staging;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 81% capacity! 6.9T of 9.1T
[1582379004] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.95, 5.17, 7.90
[1582379304] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 28.95, 18.86, 12.94
[1582379364] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 31.04, 21.38, 14.18
[1582379424] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 25.07, 21.72, 14.77
[1582379484] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 25.88, 22.40, 15.42
[1582379784] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.92, 25.35, 18.68
[1582379940] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582380239] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582380551] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.46, 5.68, 5.07
[1582380611] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 5.34, 5.36, 5.00
[1582380671] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.47, 5.56, 5.08
[1582380731] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.22, 5.11, 4.96
[1582380790] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.13, 5.18, 2.75
[1582380850] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 3.86, 4.51, 2.67
[1582381584] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.17, 4.81, 16.22
[1582382484] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.12, 0.48, 6.33
[1582387584] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582397340] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1582397400] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 245 processes with STATE = RSZDT
[1582407553] SERVICE ALERT: WINV-104;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1582407673] SERVICE ALERT: WINV-104;Memory Usage;OK;SOFT;2;Memory usage: total:129142.02 MB - used: 44787.89 MB (35%) - free: 84354.13 MB (65%)
[1582413684] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 20.56, 17.79, 8.88
[1582413744] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 23.30, 19.18, 9.92
[1582413804] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 37.23, 22.79, 11.69
[1582413864] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.71, 25.76, 13.40
[1582414164] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.92, 18.87, 14.42
[1582414765] SERVICE ALERT: WINV-101;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1582414797] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;No data was received from host!
[1582414885] SERVICE ALERT: WINV-101;Memory Usage;OK;SOFT;2;Memory usage: total:98284.48 MB - used: 41996.68 MB (43%) - free: 56287.80 MB (57%)
[1582414917] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1582415064] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.51, 2.57, 7.38
[1582418965] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 57.38, 30.87, 12.71
[1582419025] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 54.27, 34.93, 15.25
[1582419086] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 62.71, 40.72, 18.46
[1582419146] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 59.53, 43.96, 20.98
[1582420091] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.35, 6.69, 5.04
[1582420151] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 5.08, 5.95, 4.88
[1582420811] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.76, 5.68, 5.10
[1582420871] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.74, 5.56, 5.09
[1582420931] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.50, 6.17, 5.32
[1582420991] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.03, 5.85, 5.26
[1582422791] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.35, 4.39, 4.92
[1582424546] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.29, 2.65, 16.50
[1582425446] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.48, 0.38, 6.41
[1582432437] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582432557] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1582432677] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 9% (5 min average)
[1582435210] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.74, 3.79, 1.78
[1582435270] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 4.68, 3.52, 1.82
[1582435272] SERVICE ALERT: WINV-102;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1582435397] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;No data was received from host!
[1582435402] SERVICE ALERT: WINV-102;Memory Usage;OK;SOFT;2;Memory usage: total:208877.43 MB - used: 67299.84 MB (32%) - free: 141577.59 MB (68%)
[1582435405] SERVICE ALERT: WINV-102;Uptime;WARNING;SOFT;1;No data was received from host!
[1582435517] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1582435525] SERVICE ALERT: WINV-102;Uptime;OK;SOFT;2;System Uptime - 152 day(s) 10 hour(s) 39 minute(s)
[1582441046] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;No data was received from host!
[1582441166] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)