[1585378802] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 0.03, 1.98, 4.68
[1585379167] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.08, 8.50, 9.60
[1585379734] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.54, 5.42, 5.07
[1585379794] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.60, 5.47, 5.10
[1585379855] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 4.65, 5.24, 5.05
[1585379915] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.37, 5.03, 4.99
[1585379971] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 23.17, 7.90, 3.71
[1585380031] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 40.06, 15.29, 6.50
[1585380091] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 27.62, 16.64, 7.54
[1585380151] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 24.38, 17.78, 8.51
[1585380367] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.67, 11.87, 10.62
[1585380427] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.27, 11.96, 10.73
[1585380487] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.52, 12.07, 10.84
[1585380547] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.10, 12.04, 10.91
[1585380882] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.16, 7.15, 5.38
[1585380942] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.34, 7.43, 5.59
[1585381002] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.95, 7.11, 5.60
[1585381051] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.56, 25.62, 19.27
[1585381062] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.85, 7.53, 5.85
[1585381443] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 5.95, 6.20, 5.19
[1585381503] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 5.86, 6.13, 5.23
[1585381563] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 6.90, 6.41, 5.38
[1585381623] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 5.16, 6.07, 5.33
[1585384007] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1585384127] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 65% (5 min average)
[1585388223] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 4.66, 4.72, 4.99
[1585389001] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;1;/ is at 81% capacity! 348G of 457G
[1585389061] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;2;/ is at 81% capacity! 349G of 457G
[1585389121] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;3;/ is at 81% capacity! 349G of 457G
[1585389181] SERVICE ALERT: vdb-103;Root Partition;WARNING;HARD;4;/ is at 81% capacity! 350G of 457G
[1585393081] SERVICE ALERT: vdb-103;Root Partition;CRITICAL;HARD;4;/ is at 91% capacity! 391G of 457G
[1585395364] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1585396034] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585396154] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1585396351] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.19, 3.67, 15.20
[1585396362] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.59, 2.83, 4.62
[1585397252] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.74, 1.07, 6.10
[1585397464] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 90% capacity! 6.2T of 7.3T
[1585400647] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.48, 9.58, 9.97
[1585405874] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585405994] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585406114] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585406714] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1585412281] SERVICE ALERT: vdb-103;Root Partition;OK;HARD;4;/ is at 22% capacity, 94G of 457G
[1585413247] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 9.99, 10.18, 10.06
[1585413307] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.06, 10.18, 10.07
[1585413367] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 9.95, 10.14, 10.07
[1585413427] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 9.89, 10.08, 10.05
[1585414327] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.39, 9.42, 9.75
[1585417923] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 10.83, 9.41, 5.68
[1585417983] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 10.78, 9.63, 5.99
[1585418043] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 10.02, 9.63, 6.22
[1585418103] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 10.06, 9.71, 6.46
[1585418521] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 6.33, 6.31, 5.06
[1585418581] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 6.24, 6.29, 5.13
[1585418641] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 6.26, 6.27, 5.20
[1585418701] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 6.42, 6.31, 5.29
[1585422123] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 81% capacity! 5.6T of 7.3T
[1585422183] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 82% capacity! 5.6T of 7.3T
[1585422243] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;3;/mnt/data01 is at 82% capacity! 5.6T of 7.3T
[1585422301] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 4.81, 4.83, 5.00
[1585422304] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 82% capacity! 5.6T of 7.3T
[1585423253] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;No data was received from host!
[1585423373] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 28% (5 min average)
[1585424220] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (1950 MB out of 10239 MB)
[1585424280] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 19% free (1879 MB out of 10239 MB)
[1585424340] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 18% free (1809 MB out of 10239 MB)
[1585424400] SERVICE ALERT: vdb-103;Swap Usage;WARNING;HARD;4;SWAP WARNING - 18% free (1741 MB out of 10239 MB)
[1585425300] SERVICE ALERT: vdb-103;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (778 MB out of 10239 MB)
[1585425304] SERVICE ALERT: vdb-103;/mnt/data01 Partition;CRITICAL;HARD;4;/mnt/data01 is at 91% capacity! 6.3T of 7.3T
[1585425893] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 83% (5 min average)
[1585426013] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;2;CPU Load 80% (5 min average)
[1585426133] SERVICE ALERT: WINV-102;CPU Load;WARNING;HARD;3;CPU Load 80% (5 min average)
[1585426733] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 16% (5 min average)
[1585426848] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585426968] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 69% (5 min average)
[1585429503] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 2.56, 3.53, 4.80
[1585433681] SERVICE ALERT: WINV-101;Uptime;WARNING;SOFT;1;No data was received from host!
[1585433720] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585433810] SERVICE ALERT: WINV-101;Uptime;WARNING;SOFT;2;No data was received from host!
[1585433844] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585433930] SERVICE ALERT: WINV-101;Uptime;OK;HARD;3;System Uptime - 23 day(s) 15 hour(s) 57 minute(s)
[1585433964] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 17% (5 min average)
[1585438352] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 11.32, 10.01, 6.75
[1585438362] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.20, 6.75, 5.49
[1585438412] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 13.89, 10.80, 7.22
[1585438422] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.79, 6.27, 5.40
[1585438472] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 10.74, 10.37, 7.29
[1585438482] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 10.64, 7.62, 5.92
[1585438532] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.61, 10.51, 7.54
[1585438542] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.71, 7.29, 5.91
[1585440032] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.95, 29.08, 21.46
[1585440427] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.01, 11.58, 10.43
[1585440487] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.39, 11.49, 10.47
[1585440547] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.03, 11.38, 10.50
[1585440607] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.43, 11.09, 10.45
[1585441114] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 83% (5 min average)
[1585441207] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.63, 9.45, 9.94
[1585441234] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 62% (5 min average)
[1585443064] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1585445132] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 4.88, 6.84, 17.18
[1585445742] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.51, 3.22, 4.42
[1585446032] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.33, 1.45, 7.66
[1585460404] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 88% capacity! 6.0T of 7.3T
[1585463464] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 89% capacity! 6.1T of 7.3T
[1585464274] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585464394] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;2;CPU Load 98% (5 min average)
[1585464514] SERVICE ALERT: WINV-107;CPU Load;OK;HARD;3;CPU Load 23% (5 min average)