[1584688577] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 34.87, 18.94, 8.49
[1584688637] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 38.77, 22.87, 10.49
[1584688697] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 49.25, 29.05, 13.46
[1584688757] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 63.70, 36.65, 17.04
[1584689255] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.39, 7.31, 5.50
[1584689315] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.84, 7.35, 5.63
[1584689376] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.13, 7.78, 5.90
[1584689436] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 10.55, 8.54, 6.28
[1584694354] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1584694474] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1584696542] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 14.94, 12.37, 10.67
[1584696602] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.86, 12.50, 10.81
[1584696662] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.85, 12.27, 10.85
[1584696722] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.04, 12.26, 10.93
[1584698794] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584698914] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1584698971] SERVICE ALERT: vinfra-200;MySQL Connection Time;WARNING;SOFT;1;WARNING - 2.12 seconds to connect as nagios
[1584699032] SERVICE ALERT: vinfra-200;MySQL Connection Time;OK;SOFT;2;OK - 0.01 seconds to connect as nagios
[1584699122] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.18, 7.25, 9.33
[1584701057] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.30, 8.03, 19.65
[1584701434] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584701436] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.88, 2.29, 4.74
[1584701554] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 59% (5 min average)
[1584701957] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.16, 0.58, 7.62
[1584704315] SERVICE ALERT: WINV-103;OS Drive Space;OK;HARD;3;c:\ - total: 68.27 Gb - used: 51.14 Gb (75%) - free 17.12 Gb (25%)
[1584706474] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584706594] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 69% (5 min average)
[1584706714] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584706834] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 65% (5 min average)
[1584746657] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.84, 9.62, 5.18
[1584746717] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 13.15, 9.72, 5.49
[1584747078] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 43.53, 30.01, 15.34
[1584747138] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 19.01, 25.51, 14.71
[1584747198] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 9.97, 21.74, 14.09
[1584747258] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.36, 18.60, 13.49
[1584747600] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 7.25, 7.72, 5.07
[1584747660] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 6.00, 7.27, 5.08
[1584747720] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 5.68, 6.96, 5.11
[1584747780] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 5.88, 6.78, 5.17
[1584748680] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 3.25, 4.10, 4.72
[1584748758] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 72.34, 38.99, 20.66
[1584748836] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.17, 6.16, 5.34
[1584748896] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.83, 7.24, 5.76
[1584748956] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.78, 6.84, 5.72
[1584749016] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.84, 7.42, 5.99
[1584751962] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584752082] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1584752544] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1584752664] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 30% (5 min average)
[1584753216] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.20, 3.41, 4.68
[1584753858] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.25, 1.80, 16.34
[1584754758] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.12, 1.31, 6.80
[1584764202] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584764322] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584764442] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584765042] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 10% (5 min average)