[1583050488] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.33, 11.85, 5.66
[1583050548] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 19.10, 13.25, 6.53
[1583050608] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 26.07, 16.24, 7.97
[1583050668] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 27.17, 18.31, 9.20
[1583050968] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.60, 25.09, 14.92
[1583051823] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.88, 6.23, 5.05
[1583051883] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.87, 5.83, 4.98
[1583052243] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.34, 7.30, 5.87
[1583052303] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.28, 7.55, 6.05
[1583052364] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.02, 7.57, 6.15
[1583052424] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.72, 7.00, 6.04
[1583053663] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.59, 11.59, 10.35
[1583053723] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.50, 11.32, 10.34
[1583053783] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.57, 11.23, 10.37
[1583053843] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.31, 11.33, 10.46
[1583058387] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583058468] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.66, 10.36, 19.32
[1583058507] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1583059668] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.70, 28.03, 21.67
[1583061643] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.54, 8.10, 9.89
[1583062924] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 6.13, 4.75, 4.90
[1583063225] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.05, 5.26, 5.10
[1583063285] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.84, 5.13, 5.06
[1583063345] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 3.87, 4.80, 4.95
[1583063405] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.72, 5.74, 5.28
[1583063465] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.31, 5.22, 5.13
[1583063525] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 2.52, 4.53, 4.90
[1583064168] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.26, 4.08, 17.09
[1583065068] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.06, 0.47, 6.68
[1583072950] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1583073010] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583073070] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1583073130] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1583080007] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1583080127] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1583080630] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583080690] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583080750] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1583080810] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583080870] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583080930] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1583081069] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583081189] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583081309] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1583081909] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1583083700] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 259 processes with STATE = RSZDT
[1583083760] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583083821] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1583083881] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1583084336] SERVICE ALERT: WINV-104;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1583084336] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;No data was received from host!
[1583084336] SERVICE ALERT: WINV-104;Uptime;WARNING;SOFT;1;No data was received from host!
[1583084336] SERVICE ALERT: WINV-104;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1583084456] SERVICE ALERT: WINV-104;Memory Usage;OK;SOFT;2;Memory usage: total:129142.02 MB - used: 44696.40 MB (35%) - free: 84445.62 MB (65%)
[1583084456] SERVICE ALERT: WINV-104;Uptime;OK;SOFT;2;System Uptime - 157 day(s) 0 hour(s) 8 minute(s)
[1583084456] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 73% (5 min average)
[1583084456] SERVICE ALERT: WINV-104;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 731.02 Gb (78%) - free 200.30 Gb (22%)
[1583085081] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583085141] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583085201] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1583085261] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1583088962] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.50 Gb (80%) - free 185.82 Gb (20%)
[1583089082] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.61 Gb (80%) - free 185.71 Gb (20%)
[1583089202] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.71 Gb (80%) - free 185.61 Gb (20%)
[1583089802] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 741.64 Gb (80%) - free 189.68 Gb (20%)
[1583090927] SERVICE ALERT: vdb-103;Root Partition;WARNING;HARD;4;/ is at 83% capacity! 360G of 457G
[1583097028] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 80% (5 min average)
[1583097148] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1583098446] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 84% capacity! 5.7T of 7.3T
[1583098746] SERVICE ALERT: vdb-103;Root Partition;OK;HARD;4;/ is at 79% capacity, 342G of 457G
[1583100936] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.12, 8.79, 5.91
[1583100996] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.71, 10.08, 6.54
[1583101056] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 9.06, 9.41, 6.54
[1583101116] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 7.49, 8.73, 6.47
[1583102016] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.99, 2.79, 4.35
[1583102916] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.89, 6.94, 4.77
[1583102976] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 5.82, 6.25, 4.67
[1583103036] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 2.64, 5.26, 4.42
[1583103880] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 264 processes with STATE = RSZDT
[1583103940] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes with STATE = RSZDT
[1583104000] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1583104061] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 257 processes with STATE = RSZDT
[1583106997] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.80, 4.21, 3.46
[1583107057] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.12, 6.39, 4.26
[1583107117] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 10.26, 6.99, 4.61
[1583107177] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.10, 6.12, 4.46
[1583107477] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 1.95, 3.42, 3.73
[1583110480] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 25.41, 18.08, 9.04
[1583110540] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 24.50, 19.15, 9.99
[1583110600] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 28.55, 21.32, 11.31
[1583110660] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 29.35, 22.88, 12.47
[1583110960] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.21, 29.36, 18.28
[1583111680] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.90, 5.69, 5.02
[1583111740] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.49, 6.23, 5.23
[1583111800] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.27, 6.17, 5.27
[1583111860] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.33, 5.50, 5.09
[1583114590] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 246 processes with STATE = RSZDT
[1583114860] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.62, 4.62, 4.97
[1583115160] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.57, 5.83, 5.36
[1583115220] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.80, 5.82, 5.39
[1583115280] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.56, 5.74, 5.39
[1583115340] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.09, 5.74, 5.42
[1583115640] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1583115640] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.49, 4.14, 4.79
[1583116090] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1583116150] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 246 processes with STATE = RSZDT
[1583116360] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.68, 7.78, 17.48
[1583117260] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.50, 0.89, 7.03
[1583120711] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1583120771] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 247 processes with STATE = RSZDT
[1583121340] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583129080] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1583129200] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 22% (5 min average)
[1583132388] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 14.06, 5.76, 2.93
[1583132491] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 8.19, 6.41, 3.48
[1583132551] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.19, 5.59, 3.38
[1583134981] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 36.34, 22.92, 11.70
[1583135041] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 34.97, 24.99, 13.13
[1583135101] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 26.42, 24.21, 13.59
[1583135161] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.16, 22.64, 13.74
[1583135320] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583135440] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 38% (5 min average)
[1583135500] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;No data was received from host!
[1583135620] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1583135803] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.03, 3.82, 7.53