[1582455060] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 258 processes with STATE = RSZDT
[1582455120] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 245 processes with STATE = RSZDT
[1582458120] SERVICE ALERT: vapps-100;Total Processes;OK;HARD;4;PROCS OK: 236 processes with STATE = RSZDT
[1582458161] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 20.11, 13.67, 6.08
[1582458221] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 18.83, 14.59, 6.88
[1582458281] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 18.63, 15.28, 7.60
[1582458341] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 15.62, 15.08, 8.02
[1582459096] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.90, 7.69, 5.20
[1582459156] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.58, 7.67, 5.36
[1582459216] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.06, 6.90, 5.23
[1582459241] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.13, 26.60, 19.03
[1582459276] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.38, 6.43, 5.18
[1582459380] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582459440] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes with STATE = RSZDT
[1582459500] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 246 processes with STATE = RSZDT
[1582459541] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.20, 20.76, 18.57
[1582459576] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.86, 4.66, 4.79
[1582459876] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.51, 5.90, 5.22
[1582459936] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.83, 5.29, 5.05
[1582459996] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 4.08, 5.01, 4.97
[1582460141] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1582460141] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 35.25, 27.38, 21.80
[1582460357] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.38, 5.41, 5.16
[1582460417] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.90, 5.17, 5.10
[1582460477] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.94, 5.76, 5.29
[1582460537] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.50, 5.49, 5.22
[1582464761] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1582464881] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1582465541] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582469237] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.07, 4.19, 4.75
[1582470041] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.80, 3.53, 16.18
[1582470941] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.75, 0.58, 6.42
[1582476761] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;1;/media/data is at 91% capacity! 6.2T of 7.3T
[1582476821] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;2;/media/data is at 91% capacity! 6.2T of 7.3T
[1582476881] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;SOFT;3;/media/data is at 91% capacity! 6.2T of 7.3T
[1582476941] SERVICE ALERT: vonly-backup;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1582487886] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582488006] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)
[1582499741] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 13.05, 11.42, 6.34
[1582499801] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 11.86, 11.40, 6.66
[1582499861] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 15.35, 12.25, 7.23
[1582499921] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.97, 13.50, 7.96
[1582500821] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.60, 4.48, 8.00
[1582505021] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.47, 11.02, 4.98
[1582505081] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 14.19, 11.35, 5.48
[1582505141] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 30.73, 16.49, 7.61
[1582505201] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 47.30, 23.63, 10.61
[1582507337] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.42, 6.33, 5.35
[1582507397] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.33, 6.15, 5.35
[1582507457] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.90, 6.64, 5.55
[1582507517] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.34, 6.33, 5.51
[1582508717] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.42, 2.60, 4.29
[1582508801] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.09, 6.85, 18.57
[1582510001] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.61, 1.42, 6.58
[1582516001] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 20.00, 15.09, 9.22
[1582516061] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 25.15, 17.52, 10.42
[1582516121] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 23.33, 18.37, 11.16
[1582516181] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 21.84, 18.82, 11.77
[1582516481] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.12, 26.15, 16.77
[1582517681] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.96, 6.58, 15.26
[1582518581] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1582518581] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.42, 1.08, 6.66
[1582520488] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582520608] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 29% (5 min average)
[1582523681] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)