[1584255632] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584255753] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1584256567] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 93% (5 min average)
[1584256688] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1584256691] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.39, 10.44, 5.10
[1584256751] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 20.26, 12.68, 6.20
[1584256808] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 64% (5 min average)
[1584256811] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 25.18, 15.35, 7.52
[1584256871] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 27.85, 17.90, 8.89
[1584257171] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.71, 27.07, 15.24
[1584257948] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 12.65, 8.47, 5.99
[1584258008] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 10.37, 8.54, 6.17
[1584258068] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.47, 8.01, 6.13
[1584258121] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (73 MB out of 975 MB)
[1584258128] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.16, 7.80, 6.18
[1584264035] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1584264155] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 20% (5 min average)
[1584264360] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 9.85, 11.27, 10.39
[1584264420] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 8.69, 10.70, 10.25
[1584264480] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 8.01, 10.18, 10.10
[1584264540] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.52, 9.81, 9.97
[1584264840] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 15.47, 12.98, 11.26
[1584264901] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 14.59, 13.16, 11.43
[1584264961] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.11, 13.02, 11.49
[1584265021] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.20, 12.48, 11.40
[1584266554] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584266674] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1584266821] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.14, 8.59, 9.66
[1584268872] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 19.50, 14.50, 17.94
[1584269172] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 51.05, 36.26, 26.30
[1584271328] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.02, 2.85, 4.52
[1584271572] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.49, 5.35, 16.51
[1584272472] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1584272472] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.03, 0.33, 6.32
[1584277572] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584280473] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584280593] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584280713] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584281075] SERVICE ALERT: WINV-103;OS Drive Space;OK;HARD;3;c:\ - total: 68.27 Gb - used: 48.01 Gb (70%) - free 20.25 Gb (30%)
[1584281313] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 8% (5 min average)
[1584313331] SERVICE ALERT: WINV-101;Uptime;WARNING;SOFT;1;No data was received from host!
[1584313331] SERVICE ALERT: WINV-101;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1584313451] SERVICE ALERT: WINV-101;Uptime;OK;SOFT;2;System Uptime - 10 day(s) 16 hour(s) 43 minute(s)
[1584313451] SERVICE ALERT: WINV-101;Memory Usage;OK;SOFT;2;Memory usage: total:98284.48 MB - used: 43854.54 MB (45%) - free: 54429.94 MB (55%)
[1584314772] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 52.16, 28.71, 14.26
[1584314832] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 37.03, 29.08, 15.32
[1584314892] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 17.31, 24.85, 14.72
[1584314952] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.16, 21.30, 14.15
[1584315852] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.69, 1.83, 6.09
[1584316753] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 12.86, 11.29, 8.58
[1584316813] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 18.89, 13.56, 9.55
[1584316873] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 23.09, 15.72, 10.55
[1584316933] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 22.40, 16.76, 11.23
[1584317228] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.80, 7.06, 5.19
[1584317233] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1584317233] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.03, 25.50, 16.31
[1584317288] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.71, 6.43, 5.09
[1584317348] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.66, 6.10, 5.06
[1584317408] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.59, 6.82, 5.38
[1584321608] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 89% (5 min average)
[1584321728] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 23% (5 min average)
[1584321733] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.52, 9.56, 16.71
[1584321908] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 5.56, 4.22, 4.86
[1584322200] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 743.31 Gb (80%) - free 188.00 Gb (20%)
[1584323233] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.97, 4.28, 7.03
[1584327733] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584335913] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584336033] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584336153] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)