[1582704532] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.48, 5.05, 1.98
[1582704592] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 10.97, 5.64, 2.38
[1582704953] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 24.46, 14.66, 7.12
[1582705013] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 24.13, 16.34, 8.17
[1582705073] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 30.24, 19.42, 9.73
[1582705133] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.01, 21.22, 10.95
[1582705429] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.90, 8.37, 5.97
[1582705489] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.46, 7.61, 5.86
[1582705552] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 8.68, 7.70, 5.99
[1582705612] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.24, 6.90, 5.82
[1582705912] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 1.90, 3.90, 4.82
[1582706632] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.69, 6.01, 5.08
[1582706692] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.39, 6.05, 5.16
[1582706752] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.83, 5.64, 5.07
[1582706812] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.24, 5.57, 5.08
[1582709805] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 13.30, 13.11, 10.92
[1582709865] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.37, 13.16, 11.08
[1582709925] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.41, 13.34, 11.28
[1582709985] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.10, 12.87, 11.24
[1582711612] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.45, 6.39, 3.96
[1582711672] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 5.68, 5.88, 3.94
[1582711905] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1582712025] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1582712446] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 278 processes with STATE = RSZDT
[1582712506] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 290 processes with STATE = RSZDT
[1582712588] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1582712648] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582712742] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.39, 13.58, 19.65
[1582713042] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.58, 6.75, 4.83
[1582713102] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 9.12, 6.86, 5.00
[1582713162] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.55, 5.96, 4.80
[1582714542] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.66, 33.12, 24.00
[1582716042] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.92, 8.71, 9.76
[1582717614] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.80, 4.21, 4.86
[1582718742] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.93, 7.83, 17.64
[1582719643] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.15, 0.95, 7.15
[1582721348] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582721408] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1582721468] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1582721528] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1582722428] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582722728] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582722788] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1582724680] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582724740] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 257 processes with STATE = RSZDT
[1582724800] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 260 processes with STATE = RSZDT
[1582724860] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 257 processes with STATE = RSZDT
[1582726360] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1582726360] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 249 processes with STATE = RSZDT
[1582726960] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582727020] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1582727080] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1582727140] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 256 processes with STATE = RSZDT
[1582729770] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.73, 4.19, 2.83
[1582729830] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 9.25, 5.01, 3.20
[1582729890] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.13, 4.32, 3.08
[1582730172] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582730292] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 60% (5 min average)
[1582730412] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582730532] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1582730672] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1582730792] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1582730912] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1582731512] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1582731788] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.03, 5.59, 3.48
[1582731848] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.76, 6.19, 3.93
[1582731908] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.38, 5.62, 3.87
[1582732657] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1582737352] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.50 Gb (80%) - free 185.82 Gb (20%)
[1582737472] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.63 Gb (80%) - free 185.69 Gb (20%)
[1582737592] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.76 Gb (80%) - free 185.56 Gb (20%)
[1582740757] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 249 processes with STATE = RSZDT
[1582741057] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1582741117] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 256 processes with STATE = RSZDT
[1582741177] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 253 processes with STATE = RSZDT
[1582741237] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 255 processes with STATE = RSZDT
[1582741837] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582742137] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 256 processes with STATE = RSZDT
[1582742197] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 258 processes with STATE = RSZDT
[1582742257] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 264 processes with STATE = RSZDT
[1582742317] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 249 processes with STATE = RSZDT
[1582742917] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes with STATE = RSZDT
[1582742977] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 253 processes with STATE = RSZDT
[1582743037] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1582743097] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 249 processes with STATE = RSZDT
[1582743997] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1582744057] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 294 processes with STATE = RSZDT
[1582744117] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 256 processes with STATE = RSZDT
[1582744178] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1582744178] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 253 processes with STATE = RSZDT
[1582744478] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582744778] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1582744838] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes with STATE = RSZDT
[1582744898] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 254 processes with STATE = RSZDT
[1582744958] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 293 processes with STATE = RSZDT
[1582746157] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 85% (5 min average)
[1582746278] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 75% (5 min average)
[1582749592] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 741.90 Gb (80%) - free 189.42 Gb (20%)
[1582750658] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582764757] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 25.87, 13.40, 5.78
[1582764817] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 20.44, 14.12, 6.50
[1582764877] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 21.71, 15.43, 7.42
[1582764938] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 20.73, 16.19, 8.18
[1582765238] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 31.70, 24.81, 13.91
[1582765966] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582766086] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 647.07 Gb (69%) - free 284.25 Gb (31%)
[1582766215] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582766257] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.56, 5.94, 5.05
[1582766317] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.13, 5.47, 4.94
[1582766335] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 646.61 Gb (69%) - free 284.71 Gb (31%)
[1582766377] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.23, 5.75, 5.07
[1582766437] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.01, 5.49, 5.02
[1582766497] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.90, 5.85, 5.18
[1582766557] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.41, 5.47, 5.10
[1582767064] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582767184] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 645.41 Gb (69%) - free 285.91 Gb (31%)
[1582769257] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.05, 3.82, 4.76
[1582770038] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.10, 5.36, 16.37
[1582770313] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582770433] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 637.21 Gb (68%) - free 294.11 Gb (32%)
[1582770938] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.19, 0.64, 6.54
[1582771162] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582771282] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 635.83 Gb (68%) - free 295.49 Gb (32%)
[1582773211] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582773331] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 636.75 Gb (68%) - free 294.57 Gb (32%)
[1582774060] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582774180] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 636.91 Gb (68%) - free 294.41 Gb (32%)
[1582774746] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;No data was received from host!
[1582774875] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;2;No data was received from host!
[1582774995] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 7% (5 min average)
[1582776195] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1582776315] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1582776435] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582776555] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1582776675] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1582777275] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)
[1582777909] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582778029] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 638.02 Gb (69%) - free 293.30 Gb (31%)
[1582778758] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582778878] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 638.38 Gb (69%) - free 292.94 Gb (31%)
[1582779607] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1582779727] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 638.68 Gb (69%) - free 292.64 Gb (31%)
[1582783058] SERVICE ALERT: vdb-101;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 90% capacity! 3.1T of 3.6T
[1582783545] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582783580] SERVICE ALERT: WINV-102;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1582783580] SERVICE ALERT: WINV-102;OS Drive Space;WARNING;SOFT;1;No data was received from host!
[1582783665] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1582783700] SERVICE ALERT: WINV-102;OS Drive Space;OK;SOFT;2;c:\ - total: 931.41 Gb - used: 305.23 Gb (33%) - free 626.19 Gb (67%)
[1582783700] SERVICE ALERT: WINV-102;Memory Usage;OK;SOFT;2;Memory usage: total:208877.43 MB - used: 69024.90 MB (33%) - free: 139852.53 MB (67%)
[1582783785] SERVICE ALERT: WINV-107;CPU Load;OK;HARD;3;CPU Load 23% (5 min average)
[1582784857] SERVICE ALERT: vdb-103;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (205 MB out of 2047 MB)
[1582785458] SERVICE ALERT: vdb-101;/mnt/data01 Partition;CRITICAL;HARD;4;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1582789038] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.97, 6.22, 4.56
[1582789098] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 4.79, 5.63, 4.46
[1582789957] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582790058] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.02, 5.84, 4.20
[1582790077] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1582790118] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.86, 5.80, 4.30
[1582790178] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 3.88, 5.07, 4.14
[1582790292] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 246 processes with STATE = RSZDT