[1583481609] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1583482731] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 21.65, 13.03, 6.07
[1583482791] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 30.48, 16.89, 7.82
[1583482851] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 32.39, 19.78, 9.38
[1583482911] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.73, 22.12, 10.83
[1583483409] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583483469] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583483529] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1583483565] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.87, 12.81, 10.35
[1583483626] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.62, 12.75, 10.48
[1583483686] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.02, 12.86, 10.67
[1583483746] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.64, 13.10, 10.89
[1583483804] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.76, 6.40, 5.18
[1583483865] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.58, 6.97, 5.45
[1583483925] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.83, 6.98, 5.54
[1583483985] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.07, 6.28, 5.39
[1583484789] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583484849] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583484909] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 249 processes with STATE = RSZDT
[1583486169] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583486229] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1583489811] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.25, 14.55, 18.91
[1583490711] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.95, 23.98, 21.18
[1583491723] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583491843] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 24% (5 min average)
[1583494065] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583494185] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1583494246] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.36, 8.68, 9.94
[1583494305] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1583496585] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.82, 3.35, 4.54
[1583497311] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.10, 1.97, 14.54
[1583498211] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.39, 0.62, 5.99
[1583504289] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583504349] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1583537511] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 26.49, 15.95, 9.00
[1583537571] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 35.47, 19.73, 10.69
[1583537631] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 32.93, 22.79, 12.37
[1583537691] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 45.92, 27.73, 14.69
[1583537992] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.74, 19.56, 15.35
[1583540092] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.77, 4.64, 7.47
[1583540508] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583540628] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1583542192] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 26.65, 9.02, 4.17
[1583542253] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 21.14, 10.61, 5.03
[1583542313] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 20.25, 12.06, 5.87
[1583542373] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.3% change >= 20.0% threshold)
[1583542373] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 16.31, 12.46, 6.39
[1583542846] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.29, 11.87, 10.06
[1583542906] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.67, 11.79, 10.15
[1583542966] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.08, 11.85, 10.28
[1583542973] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.22, 24.16, 15.30
[1583543026] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.17, 12.13, 10.47
[1583543685] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.22, 6.41, 5.21
[1583543745] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.33, 5.87, 5.10
[1583543805] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.41, 5.76, 5.11
[1583543865] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.56, 5.17, 4.94
[1583545065] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.57, 6.22, 5.50
[1583545125] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.67, 5.95, 5.46
[1583545185] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.63, 5.54, 5.34
[1583545245] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.37, 5.32, 5.28
[1583546983] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 736.64 Gb (79%) - free 194.67 Gb (21%)
[1583547946] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.70, 4.24, 4.97
[1583548373] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.23, 4.51, 14.90
[1583548726] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.12, 7.40, 9.43
[1583548973] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.56, 0.99, 7.98
[1583552710] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583552770] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 246 processes with STATE = RSZDT
[1583554373] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1583559949] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583560069] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583560189] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 10% (5 min average)
[1583561698] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1583561818] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 48% (5 min average)
[1583567338] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583567458] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583567578] SERVICE ALERT: WINV-102;CPU Load;WARNING;HARD;3;CPU Load 83% (5 min average)