[1584861199] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.76, 8.74, 3.90
[1584861259] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 18.09, 10.50, 4.80
[1584861319] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 26.63, 14.48, 6.54
[1584861379] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.44, 18.57, 8.42
[1584861996] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.25, 7.18, 5.10
[1584862056] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.09, 7.30, 5.26
[1584862116] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.19, 7.63, 5.51
[1584862177] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 10.21, 8.05, 5.78
[1584862983] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.19, 12.39, 10.40
[1584863043] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.48, 12.14, 10.43
[1584863103] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.76, 12.11, 10.53
[1584863163] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.47, 12.27, 10.69
[1584865362] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1584865482] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)
[1584866163] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.78, 8.74, 9.98
[1584868279] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.39, 11.49, 19.75
[1584868863] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 15.27, 12.54, 10.18
[1584868879] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 19.88, 24.47, 22.79
[1584868923] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 16.90, 13.45, 10.65
[1584868983] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 16.20, 13.93, 10.99
[1584869043] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 14.22, 13.80, 11.13
[1584870979] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.60, 13.91, 18.36
[1584871279] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1584871279] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 40.25, 26.57, 21.89
[1584872043] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.23, 6.34, 8.88
[1584873380] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 16.25, 12.09, 17.40
[1584873980] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.98, 26.13, 22.02
[1584876877] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.35, 2.69, 4.69
[1584876980] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.34, 5.60, 16.97
[1584877880] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.28, 0.57, 6.62
[1584878724] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.26, 5.28, 5.09
[1584878784] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.55, 5.36, 5.13
[1584878844] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.07, 5.51, 5.19
[1584878904] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.51, 5.46, 5.20
[1584879804] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.43, 4.03, 4.63
[1584882980] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584883104] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.44, 5.46, 5.06
[1584883164] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 4.10, 5.05, 4.94
[1584887424] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.84, 5.45, 5.09
[1584887484] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.98, 5.58, 5.16
[1584887544] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.85, 5.62, 5.20
[1584887604] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.46, 5.56, 5.21
[1584888205] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.15, 4.68, 4.93
[1584891205] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.23, 5.30, 5.05
[1584891265] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.75, 5.42, 5.11
[1584891325] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.88, 5.54, 5.17
[1584891385] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 4.59, 5.19, 5.07
[1584891686] SERVICE FLAPPING ALERT: vonly-database;Current Load;STARTED; Service appears to have started flapping (21.7% change >= 20.0% threshold)
[1584891686] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.10, 4.75, 4.94
[1584894386] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.14, 5.79, 5.11
[1584894446] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.00, 5.83, 5.16
[1584894506] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.56, 5.73, 5.17
[1584894566] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.68, 5.75, 5.21
[1584895466] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.70, 4.71, 4.99
[1584900566] SERVICE FLAPPING ALERT: vonly-database;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1584910766] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.82, 5.98, 5.35
[1584910826] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.99, 5.98, 5.39
[1584910886] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.32, 6.10, 5.47
[1584910946] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.16, 6.09, 5.50
[1584912146] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 5.43, 4.70, 4.94
[1584912446] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.70, 5.27, 5.12
[1584912506] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 4.34, 4.99, 5.03
[1584912566] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 4.52, 4.93, 5.01
[1584912626] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.61, 4.90, 4.99
[1584912926] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.23, 5.17, 5.07
[1584912986] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 4.64, 5.05, 5.04
[1584913046] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 4.71, 4.98, 5.01
[1584913106] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.19, 4.76, 4.93
[1584913406] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.76, 5.33, 5.12
[1584913466] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 4.92, 5.19, 5.08
[1584913526] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;3;OK - load average: 4.14, 4.92, 5.00
[1584914787] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.78, 5.61, 5.19
[1584914847] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.78, 5.65, 5.23
[1584914907] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.18, 5.77, 5.30
[1584914967] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.87, 5.74, 5.31
[1584916467] SERVICE FLAPPING ALERT: vonly-database;Current Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1584916467] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.64, 4.62, 4.94
[1584919880] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 26.49, 14.93, 8.32
[1584919940] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 22.89, 16.37, 9.25
[1584920000] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 35.70, 21.64, 11.54
[1584920061] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 42.46, 26.29, 13.78
[1584920077] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.42, 7.88, 5.42
[1584920137] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.98, 7.56, 5.46
[1584920198] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.23, 7.63, 5.62
[1584920258] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.51, 7.80, 5.79
[1584920361] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 15.55, 22.14, 15.61
[1584920967] SERVICE FLAPPING ALERT: vonly-database;Current Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1584921561] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.68, 25.46, 19.28
[1584926258] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.53, 3.11, 4.76
[1584926361] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.43, 4.23, 17.11
[1584927561] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 3.17, 2.83, 6.64
[1584943833] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;No data was received from host!
[1584943953] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 14% (5 min average)