[1582963213] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 6.61, 7.80, 7.43
[1582963945] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 1.90, 2.26, 4.43
[1582964111] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 14.23, 10.27, 5.20
[1582964171] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 12.60, 10.48, 5.59
[1582964231] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 16.15, 11.79, 6.34
[1582964291] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 24.67, 15.19, 7.86
[1582964891] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.27, 25.94, 17.23
[1582965074] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582965194] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1582965910] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.61, 5.74, 5.08
[1582965913] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1582965913] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.68, 3.91, 4.80
[1582965970] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.75, 6.11, 5.23
[1582966030] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.85, 6.81, 5.53
[1582966090] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.67, 6.27, 5.43
[1582966213] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 3.88, 6.25, 5.89
[1582966273] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 3.10, 5.57, 5.67
[1582966333] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 2.81, 5.05, 5.49
[1582966345] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.70, 12.16, 10.45
[1582966393] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 2.74, 4.61, 5.30
[1582966405] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.45, 11.55, 10.35
[1582966465] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.04, 11.79, 10.50
[1582966525] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.87, 12.10, 10.69
[1582966693] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 2.71, 3.50, 4.63
[1582967110] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;1;/ is at 81% capacity! 351G of 457G
[1582967110] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 6.45, 6.89, 5.09
[1582967170] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;2;/ is at 82% capacity! 352G of 457G
[1582967171] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 6.57, 6.86, 5.19
[1582967230] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;3;/ is at 82% capacity! 353G of 457G
[1582967231] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 6.55, 6.79, 5.27
[1582967290] SERVICE ALERT: vdb-103;Root Partition;WARNING;HARD;4;/ is at 82% capacity! 354G of 457G
[1582967291] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 6.64, 6.82, 5.38
[1582967793] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 333 processes with STATE = RSZDT
[1582967853] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes with STATE = RSZDT
[1582968095] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.22, 6.81, 5.18
[1582968095] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1582968145] SERVICE FLAPPING ALERT: vdb-staging;Current Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1582968155] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.07, 7.06, 5.68
[1582968155] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 248 processes with STATE = RSZDT
[1582968215] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.05, 6.32, 5.52
[1582968275] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 3.53, 5.74, 5.36
[1582968525] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582968585] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 247 processes with STATE = RSZDT
[1582968875] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 2.87, 4.09, 4.86
[1582969127] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582969175] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 4.22, 5.84, 5.57
[1582969235] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 3.74, 5.44, 5.44
[1582969247] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 20% (5 min average)
[1582969295] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 3.12, 4.96, 5.28
[1582969355] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 2.66, 4.50, 5.10
[1582969655] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.81, 3.96, 4.67
[1582970327] SERVICE ALERT: vdb-103;Root Partition;CRITICAL;HARD;4;/ is at 91% capacity! 391G of 457G
[1582971101] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.09, 8.70, 9.75
[1582971155] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 14.24, 6.79, 4.98
[1582971215] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.92, 7.94, 5.50
[1582971227] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.02, 11.50, 18.85
[1582971275] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 6.96, 7.13, 5.38
[1582971335] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.95, 6.50, 5.27
[1582971635] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.36, 4.52, 4.74
[1582971701] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.67, 10.66, 10.14
[1582971761] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.48, 10.54, 10.13
[1582971821] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.77, 10.61, 10.18
[1582971881] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 9.18, 10.20, 10.07
[1582972181] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1582972181] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.23, 9.04, 9.56
[1582972481] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.85, 10.39, 10.02
[1582972542] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.74, 10.75, 10.17
[1582972602] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.99, 11.28, 10.39
[1582972662] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.48, 11.45, 10.50
[1582972727] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.92, 28.00, 19.50
[1582973027] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.59, 20.34, 18.89
[1582973327] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 4.85, 4.59, 4.91
[1582974762] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.18, 8.59, 9.91
[1582974827] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.06, 4.34, 4.77
[1582975127] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1582975127] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.07, 23.28, 17.47
[1582975427] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.62, 5.88, 5.24
[1582975487] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.96, 6.25, 5.40
[1582975548] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.10, 5.74, 5.28
[1582975608] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 3.77, 5.26, 5.14
[1582975727] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 5.19, 5.13, 5.03
[1582975787] SERVICE ALERT: vdb-103;Current Load;OK;SOFT;2;OK - load average: 4.93, 5.05, 5.00
[1582975847] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 5.08, 5.07, 5.01
[1582975907] SERVICE ALERT: vdb-103;Current Load;OK;SOFT;2;OK - load average: 4.65, 4.98, 4.98
[1582976508] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.05, 4.53, 4.78
[1582976867] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 5.59, 5.17, 5.01
[1582976927] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 5.74, 5.26, 5.05
[1582976987] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 6.19, 5.47, 5.14
[1582977047] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 6.53, 5.72, 5.24
[1582977335] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582977647] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 4.36, 4.73, 4.98
[1582978570] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582978630] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 252 processes with STATE = RSZDT
[1582978662] SERVICE FLAPPING ALERT: vonly-replica;Current Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1582978690] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1582978750] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1582979027] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.26, 3.77, 17.69
[1582979927] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.26, 0.55, 6.96
[1582982050] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582985027] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582985187] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582985307] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1582985427] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1582986027] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1582995722] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.42 Gb (80%) - free 185.90 Gb (20%)
[1582995842] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.54 Gb (80%) - free 185.78 Gb (20%)
[1582995962] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.68 Gb (80%) - free 185.63 Gb (20%)
[1583000162] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 734.50 Gb (79%) - free 196.82 Gb (21%)
[1583019227] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.44, 20.92, 10.82
[1583019287] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 33.04, 23.31, 12.28
[1583019347] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.71, 24.18, 13.28
[1583019407] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.66, 22.06, 13.25
[1583020007] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.29, 4.36, 7.73
[1583023607] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 37.27, 19.74, 9.38
[1583023667] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 40.20, 23.82, 11.43
[1583023727] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 28.14, 23.15, 11.96
[1583023787] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 23.99, 23.02, 12.63
[1583024087] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.2% change >= 20.0% threshold)
[1583024087] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 35.66, 26.79, 16.90
[1583024523] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.10, 6.28, 5.16
[1583024583] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.48, 5.90, 5.10
[1583024643] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.00, 5.77, 5.10
[1583024703] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.96, 6.38, 5.37
[1583027703] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.67, 4.43, 5.00
[1583028303] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.35, 5.26, 5.07
[1583028363] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 3.45, 4.73, 4.90
[1583028587] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 21.74, 18.23, 19.69
[1583028887] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.98, 25.00, 22.20
[1583029367] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 93% (5 min average)
[1583029487] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 21% (5 min average)
[1583029488] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.35, 8.55, 16.32
[1583030388] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.37, 0.99, 6.55
[1583035488] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583047829] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583047949] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)