[1583745175] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.80, 15.00, 19.30
[1583746075] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 29.68, 22.48, 20.58
[1583748001] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 85% (5 min average)
[1583748121] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 25% (5 min average)
[1583748241] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583748361] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583748481] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1583749081] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 6% (5 min average)
[1583750275] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.69, 13.21, 19.26
[1583752975] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.51, 21.90, 20.46
[1583753875] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.56, 16.17, 19.89
[1583754403] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1583754407] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.70, 4.03, 4.89
[1583754523] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1583755075] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1583755075] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.46, 1.34, 7.40
[1583756928] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.85, 9.32, 9.87
[1583758663] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.06 Gb (80%) - free 186.26 Gb (20%)
[1583758783] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.22 Gb (80%) - free 186.10 Gb (20%)
[1583758903] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.40 Gb (80%) - free 185.92 Gb (20%)
[1583759875] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583762628] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.03, 10.18, 10.08
[1583762688] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.77, 10.08, 10.05
[1583762748] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.42, 10.23, 10.11
[1583762808] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.32, 10.23, 10.11
[1583765208] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.69, 9.76, 9.94
[1583797008] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 15.73, 12.67, 10.48
[1583797068] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 17.49, 13.69, 10.97
[1583797075] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 38.51, 22.69, 12.10
[1583797128] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 17.75, 14.48, 11.42
[1583797135] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.35, 22.68, 12.78
[1583797188] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 14.43, 14.29, 11.55
[1583797195] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 15.18, 20.12, 12.52
[1583797255] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.60, 18.21, 12.34
[1583797788] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.51, 8.51, 9.68
[1583798455] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.04, 25.17, 20.19
[1583798807] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.38, 6.22, 5.28
[1583798867] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.18, 5.78, 5.19
[1583798927] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.83, 5.43, 5.10
[1583798987] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.71, 5.27, 5.06
[1583802655] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.26, 13.95, 19.58
[1583803187] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.79, 2.34, 4.20
[1583803855] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.14, 0.87, 6.78
[1583814875] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 11% free (106 MB out of 975 MB)
[1583814935] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 9% free (83 MB out of 975 MB)
[1583814995] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 7% free (61 MB out of 975 MB)
[1583815055] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 5% free (39 MB out of 975 MB)