[1585120345] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 18.37, 9.38, 4.08
[1585120405] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 23.77, 12.58, 5.51
[1585120466] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 22.58, 14.41, 6.59
[1585120526] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 27.39, 17.07, 7.99
[1585120826] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.08, 27.19, 15.18
[1585121167] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.57, 11.71, 10.06
[1585121227] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.57, 11.75, 10.18
[1585121287] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.58, 11.45, 10.18
[1585121320] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.10, 7.38, 5.48
[1585121347] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.59, 11.58, 10.30
[1585121380] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.41, 7.01, 5.47
[1585121440] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.50, 7.51, 5.73
[1585121500] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.50, 7.11, 5.71
[1585127195] SERVICE ALERT: WINV-104;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1585127315] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 29% (5 min average)
[1585132226] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.92, 9.99, 16.88
[1585132234] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585132354] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)
[1585132983] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (5719 MB out of 57111 MB)
[1585133283] SERVICE ALERT: vonly-database;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (5677 MB out of 57111 MB)
[1585133584] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (5742 MB out of 57111 MB)
[1585133726] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.68, 25.43, 19.56
[1585136426] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.74, 9.00, 17.86
[1585136500] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.21, 3.49, 4.78
[1585136584] SERVICE FLAPPING ALERT: vonly-database;Swap Usage;STARTED; Service appears to have started flapping (20.7% change >= 20.0% threshold)
[1585136584] SERVICE ALERT: vonly-database;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (5662 MB out of 57111 MB)
[1585136647] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.15, 9.38, 9.98
[1585137626] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1585137626] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.92, 1.89, 6.37
[1585137635] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1585137755] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 56% (5 min average)
[1585137810] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.11, 5.67, 5.04
[1585137872] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.90, 5.96, 5.18
[1585137932] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 7.03, 6.18, 5.31
[1585137992] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.18, 6.10, 5.34
[1585138086] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (5712 MB out of 57111 MB)
[1585138386] SERVICE ALERT: vonly-database;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (5621 MB out of 57111 MB)
[1585139192] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.18, 4.59, 4.96
[1585141686] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (5742 MB out of 57111 MB)
[1585141892] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 7.99, 6.60, 5.58
[1585141952] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 7.03, 6.56, 5.62
[1585142012] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 7.15, 6.68, 5.73
[1585142072] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.68, 6.64, 5.77
[1585146156] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585146279] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585146403] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1585147003] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 19% (5 min average)
[1585148372] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.31, 4.03, 4.82
[1585149380] SERVICE ALERT: WINV-101;Uptime;WARNING;SOFT;1;No data was received from host!
[1585149388] SERVICE ALERT: WINV-106;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585149503] SERVICE ALERT: WINV-101;Uptime;OK;SOFT;2;System Uptime - 20 day(s) 8 hour(s) 57 minute(s)
[1585149508] SERVICE ALERT: WINV-106;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 589.23 Gb (63%) - free 342.09 Gb (37%)
[1585149632] SERVICE ALERT: WINV-101;Uptime;WARNING;SOFT;1;No data was received from host!
[1585149653] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;No data was received from host!
[1585149730] SERVICE ALERT: WINV-101;OS Drive Space;WARNING;SOFT;1;No data was received from host!
[1585149752] SERVICE ALERT: WINV-101;Uptime;OK;SOFT;2;System Uptime - 20 day(s) 9 hour(s) 1 minute(s)
[1585149773] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585149850] SERVICE ALERT: WINV-101;OS Drive Space;OK;SOFT;2;c:\ - total: 68.27 Gb - used: 43.50 Gb (64%) - free 24.77 Gb (36%)
[1585149893] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 8% (5 min average)
[1585151886] SERVICE ALERT: vonly-database;Swap Usage;OK;HARD;4;SWAP OK - 100% free (56880 MB out of 57111 MB)
[1585159003] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1585159123] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;2;CPU Load 89% (5 min average)
[1585159243] SERVICE ALERT: WINV-102;CPU Load;WARNING;HARD;3;CPU Load 88% (5 min average)
[1585159843] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CPU Load 90% (5 min average)
[1585160443] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 33% (5 min average)
[1585166693] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 85% (5 min average)
[1585166813] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 29% (5 min average)
[1585167131] SERVICE ALERT: WINV-105;Data Drive Space;WARNING;SOFT;1;No data was received from host!
[1585167251] SERVICE ALERT: WINV-105;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 574.34 Gb (62%) - free 356.98 Gb (38%)
[1585173533] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585173653] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1585173773] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 49% (5 min average)
[1585179026] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 13.63, 10.15, 6.66
[1585179086] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 12.77, 10.52, 7.01
[1585179146] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 10.23, 10.21, 7.12
[1585179206] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.27, 10.50, 7.42
[1585179547] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.71, 10.90, 10.02
[1585179607] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.85, 11.10, 10.14
[1585179667] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 14.12, 11.88, 10.47
[1585179700] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.71, 7.05, 5.48
[1585179727] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 15.65, 12.71, 10.84
[1585179760] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.82, 7.44, 5.70
[1585179820] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.80, 6.82, 5.59
[1585179880] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.75, 6.98, 5.73
[1585180706] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.36, 20.05, 15.03
[1585183987] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (11329 MB out of 57111 MB)
[1585184047] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (11285 MB out of 57111 MB)
[1585184107] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 20% free (11266 MB out of 57111 MB)
[1585184167] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 20% free (11230 MB out of 57111 MB)
[1585185806] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.63, 9.06, 18.12
[1585186780] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.40, 3.55, 4.63
[1585187006] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.59, 4.51, 7.85
[1585189327] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.14, 9.40, 9.90
[1585200667] SERVICE ALERT: vonly-database;Swap Usage;OK;HARD;4;SWAP OK - 21% free (11446 MB out of 57111 MB)
[1585202167] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (11328 MB out of 57111 MB)
[1585202227] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (11284 MB out of 57111 MB)
[1585202287] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 20% free (11290 MB out of 57111 MB)
[1585202347] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 20% free (11292 MB out of 57111 MB)
[1585203847] SERVICE ALERT: vonly-database;Swap Usage;OK;HARD;4;SWAP OK - 21% free (11763 MB out of 57111 MB)
[1585203872] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.47, 6.42, 5.50
[1585203933] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.57, 6.48, 5.58
[1585203993] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 7.07, 6.63, 5.69
[1585204053] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.40, 6.51, 5.71