[1583224006] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.15, 11.33, 6.00
[1583224066] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 29.38, 15.92, 7.90
[1583224126] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 31.74, 19.17, 9.53
[1583224186] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 24.88, 19.69, 10.33
[1583224786] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.60, 26.25, 18.54
[1583224966] HOST ALERT: WINV-107;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.3.107)
[1583225030] HOST ALERT: WINV-107;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.34 ms
[1583225323] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.46, 6.40, 5.38
[1583225383] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.89, 5.92, 5.28
[1583225443] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.25, 6.21, 5.43
[1583225503] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 3.84, 5.53, 5.24
[1583226344] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.90, 12.98, 10.23
[1583226404] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.22, 13.10, 10.45
[1583226464] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.64, 13.26, 10.67
[1583226524] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.70, 12.57, 10.60
[1583227124] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.25, 8.95, 9.61
[1583228863] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;1;/ is at 81% capacity! 347G of 457G
[1583228923] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;2;/ is at 81% capacity! 348G of 457G
[1583228983] SERVICE ALERT: vdb-103;Root Partition;WARNING;SOFT;3;/ is at 81% capacity! 349G of 457G
[1583229043] SERVICE ALERT: vdb-103;Root Partition;WARNING;HARD;4;/ is at 81% capacity! 349G of 457G
[1583229943] SERVICE ALERT: vdb-103;Root Partition;OK;HARD;4;/ is at 20% capacity, 86G of 457G
[1583237203] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.18, 3.58, 4.59
[1583237386] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.30, 5.84, 15.28
[1583238286] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.65, 0.85, 6.06
[1583246605] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583246725] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 14% (5 min average)
[1583278486] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 31.37, 16.60, 8.77
[1583278546] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 40.31, 21.66, 10.99
[1583278606] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 36.52, 23.83, 12.40
[1583278667] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.33, 26.62, 14.11
[1583278967] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.61, 11.93, 11.26
[1583279567] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.06, 3.00, 6.67
[1583282847] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583282867] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.11, 12.15, 6.22
[1583282927] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 25.29, 15.02, 7.59
[1583282968] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583282988] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 25.80, 16.81, 8.66
[1583283048] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.0% change >= 20.0% threshold)
[1583283048] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.01, 20.16, 10.30
[1583283088] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 9% (5 min average)
[1583284003] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.53, 5.99, 5.05
[1583284063] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 5.08, 5.67, 5.00
[1583284123] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.54, 6.26, 5.24
[1583284183] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.56, 5.92, 5.19
[1583284243] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.47, 5.50, 5.09
[1583284303] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.90, 5.40, 5.08
[1583285803] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 5.20, 4.56, 4.93
[1583286048] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1583286103] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.71, 6.05, 5.50
[1583286163] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.84, 6.89, 5.82
[1583286223] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.74, 7.01, 5.94
[1583286283] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.20, 6.41, 5.80
[1583287548] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 7.80, 9.63, 17.01
[1583288383] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1583288383] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 5.02, 3.98, 4.54
[1583288448] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 43.02, 22.14, 19.53
[1583288748] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.8% change >= 20.0% threshold)
[1583288748] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.37, 18.65, 19.20
[1583290248] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.79, 1.75, 6.71
[1583292283] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1583294748] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1583299888] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1583300008] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 44% (5 min average)
[1583301343] SERVICE ALERT: vdb-103;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 61% capacity, 4.2T of 7.3T