[1582016163] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.00, 6.01, 5.14
[1582016223] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.04, 6.02, 5.20
[1582016283] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.66, 5.96, 5.23
[1582016343] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.48, 5.86, 5.24
[1582024143] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 2.58, 3.37, 4.46
[1582028343] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 7.03, 6.41, 5.44
[1582028403] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.88, 6.51, 5.54
[1582028463] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.94, 6.59, 5.63
[1582028523] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.78, 6.61, 5.70
[1582030750] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 2.05, 6.52, 8.80
[1582030810] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 2.02, 5.69, 8.37
[1582030850] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 5.17, 5.70, 6.15
[1582030870] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 1.50, 4.86, 7.91
[1582030930] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.71, 4.30, 7.53
[1582030937] SERVICE ALERT: vapps-100;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 1% free (0 MB out of 974 MB)
[1582030975] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1457 processes with STATE = RSZDT
[1582030997] SERVICE ALERT: vapps-100;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 1% free (0 MB out of 974 MB)
[1582031036] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1457 processes with STATE = RSZDT
[1582031057] SERVICE ALERT: vapps-100;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 1% free (0 MB out of 974 MB)
[1582031096] SERVICE ALERT: vapps-100;Total Processes;OK;SOFT;3;PROCS OK: 1459 processes with STATE = RSZDT
[1582031117] SERVICE ALERT: vapps-100;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 1% free (0 MB out of 974 MB)
[1582031350] SERVICE ALERT: vapps-100;Swap Usage;OK;HARD;4;SWAP OK - 92% free (10240 MB out of 11214 MB)
[1582031456] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;1;PROCS WARNING: 1510 processes with STATE = RSZDT
[1582031514] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.21, 2.21, 4.90
[1582031516] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;2;PROCS WARNING: 1511 processes with STATE = RSZDT
[1582031577] SERVICE ALERT: vapps-100;Total Processes;OK;SOFT;3;PROCS OK: 1510 processes with STATE = RSZDT
[1582041537] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;1;PROCS WARNING: 2008 processes with STATE = RSZDT
[1582041597] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;2;PROCS WARNING: 2008 processes with STATE = RSZDT
[1582041658] SERVICE ALERT: vapps-100;Total Processes;WARNING;SOFT;3;PROCS WARNING: 2008 processes with STATE = RSZDT
[1582041718] SERVICE ALERT: vapps-100;Total Processes;WARNING;HARD;4;PROCS WARNING: 2009 processes with STATE = RSZDT
[1582043642] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 5.55, 5.61, 5.12
[1582043702] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 5.26, 5.53, 5.12
[1582043762] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.51, 5.55, 5.16
[1582043822] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.52, 5.55, 5.18
[1582048023] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.13, 4.59, 4.99
[1582054923] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.38, 5.20, 4.66
[1582054983] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 6.78, 5.47, 4.79
[1582067814] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 20.03, 12.73, 6.83
[1582067874] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 22.85, 14.65, 7.84
[1582067934] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 25.77, 16.81, 9.00
[1582067995] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 41.25, 22.68, 11.50
[1582068295] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.46, 10.36, 9.50
[1582068595] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.49, 4.39, 7.21
[1582072018] SERVICE ALERT: vapps-100;Total Processes;OK;HARD;4;PROCS OK: 211 processes with STATE = RSZDT
[1582073695] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.96, 16.23, 8.07
[1582073755] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 16.89, 15.71, 8.40
[1582073815] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 17.41, 16.31, 9.08
[1582073875] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 16.34, 16.07, 9.44
[1582074475] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.43, 21.17, 13.84
[1582078375] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.11, 13.37, 18.57
[1582080475] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.28, 1.35, 6.19
[1582096079] SERVICE ALERT: vinfra-200;Current Users;CRITICAL;SOFT;1;NRPE: Command 'check_users' not defined
[1582096139] SERVICE ALERT: vinfra-200;Current Users;CRITICAL;SOFT;2;NRPE: Command 'check_users' not defined
[1582096199] SERVICE ALERT: vinfra-200;Current Users;CRITICAL;SOFT;3;NRPE: Command 'check_users' not defined
[1582096259] SERVICE ALERT: vinfra-200;Current Users;CRITICAL;HARD;4;NRPE: Command 'check_users' not defined
[1582097306] SERVICE ALERT: vinfra-200;Current Users;OK;HARD;4;USERS OK - 2 users currently logged in
[1582098206] SERVICE ALERT: vapps-109;/mnt/data01 Partition;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098244] SERVICE ALERT: vapps-109;Current Load;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098266] SERVICE ALERT: vapps-109;/mnt/data01 Partition;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098281] SERVICE ALERT: vapps-109;Current Users;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098304] SERVICE ALERT: vapps-109;Current Load;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098319] SERVICE ALERT: vapps-109;PING;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098326] SERVICE ALERT: vapps-109;/mnt/data01 Partition;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098341] SERVICE ALERT: vapps-109;Current Users;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098356] SERVICE ALERT: vapps-109;Root Partition;CRITICAL;SOFT;1;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098364] SERVICE ALERT: vapps-109;Current Load;CRITICAL;SOFT;3;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098379] SERVICE ALERT: vapps-109;PING;CRITICAL;SOFT;2;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098386] SERVICE ALERT: vapps-109;/mnt/data01 Partition;CRITICAL;HARD;4;CHECK_NRPE: Error - Could not complete SSL handshake.
[1582098401] SERVICE ALERT: vapps-109;Current Users;OK;SOFT;3;USERS OK - 2 users currently logged in
[1582098417] SERVICE ALERT: vapps-109;Root Partition;OK;SOFT;2;/ is at 39% capacity, 16G of 43G
[1582098424] SERVICE ALERT: vapps-109;Current Load;OK;HARD;4;OK - load average: 0.02, 0.04, 0.04
[1582098444] SERVICE ALERT: vapps-109;PING;OK;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.04 ms
[1582098444] SERVICE ALERT: vapps-109;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 59% capacity, 511G of 917G