[1586676058] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 51.20, 29.82, 12.96
[1586676119] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 42.99, 31.54, 14.62
[1586676179] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 38.59, 32.61, 16.07
[1586676239] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 39.83, 34.11, 17.62
[1586676482] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.44, 8.44, 5.55
[1586676542] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 10.15, 8.68, 5.81
[1586676602] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.45, 8.47, 5.92
[1586676662] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.54, 7.95, 5.90
[1586677572] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 14.30, 12.40, 10.02
[1586677632] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.33, 12.22, 10.11
[1586677693] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.17, 12.45, 10.32
[1586677753] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.70, 12.12, 10.34
[1586682325] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 21.01, 10.80, 5.27
[1586682385] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 21.46, 12.79, 6.30
[1586682445] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 22.90, 14.76, 7.39
[1586682505] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 22.45, 16.19, 8.34
[1586682605] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586682725] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1586682853] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.73, 5.53, 8.60
[1586687339] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.14, 1.67, 15.22
[1586688062] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.91, 2.61, 4.27
[1586688239] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.16, 0.39, 6.01
[1586690005] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.05, 5.87, 7.86
[1586691805] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.86, 4.28, 4.94
[1586706245] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586706365] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 46% (5 min average)
[1586721320] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 87% (5 min average)
[1586721440] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 31% (5 min average)
[1586735339] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.70, 17.22, 7.59
[1586735399] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 36.84, 21.17, 9.55
[1586735459] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 47.93, 27.09, 12.31
[1586735519] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 44.68, 30.27, 14.35
[1586736062] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.55, 8.24, 5.56
[1586736122] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.11, 8.30, 5.75
[1586736182] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.70, 7.91, 5.78
[1586736242] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.52, 8.31, 6.04
[1586737919] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.70, 6.07, 19.73
[1586738042] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.70, 2.41, 4.36
[1586738819] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.06, 0.61, 7.72
[1586744550] SERVICE ALERT: vapps-109;Root Partition;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1586744551] HOST ALERT: vapps-109;DOWN;SOFT;1;(Host check timed out after 31.03 seconds)
[1586744551] SERVICE ALERT: vapps-109;Current Users;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1586744551] SERVICE ALERT: vapps-109;PING;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1586744551] SERVICE ALERT: vapps-109;Total Processes;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1586744552] SERVICE ALERT: vapps-109;/mnt/data01 Partition;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1586744552] SERVICE ALERT: vapps-109;Swap Usage;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1586744552] SERVICE ALERT: vapps-109;Current Load;CRITICAL;HARD;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1586744611] SERVICE ALERT: vapps-109;Root Partition;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.3.109 port 5666: No route to host
[1586744611] HOST ALERT: vapps-109;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.3.109)
[1586744640] SERVICE ALERT: vapps-109;SSH;CRITICAL;HARD;1;connect to address 192.168.3.109 and port 22: No route to host
[1586744674] HOST ALERT: vapps-109;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.3.109)
[1586744737] HOST ALERT: vapps-109;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.3.109)
[1586744800] HOST ALERT: vapps-109;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.3.109)
[1586744863] HOST ALERT: vapps-109;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.3.109)
[1586744926] HOST ALERT: vapps-109;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.3.109)
[1586744989] HOST ALERT: vapps-109;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.3.109)
[1586745052] HOST ALERT: vapps-109;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.3.109)
[1586745115] HOST ALERT: vapps-109;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.3.109)
[1586747735] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586747855] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586747975] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 28% (5 min average)
[1586752782] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586752902] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1586755581] SERVICE ALERT: WINV-105;Data Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586755582] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586755594] HOST ALERT: WINV-105;DOWN;SOFT;1;PING CRITICAL - Packet loss = 0%, RTA = 9504.89 ms
[1586755636] HOST ALERT: WINV-105;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1586755701] SERVICE ALERT: WINV-105;Data Drive Space;OK;SOFT;2;d:\ - total: 931.32 Gb - used: 606.72 Gb (65%) - free 324.60 Gb (35%)
[1586755702] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)