[1586502723] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 78% capacity, 2.7T of 3.6T
[1586503437] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 39.86, 22.58, 10.48
[1586503497] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 39.15, 25.75, 12.35
[1586503557] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 31.38, 26.16, 13.34
[1586503617] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 38.17, 29.39, 15.28
[1586503679] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.05, 7.29, 5.08
[1586503739] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.67, 7.89, 5.43
[1586503799] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.82, 7.22, 5.36
[1586503860] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.27, 7.86, 5.71
[1586504199] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 8.98, 8.23, 5.68
[1586504259] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 9.34, 8.46, 5.92
[1586504319] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 9.07, 8.52, 6.10
[1586504379] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 8.95, 8.59, 6.28
[1586504711] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 16.41, 14.38, 10.18
[1586504771] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 15.04, 14.39, 10.45
[1586504831] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 15.87, 14.72, 10.81
[1586504892] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 15.34, 14.80, 11.08
[1586505193] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586505313] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586505433] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1586506324] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586506444] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1586507233] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1586507353] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1586511492] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.53, 7.72, 9.65
[1586512023] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 2.8T of 3.6T
[1586512083] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 2.8T of 3.6T
[1586512143] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 2.8T of 3.6T
[1586512203] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 2.8T of 3.6T
[1586513161] SERVICE ALERT: vdb-103;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (1011 MB out of 10239 MB)
[1586513764] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1586513884] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 78% (5 min average)
[1586514360] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.19, 2.19, 4.22
[1586514717] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.56, 2.25, 16.42
[1586515392] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.34, 11.26, 10.27
[1586515452] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.32, 11.27, 10.34
[1586515512] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.37, 11.29, 10.40
[1586515572] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.63, 11.38, 10.49
[1586515617] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.55, 0.51, 6.50
[1586516703] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 2.7T of 3.6T
[1586527804] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586527924] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586528044] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 23% (5 min average)
[1586535304] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 2.8T of 3.6T
[1586535364] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 2.8T of 3.6T
[1586535424] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 2.8T of 3.6T
[1586535484] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 2.8T of 3.6T
[1586542779] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.78, 4.89, 4.98
[1586543979] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.30, 5.09, 5.03
[1586544039] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.52, 5.20, 5.07
[1586544099] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.43, 5.22, 5.09
[1586544159] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 5.17, 5.18, 5.08
[1586545659] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.88, 4.94, 4.99
[1586545959] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.26, 5.05, 5.02
[1586546019] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;2;OK - load average: 4.88, 4.98, 4.99
[1586546172] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.60, 8.21, 9.44
[1586546379] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.21, 5.04, 5.01
[1586546439] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.12, 5.04, 5.01
[1586546499] SERVICE ALERT: vonly-database;Current Load;OK;SOFT;3;OK - load average: 4.36, 4.85, 4.95
[1586555284] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 74% capacity, 2.5T of 3.6T
[1586563017] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 44.50, 33.80, 15.66
[1586563077] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 38.67, 34.34, 17.00
[1586563137] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 37.74, 34.69, 18.19
[1586563197] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 43.87, 36.84, 19.95
[1586563861] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.50, 7.08, 5.46
[1586563921] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.61, 6.78, 5.46
[1586563981] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.63, 7.26, 5.70
[1586564041] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.51, 6.89, 5.67
[1586564125] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (1044 MB out of 11215 MB)
[1586565597] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.63, 5.74, 17.92
[1586565841] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.04, 3.44, 4.92
[1586565927] SERVICE ALERT: vdb-103;/mnt/data01 Partition;CRITICAL;SOFT;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586565930] HOST ALERT: vdb-103;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.4.103)
[1586565964] SERVICE ALERT: vdb-103;Current Load;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586565990] SERVICE ALERT: vdb-103;/mnt/data01 Partition;CRITICAL;HARD;4;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586565993] HOST ALERT: vdb-103;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.4.103)
[1586566023] SERVICE ALERT: vdb-103;Current Users;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586566023] SERVICE ALERT: vdb-103;Total Processes;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586566023] SERVICE ALERT: vdb-103;PING;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586566023] SERVICE ALERT: vdb-103;SSH;CRITICAL;HARD;1;connect to address 192.168.4.103 and port 22: No route to host
[1586566056] HOST ALERT: vdb-103;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.4.103)
[1586566084] SERVICE ALERT: vdb-103;Root Partition;CRITICAL;HARD;1;(No output on stdout) stderr: connect to address 192.168.4.103 port 5666: No route to host
[1586566119] HOST ALERT: vdb-103;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.4.103)
[1586566182] HOST ALERT: vdb-103;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.4.103)
[1586566245] HOST ALERT: vdb-103;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.4.103)
[1586566308] HOST ALERT: vdb-103;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.4.103)
[1586566371] HOST ALERT: vdb-103;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.4.103)
[1586566434] HOST ALERT: vdb-103;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.4.103)
[1586566497] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.03, 0.58, 7.19
[1586566497] HOST ALERT: vdb-103;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.4.103)
[1586578877] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586579006] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;2;No data was received from host!
[1586579135] SERVICE ALERT: WINV-102;CPU Load;WARNING;HARD;3;No data was received from host!
[1586579735] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1586584445] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586584565] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586584685] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1586585285] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)