[1581584162] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.95, 5.94, 5.06
[1581584222] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.00, 5.98, 5.13
[1581584282] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.78, 5.91, 5.16
[1581584342] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.51, 5.80, 5.17
[1581592322] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.76, 4.01, 2.75
[1581592382] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.16, 5.68, 3.41
[1581592442] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 6.81, 5.49, 3.49
[1581593043] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.21, 4.04, 4.68
[1581595143] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.46, 5.48, 5.13
[1581595203] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.26, 5.66, 5.21
[1581595263] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.04, 5.68, 5.24
[1581595323] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 7.51, 6.16, 5.43
[1581596144] SERVICE ALERT: vdb-100;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 77% capacity, 2.7T of 3.7T
[1581608702] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 6.66, 6.53, 4.89
[1581608762] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 5.05, 6.10, 4.84
[1581608822] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.29, 5.68, 4.78
[1581609423] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.89, 4.93, 4.98
[1581609723] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.52, 5.26, 5.10
[1581609783] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 5.81, 5.39, 5.16
[1581609843] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.28, 5.32, 5.15
[1581609903] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 5.77, 5.47, 5.21
[1581610503] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.35, 4.63, 4.92
[1581634934] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 4.05, 4.44, 4.90
[1581635834] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 6.85, 5.92, 5.35
[1581635894] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 6.57, 6.04, 5.44
[1581635954] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 6.56, 6.12, 5.50
[1581636014] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 6.11, 6.07, 5.52
[1581637514] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 4.52, 4.85, 4.97
[1581640214] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 6.12, 5.25, 5.03
[1581640275] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 5.15, 5.18, 5.02
[1581640335] SERVICE ALERT: vdb-staging;Current Load;OK;SOFT;3;OK - load average: 4.82, 5.09, 5.00
[1581641595] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 5.22, 5.22, 5.05
[1581641655] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 5.57, 5.32, 5.09
[1581641715] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 6.53, 5.57, 5.19
[1581641775] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 7.15, 5.95, 5.34
[1581642375] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 4.51, 4.68, 4.98
[1581642675] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 5.20, 5.49, 5.28
[1581642735] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 5.18, 5.43, 5.27
[1581642795] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 5.18, 5.37, 5.26
[1581642855] SERVICE FLAPPING ALERT: vdb-staging;Current Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1581642855] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 4.94, 5.25, 5.22
[1581644655] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 4.09, 4.22, 4.76
[1581648855] SERVICE FLAPPING ALERT: vdb-staging;Current Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)