[1586157953] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.15, 9.19, 4.31
[1586158013] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 13.71, 9.93, 4.87
[1586158073] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 23.70, 13.37, 6.37
[1586158133] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 22.17, 14.91, 7.34
[1586158193] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 21.94, 16.05, 8.20
[1586158211] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.22, 27.42, 25.76
[1586158253] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 19.68, 16.49, 8.85
[1586158675] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.74, 7.31, 5.18
[1586158735] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.70, 7.61, 5.41
[1586158795] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.64, 7.40, 5.47
[1586158853] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.40, 27.80, 17.96
[1586158855] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 11.62, 8.50, 5.96
[1586161356] SERVICE ALERT: WINV-112;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.51 Gb - used: 748.34 Gb (80%) - free 183.17 Gb (20%)
[1586161476] SERVICE ALERT: WINV-112;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.51 Gb - used: 748.78 Gb (80%) - free 182.73 Gb (20%)
[1586161596] SERVICE ALERT: WINV-112;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.51 Gb - used: 749.25 Gb (80%) - free 182.26 Gb (20%)
[1586166994] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1586167114] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1586174796] SERVICE ALERT: WINV-112;Data Drive Space;CRITICAL;HARD;3;d:\ - total: 931.51 Gb - used: 839.30 Gb (90%) - free 92.21 Gb (10%)
[1586175055] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.95, 2.75, 4.41
[1586175353] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.09, 3.65, 19.95
[1586176253] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.35, 0.62, 7.89
[1586178396] SERVICE ALERT: WINV-112;Data Drive Space;OK;HARD;3;d:\ - total: 931.51 Gb - used: 614.08 Gb (66%) - free 317.43 Gb (34%)
[1586179354] SERVICE ALERT: WINV-108;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1586179474] SERVICE ALERT: WINV-108;CPU Load;WARNING;SOFT;2;CPU Load 82% (5 min average)
[1586179594] SERVICE ALERT: WINV-108;CPU Load;OK;HARD;3;CPU Load 79% (5 min average)
[1586182083] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;1;WARNING - load average: 6.95, 7.66, 5.62
[1586182143] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;2;WARNING - load average: 6.99, 7.54, 5.70
[1586182203] SERVICE ALERT: vdb-101;Current Load;WARNING;SOFT;3;WARNING - load average: 6.37, 7.28, 5.73
[1586182263] SERVICE ALERT: vdb-101;Current Load;WARNING;HARD;4;WARNING - load average: 6.42, 7.11, 5.77
[1586184063] SERVICE ALERT: vdb-101;Current Load;OK;HARD;4;OK - load average: 1.44, 3.04, 4.59
[1586186196] SERVICE ALERT: WINV-107;Data Drive Space;OK;HARD;3;d:\ - total: 931.51 Gb - used: 738.35 Gb (79%) - free 193.16 Gb (21%)
[1586191511] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 24.28, 24.48, 24.95
[1586193034] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1586193154] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;2;CPU Load 86% (5 min average)
[1586193274] SERVICE ALERT: WINV-107;CPU Load;OK;HARD;3;CPU Load 13% (5 min average)
[1586196996] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.51 Gb - used: 745.65 Gb (80%) - free 185.86 Gb (20%)
[1586197116] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.51 Gb - used: 745.65 Gb (80%) - free 185.86 Gb (20%)
[1586197236] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.51 Gb - used: 745.65 Gb (80%) - free 185.86 Gb (20%)
[1586200865] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 5.5T of 7.3T
[1586201701] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 19% free (1914 MB out of 10239 MB)
[1586201761] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 19% free (1914 MB out of 10239 MB)
[1586201821] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 19% free (1914 MB out of 10239 MB)
[1586201881] SERVICE ALERT: vdb-103;Swap Usage;WARNING;HARD;4;SWAP WARNING - 18% free (1833 MB out of 10239 MB)
[1586211301] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 2.8T of 3.6T
[1586211361] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 2.8T of 3.6T
[1586211421] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 2.8T of 3.6T
[1586211481] SERVICE ALERT: vonly-database;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 2.8T of 3.6T
[1586211953] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 26.39, 19.58, 10.66
[1586212013] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 11.04, 16.43, 10.14
[1586212073] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 5.28, 13.82, 9.63
[1586212133] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.04, 11.34, 9.04
[1586212433] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.14, 4.37, 6.65
[1586216933] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.31, 23.08, 10.50
[1586216993] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 29.09, 24.67, 11.89
[1586217053] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 22.36, 22.96, 12.07
[1586217113] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 20.84, 22.40, 12.56
[1586217414] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1586217414] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.73, 29.42, 18.39
[1586218211] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.76, 25.94, 25.05
[1586218314] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 15.16, 20.10, 19.96
[1586218555] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.77, 5.88, 5.11
[1586218615] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.26, 5.64, 5.07
[1586218676] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.49, 5.84, 5.19
[1586218736] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.84, 5.94, 5.27
[1586219214] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.68, 21.81, 20.21
[1586220714] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.29, 17.28, 19.82
[1586225111] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 24.65, 24.48, 24.97
[1586225214] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.09, 18.20, 13.48
[1586225514] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.12, 16.55, 14.12
[1586226414] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.88, 1.84, 6.18
[1586226536] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.81, 2.74, 4.19
[1586230081] SERVICE ALERT: vdb-103;Swap Usage;OK;HARD;4;SWAP OK - 21% free (2102 MB out of 10239 MB)
[1586231514] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1586231583] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;1;WARNING - load average: 5.74, 5.82, 5.06
[1586231643] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;2;WARNING - load average: 5.33, 5.72, 5.08
[1586231703] SERVICE ALERT: vdb-103;Current Load;WARNING;SOFT;3;WARNING - load average: 4.61, 5.47, 5.03
[1586231763] SERVICE ALERT: vdb-103;Current Load;WARNING;HARD;4;WARNING - load average: 4.98, 5.38, 5.02
[1586232063] SERVICE ALERT: vdb-103;Current Load;OK;HARD;4;OK - load average: 4.41, 4.82, 4.88
[1586242681] SERVICE ALERT: vonly-database;/media/data Partition;OK;HARD;4;/media/data is at 75% capacity, 2.6T of 3.6T