[1583310348] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 21.54, 10.22, 4.18
[1583310408] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 14.45, 10.35, 4.61
[1583310468] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 16.27, 11.54, 5.38
[1583310528] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.76, 11.68, 5.81
[1583311003] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1583311123] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1583311428] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.12, 25.30, 16.64
[1583311783] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.46, 6.03, 5.06
[1583311843] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 5.19, 5.59, 4.96
[1583312203] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.16, 5.83, 5.33
[1583312263] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.55, 5.94, 5.40
[1583312323] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.49, 6.07, 5.49
[1583312383] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.00, 5.80, 5.44
[1583312924] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 16.14, 12.60, 10.69
[1583312984] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 16.96, 13.45, 11.11
[1583313044] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 15.22, 13.61, 11.31
[1583313104] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 16.21, 14.13, 11.63
[1583317164] SERVICE ALERT: WINV-103;CPU Load;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1583317284] SERVICE ALERT: WINV-103;CPU Load;OK;SOFT;2;CPU Load 56% (5 min average)
[1583317843] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583317904] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 6.86, 7.94, 9.37
[1583317963] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1583319829] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.84, 15.83, 18.99
[1583319883] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.81, 4.53, 4.95
[1583320729] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 24.72, 23.42, 21.28
[1583321084] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.38, 5.13, 5.07
[1583321144] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.62, 5.56, 5.24
[1583321204] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 3.92, 5.10, 5.10
[1583321264] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.45, 5.44, 5.22
[1583324864] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.59, 2.50, 4.46
[1583325229] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.99, 3.37, 15.86
[1583326129] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.11, 0.28, 6.07
[1583333743] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.74 Gb (80%) - free 185.58 Gb (20%)
[1583333863] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.89 Gb (80%) - free 185.43 Gb (20%)
[1583333983] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 746.10 Gb (80%) - free 185.22 Gb (20%)
[1583358336] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583358465] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;2;No data was received from host!
[1583358585] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 28% (5 min average)
[1583364529] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 28.21, 16.41, 8.87
[1583364589] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 31.56, 19.40, 10.37
[1583364649] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 39.87, 24.01, 12.52
[1583364709] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 40.56, 27.80, 14.59
[1583365009] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.81, 15.21, 13.24
[1583365609] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.50, 2.88, 7.52
[1583369509] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.83, 10.65, 4.38
[1583369569] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 21.67, 11.86, 5.20
[1583369629] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 16.74, 12.40, 5.82
[1583369689] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 16.60, 13.31, 6.55
[1583370290] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (24.4% change >= 20.0% threshold)
[1583370290] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.84, 26.82, 16.06
[1583370764] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.84, 6.60, 5.13
[1583370824] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.74, 6.75, 5.27
[1583370884] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.30, 6.20, 5.17
[1583370944] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.98, 5.95, 5.15
[1583374544] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.97, 4.21, 4.91
[1583375090] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 17.69, 16.31, 19.50
[1583376290] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.27, 1.77, 6.41
[1583381090] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583389175] SERVICE ALERT: WINV-101;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1583389178] HOST ALERT: WINV-101;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.3.101)
[1583389188] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;1;connect to address 192.168.3.101 and port 12489: No route to host
[1583389239] HOST ALERT: WINV-101;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.3.101)
[1583389304] HOST ALERT: WINV-101;UP;SOFT;1;PING OK - Packet loss = 20%, RTA = 0.30 ms
[1583389775] SERVICE ALERT: WINV-101;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 751.04 Gb (81%) - free 180.27 Gb (19%)
[1583389785] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;1;CPU Load 0% (5 min average)
[1583390385] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583390505] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583390625] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 53% (5 min average)
[1583391548] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.65, 7.78, 3.93
[1583391608] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 10.91, 8.11, 4.29
[1583391668] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 11.16, 8.64, 4.71
[1583391728] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 8.80, 8.40, 4.87
[1583392628] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 2.18, 3.96, 4.84
[1583392866] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583392926] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1583393286] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583393346] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583393406] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1583393466] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1583394066] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583394126] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583394187] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1583394247] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1583394275] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 29.28, 17.94, 8.20
[1583394335] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 27.39, 19.67, 9.42
[1583394395] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 26.63, 20.84, 10.47
[1583394455] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 24.68, 21.38, 11.31
[1583394547] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583394607] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583394667] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1583394727] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1583395027] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT