[1583996554] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1583996674] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1583997605] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.03, 10.74, 5.44
[1583997666] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 19.51, 12.76, 6.47
[1583997726] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 21.48, 14.60, 7.50
[1583997786] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 23.82, 16.31, 8.52
[1583998386] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.52, 32.72, 20.60
[1583998505] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.85, 7.13, 5.11
[1583998565] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 10.29, 7.99, 5.54
[1583998625] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.36, 7.90, 5.66
[1583998685] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 10.11, 8.35, 5.95
[1583998893] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1584000737] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584000754] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1584000857] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584000887] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1584000983] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 9% (5 min average)
[1584001986] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.56, 13.94, 19.21
[1584002586] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (20.8% change >= 20.0% threshold)
[1584002586] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.36, 21.08, 20.22
[1584003383] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584003503] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 18% (5 min average)
[1584004954] SERVICE ALERT: WINV-112;CPU Load;WARNING;SOFT;1;CPU Load 86% (5 min average)
[1584005074] SERVICE ALERT: WINV-112;CPU Load;OK;SOFT;2;CPU Load 15% (5 min average)
[1584005820] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.11, 11.28, 10.34
[1584005880] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.52, 11.23, 10.38
[1584005940] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.27, 11.51, 10.53
[1584006000] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.73, 11.20, 10.49
[1584007800] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 7.15, 8.41, 10.00
[1584007986] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1584009486] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.80, 14.37, 18.71
[1584010686] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 29.07, 24.29, 20.41
[1584012186] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.07, 11.51, 18.17
[1584012186] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 6.01, 4.21, 4.91
[1584013386] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1584013386] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.10, 1.11, 6.41
[1584018186] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584019320] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.58 Gb (80%) - free 185.74 Gb (20%)
[1584019440] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.69 Gb (80%) - free 185.63 Gb (20%)
[1584019560] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.83 Gb (80%) - free 185.49 Gb (20%)
[1584025321] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.03, 6.93, 5.34
[1584025383] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.04, 6.46, 5.27
[1584025444] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.74, 6.62, 5.40
[1584025504] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.02, 6.45, 5.41
[1584027005] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.11, 4.27, 4.82
[1584055987] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 33.57, 24.19, 12.83
[1584056047] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 22.40, 23.07, 13.18
[1584056107] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 12.34, 20.13, 12.79
[1584056167] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.52, 17.00, 12.17
[1584056886] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.74, 5.85, 5.23
[1584056946] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.75, 5.52, 5.15
[1584057006] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.01, 5.76, 5.24
[1584057067] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.49, 4.97, 5.00
[1584057367] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.21, 5.70, 5.22
[1584057427] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.38, 5.78, 5.28
[1584057487] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.41, 5.29, 5.14
[1584057547] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.26, 5.89, 5.37
[1584057967] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 31.47, 23.65, 16.41
[1584058354] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;1;c:\ - total: 68.27 Gb - used: 54.66 Gb (80%) - free 13.61 Gb (20%)
[1584058474] SERVICE ALERT: WINV-103;OS Drive Space;OK;SOFT;2;c:\ - total: 68.27 Gb - used: 54.60 Gb (80%) - free 13.66 Gb (20%)
[1584058594] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;1;c:\ - total: 68.27 Gb - used: 54.64 Gb (80%) - free 13.62 Gb (20%)
[1584058714] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;SOFT;2;c:\ - total: 68.27 Gb - used: 54.68 Gb (80%) - free 13.58 Gb (20%)
[1584058834] SERVICE ALERT: WINV-103;OS Drive Space;WARNING;HARD;3;c:\ - total: 68.27 Gb - used: 54.68 Gb (80%) - free 13.59 Gb (20%)
[1584060120] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1584060240] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1584062647] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.59, 3.83, 4.97
[1584062767] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.74, 11.03, 17.95
[1584063968] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.04, 0.80, 6.56
[1584074760] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 80% (5 min average)
[1584074885] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 97% (5 min average)
[1584075005] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 68% (5 min average)
[1584075027] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584075152] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584075272] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 8% (5 min average)
[1584077072] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584077192] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584077312] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1584077912] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 65% (5 min average)