[1584602116] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 33.93, 19.03, 8.25
[1584602176] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 45.60, 25.08, 11.00
[1584602236] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 43.18, 28.14, 12.93
[1584602297] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 46.67, 31.91, 15.19
[1584603095] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.53, 7.07, 5.59
[1584603155] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.68, 7.73, 5.92
[1584603215] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.82, 8.05, 6.13
[1584603275] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 13.14, 9.68, 6.83
[1584604381] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 14.74, 14.76, 10.88
[1584604441] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 5.5T of 7.3T
[1584604442] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.72, 14.54, 11.06
[1584604502] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 14.57, 14.49, 11.25
[1584604562] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 12.81, 14.20, 11.37
[1584606062] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.33, 8.06, 9.93
[1584608641] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 5.5T of 7.3T
[1584608701] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 5.5T of 7.3T
[1584608761] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 5.5T of 7.3T
[1584608821] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 5.6T of 7.3T
[1584609394] SERVICE ALERT: WINV-112;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1584609514] SERVICE ALERT: WINV-112;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1584610262] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.92, 12.40, 10.66
[1584610322] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.46, 12.36, 10.76
[1584610382] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.91, 12.53, 10.92
[1584610442] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 14.20, 12.91, 11.15
[1584612842] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 2.80, 5.24, 8.58
[1584616175] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.37, 2.06, 4.13
[1584616397] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.73, 2.78, 14.74
[1584616997] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.23, 0.87, 7.98
[1584626440] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584626560] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1584626680] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 54% (5 min average)
[1584628386] SERVICE ALERT: vonly-database;SSH;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1584628454] SERVICE ALERT: vonly-database;SSH;OK;SOFT;2;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1584648422] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 5.5T of 7.3T
[1584651282] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584651402] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 4% (5 min average)
[1584661097] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 36.23, 20.90, 11.20
[1584661157] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 36.79, 24.15, 12.95
[1584661217] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 38.03, 27.24, 14.73
[1584661277] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.97, 29.17, 16.18
[1584661475] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.74, 7.36, 5.71
[1584661535] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.27, 6.76, 5.60
[1584661595] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.48, 6.52, 5.59
[1584661655] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 9.28, 7.13, 5.84
[1584662822] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 5.5T of 7.3T
[1584662882] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 5.5T of 7.3T
[1584662942] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 5.5T of 7.3T
[1584663002] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 5.6T of 7.3T
[1584666254] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 93% (5 min average)
[1584666374] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1584666455] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.74, 3.29, 4.58
[1584666977] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.83, 2.46, 18.48
[1584667877] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.17, 0.49, 7.22
[1584683297] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584683424] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584683544] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584683922] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584684042] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584684144] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 9% (5 min average)
[1584684162] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1584687154] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1584687274] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 47% (5 min average)