[1585465765] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585465887] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585466007] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585466432] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 29.35, 11.60, 4.77
[1585466492] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 26.30, 14.10, 6.07
[1585466552] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 19.87, 14.30, 6.64
[1585466607] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1585466612] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 24.78, 16.55, 7.88
[1585467007] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 14.54, 13.53, 10.70
[1585467067] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 14.00, 13.58, 10.90
[1585467128] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 14.16, 13.63, 11.08
[1585467188] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 16.21, 14.33, 11.48
[1585467212] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.58, 26.88, 17.06
[1585467342] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.05, 6.91, 5.32
[1585467402] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.92, 6.61, 5.31
[1585467462] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.19, 7.53, 5.72
[1585467522] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.96, 7.21, 5.72
[1585469333] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1585469453] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 78% (5 min average)
[1585470364] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1585472315] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 7.06, 6.36, 5.27
[1585472375] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 7.08, 6.47, 5.38
[1585472435] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.45, 6.43, 5.44
[1585472495] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 7.04, 6.59, 5.55
[1585472607] SERVICE ALERT: WINV-101;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1585472727] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 45% (5 min average)
[1585474954] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 87% (5 min average)
[1585475074] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1585478012] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.48, 12.75, 18.92
[1585478022] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.24, 3.68, 4.98
[1585478312] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 34.71, 32.50, 26.22
[1585478322] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.36, 4.73, 5.07
[1585478382] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 3.64, 4.34, 4.91
[1585478495] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.53, 4.12, 4.91
[1585479812] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.99, 11.64, 18.30
[1585479942] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.80, 5.64, 5.09
[1585480003] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.90, 5.29, 5.00
[1585480063] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.07, 5.65, 5.15
[1585480112] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1585480112] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.45, 27.80, 24.04
[1585480123] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.20, 5.52, 5.14
[1585480183] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.47, 5.26, 5.07
[1585480243] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.73, 5.32, 5.11
[1585480543] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.63, 3.79, 4.56
[1585483412] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.49, 8.53, 18.07
[1585484314] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585484434] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 56% (5 min average)
[1585484847] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585484912] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 9.43, 7.35, 7.99
[1585484967] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1585485087] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 16% (5 min average)
[1585485212] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 8.27, 7.89, 8.13
[1585485272] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 6.14, 7.25, 7.89
[1585485632] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 8.79, 8.88, 8.37
[1585485643] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.56, 5.65, 5.07
[1585485692] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 10.40, 9.01, 8.42
[1585485703] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.68, 5.97, 5.20
[1585485752] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 5.21, 7.80, 8.04
[1585485763] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.12, 5.54, 5.10
[1585485812] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.95, 8.14, 8.17
[1585485823] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.32, 5.73, 5.19
[1585486112] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.16, 6.58, 7.61
[1585486423] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.24, 3.46, 4.55
[1585490525] SERVICE ALERT: WINV-107;Uptime;CRITICAL;SOFT;1;connect to address 192.168.3.107 and port 12489: No route to host
[1585490528] HOST ALERT: WINV-107;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.3.107)
[1585490557] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;HARD;1;connect to address 192.168.3.107 and port 12489: No route to host
[1585490557] SERVICE ALERT: WINV-107;OS Drive Space;CRITICAL;HARD;1;connect to address 192.168.3.107 and port 12489: No route to host
[1585490557] SERVICE ALERT: WINV-107;Memory Usage;CRITICAL;HARD;1;connect to address 192.168.3.107 and port 12489: No route to host
[1585490591] HOST ALERT: WINV-107;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.3.107)
[1585490646] SERVICE ALERT: WINV-107;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1585490650] HOST ALERT: WINV-107;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.33 ms
[1585490976] HOST ALERT: WINV-107;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1585491037] HOST ALERT: WINV-107;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.3.107)
[1585491101] HOST FLAPPING ALERT: WINV-107;STARTED; Host appears to have started flapping (22.6% change > 20.0% threshold)
[1585491101] HOST ALERT: WINV-107;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.37 ms
[1585491154] SERVICE ALERT: WINV-107;OS Drive Space;OK;SOFT;1;c:\ - total: 67.45 Gb - used: 35.02 Gb (52%) - free 32.43 Gb (48%)
[1585491154] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;1;CPU Load 0% (5 min average)
[1585491154] SERVICE ALERT: WINV-107;Memory Usage;OK;SOFT;1;Memory usage: total:79143.40 MB - used: 36115.14 MB (46%) - free: 43028.26 MB (54%)
[1585491813] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585492566] HOST FLAPPING ALERT: WINV-107;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1585495954] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585496074] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 27% (5 min average)
[1585508795] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 5.48, 6.02, 5.09
[1585508855] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 5.37, 5.88, 5.10
[1585508915] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 5.41, 5.79, 5.11
[1585508975] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.06, 5.92, 5.20
[1585510088] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 11% free (6020 MB out of 57111 MB)
[1585517764] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 90% capacity! 6.2T of 7.3T
[1585519475] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.55, 4.12, 4.68
[1585521223] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.51, 7.12, 5.09
[1585521283] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.49, 7.04, 5.19
[1585521343] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.97, 6.76, 5.21
[1585521403] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.00, 6.74, 5.31
[1585521512] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 11.92, 11.19, 7.65
[1585521573] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 11.31, 11.26, 7.90
[1585521633] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 11.29, 11.22, 8.10
[1585521693] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.11, 11.06, 8.23
[1585522594] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 98% (5 min average)
[1585522714] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1585522903] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.74, 2.90, 4.62
[1585523193] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.41, 1.74, 6.24
[1585525293] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 28.98, 14.23, 6.46
[1585525353] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.46, 16.70, 7.81
[1585525414] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.57, 19.20, 9.22
[1585525474] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 35.92, 22.49, 10.95
[1585526203] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.92, 7.10, 5.45
[1585526263] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.62, 7.20, 5.59
[1585526323] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.59, 6.71, 5.52
[1585526383] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.77, 6.26, 5.44
[1585527965] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1585531973] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585532074] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.23, 6.43, 15.17
[1585532084] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.67, 2.42, 4.36
[1585532093] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1585532975] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.40, 1.76, 6.64
[1585539788] SERVICE ALERT: vonly-database;Swap Usage;OK;HARD;4;SWAP OK - 21% free (11612 MB out of 57111 MB)
[1585540565] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 90% capacity! 6.2T of 7.3T