[1585207407] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 21.10, 16.32, 7.58
[1585207467] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 25.33, 18.39, 8.84
[1585207527] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 21.61, 18.66, 9.53
[1585207587] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 21.50, 18.99, 10.21
[1585207653] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.52, 4.56, 4.99
[1585207748] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (11335 MB out of 57111 MB)
[1585207808] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (11303 MB out of 57111 MB)
[1585207868] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 20% free (11332 MB out of 57111 MB)
[1585207928] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 20% free (11348 MB out of 57111 MB)
[1585208179] SERVICE ALERT: WINV-101;OS Drive Space;WARNING;SOFT;1;No data was received from host!
[1585208180] SERVICE ALERT: WINV-101;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1585208187] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.83, 26.28, 18.24
[1585208299] SERVICE ALERT: WINV-101;OS Drive Space;OK;SOFT;2;c:\ - total: 68.27 Gb - used: 43.49 Gb (64%) - free 24.77 Gb (36%)
[1585208300] SERVICE ALERT: WINV-101;Memory Usage;OK;SOFT;2;Memory usage: total:98284.48 MB - used: 45769.76 MB (47%) - free: 52514.72 MB (53%)
[1585208380] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 10.06, 7.67, 5.63
[1585208440] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.16, 7.29, 5.63
[1585208500] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.91, 7.13, 5.68
[1585208527] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.79, 12.70, 10.59
[1585208560] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.12, 7.32, 5.85
[1585208573] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585208587] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.37, 12.84, 10.77
[1585208647] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.59, 12.42, 10.75
[1585208693] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585208707] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.21, 12.17, 10.77
[1585208813] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1585209128] SERVICE ALERT: vonly-database;Swap Usage;OK;HARD;4;SWAP OK - 21% free (11428 MB out of 57111 MB)
[1585209274] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 80% (5 min average)
[1585209394] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 32% (5 min average)
[1585209514] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1585209634] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 35% (5 min average)
[1585209728] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (11360 MB out of 57111 MB)
[1585209788] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (11378 MB out of 57111 MB)
[1585209848] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 20% free (11369 MB out of 57111 MB)
[1585209908] SERVICE ALERT: vonly-database;Swap Usage;OK;HARD;4;SWAP OK - 21% free (11436 MB out of 57111 MB)
[1585212413] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1585212533] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1585215308] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (11332 MB out of 57111 MB)
[1585215368] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (11328 MB out of 57111 MB)
[1585215428] SERVICE ALERT: vonly-database;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 20% free (11309 MB out of 57111 MB)
[1585215488] SERVICE ALERT: vonly-database;Swap Usage;WARNING;HARD;4;SWAP WARNING - 20% free (11331 MB out of 57111 MB)
[1585216587] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.69, 13.92, 19.49
[1585218387] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 29.76, 23.56, 20.54
[1585218907] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.32, 8.46, 9.79
[1585219587] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 18.10, 17.80, 19.37
[1585220407] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.80, 10.70, 10.12
[1585220467] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.28, 10.58, 10.12
[1585220527] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.17, 10.79, 10.22
[1585220587] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.98, 10.83, 10.27
[1585220787] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1585220787] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 35.50, 28.39, 22.87
[1585225587] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1585227160] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.22, 3.64, 4.82
[1585227487] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.14, 8.85, 9.99
[1585227687] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.36, 3.99, 15.13
[1585228588] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.08, 0.70, 6.21
[1585231960] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 12.67, 9.86, 6.35
[1585232020] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 10.55, 9.68, 6.50
[1585232080] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.42, 8.75, 6.38
[1585232140] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 3.38, 7.44, 6.08
[1585233689] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 11.69, 13.69, 10.45
[1585233749] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 9.87, 12.82, 10.36
[1585233809] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 11.81, 12.79, 10.50
[1585233869] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.47, 11.91, 10.34
[1585234687] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.22, 10.89, 10.03
[1585234747] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.20, 10.95, 10.11
[1585234807] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.47, 11.07, 10.20
[1585234867] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.44, 11.13, 10.28
[1585234954] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585235074] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1585237169] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.63, 6.72, 7.70
[1585237769] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 9.03, 8.73, 8.16
[1585237830] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 8.88, 8.77, 8.21
[1585237890] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 9.04, 8.75, 8.23
[1585237950] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 8.77, 8.78, 8.28
[1585239450] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1585239450] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 5.30, 6.47, 7.58
[1585240240] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.24, 3.45, 4.49
[1585240853] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585240973] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1585241167] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.34, 9.36, 9.88
[1585243243] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1585243363] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;2;CPU Load 89% (5 min average)
[1585243483] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585243950] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.3% change < 5.0% threshold)
[1585244083] SERVICE ALERT: WINV-102;CPU Load;WARNING;HARD;3;CPU Load 88% (5 min average)
[1585245283] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 25% (5 min average)
[1585255953] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.53, 6.28, 5.29
[1585256013] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 6.33, 6.26, 5.34
[1585256073] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 6.79, 6.40, 5.45
[1585256133] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 6.55, 6.42, 5.52
[1585257934] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.19, 3.82, 4.74
[1585264802] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 5.5T of 7.3T
[1585265467] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.79, 11.00, 10.23
[1585265527] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.56, 11.21, 10.35
[1585265550] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 12.81, 12.15, 7.67
[1585265587] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.91, 11.36, 10.46
[1585265610] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 14.89, 12.58, 8.09
[1585265647] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.37, 11.33, 10.51
[1585265670] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 10.35, 11.77, 8.09
[1585265730] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.05, 11.72, 8.31
[1585265740] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.49, 7.32, 5.49
[1585265800] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 10.42, 8.29, 5.93
[1585265861] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.30, 7.52, 5.82
[1585265921] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 7.77, 7.62, 5.95
[1585267088] SERVICE ALERT: vonly-database;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (5710 MB out of 57111 MB)
[1585267202] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 5.5T of 7.3T
[1585267230] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 43.36, 30.97, 19.45
[1585267263] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 5.5T of 7.3T
[1585267323] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 5.5T of 7.3T
[1585267383] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 5.5T of 7.3T
[1585272283] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585272330] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 4.51, 9.75, 18.32
[1585272403] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 16% (5 min average)
[1585273121] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.43, 3.50, 4.56
[1585273530] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.17, 2.61, 7.31
[1585275423] SERVICE ALERT: vinfra-200;Root Partition;WARNING;SOFT;1;/ is at 81% capacity! 140G of 182G
[1585275483] SERVICE ALERT: vinfra-200;Root Partition;WARNING;SOFT;2;/ is at 82% capacity! 141G of 182G
[1585275487] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 18% free (1979 MB out of 11215 MB)
[1585275543] SERVICE ALERT: vinfra-200;Root Partition;WARNING;SOFT;3;/ is at 83% capacity! 143G of 182G
[1585275547] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 17% free (1897 MB out of 11215 MB)
[1585275603] SERVICE ALERT: vinfra-200;Root Partition;WARNING;HARD;4;/ is at 83% capacity! 144G of 182G
[1585275607] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 17% free (1860 MB out of 11215 MB)
[1585275667] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;HARD;4;SWAP WARNING - 17% free (1837 MB out of 11215 MB)
[1585275847] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.34, 9.16, 9.97
[1585276203] SERVICE ALERT: vinfra-200;Root Partition;OK;HARD;4;/ is at 67% capacity, 115G of 182G
[1585281729] SERVICE ALERT: WINV-104;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1585281729] SERVICE ALERT: WINV-104;OS Drive Space;WARNING;SOFT;1;No data was received from host!
[1585281730] SERVICE ALERT: WINV-104;Uptime;WARNING;SOFT;1;No data was received from host!
[1585281849] SERVICE ALERT: WINV-104;OS Drive Space;OK;SOFT;2;c:\ - total: 68.27 Gb - used: 47.81 Gb (70%) - free 20.45 Gb (30%)
[1585281849] SERVICE ALERT: WINV-104;Memory Usage;OK;SOFT;2;Memory usage: total:129142.02 MB - used: 45167.97 MB (35%) - free: 83974.04 MB (65%)
[1585281850] SERVICE ALERT: WINV-104;Uptime;OK;SOFT;2;System Uptime - 182 day(s) 10 hour(s) 30 minute(s)
[1585281978] SERVICE ALERT: WINV-104;Memory Usage;WARNING;SOFT;1;No data was received from host!
[1585281978] SERVICE ALERT: WINV-104;OS Drive Space;WARNING;SOFT;1;No data was received from host!
[1585281979] SERVICE ALERT: WINV-104;Uptime;WARNING;SOFT;1;No data was received from host!
[1585282098] SERVICE ALERT: WINV-104;OS Drive Space;OK;SOFT;2;c:\ - total: 68.27 Gb - used: 47.81 Gb (70%) - free 20.45 Gb (30%)
[1585282098] SERVICE ALERT: WINV-104;Memory Usage;OK;SOFT;2;Memory usage: total:129142.02 MB - used: 44445.94 MB (34%) - free: 84696.07 MB (66%)
[1585282099] SERVICE ALERT: WINV-104;Uptime;OK;SOFT;2;System Uptime - 182 day(s) 10 hour(s) 35 minute(s)
[1585286283] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 5.5T of 7.3T