[1584515956] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 61.52, 36.23, 15.83
[1584516016] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 51.64, 38.01, 17.72
[1584516076] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 50.98, 40.72, 19.95
[1584516136] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 48.73, 41.66, 21.57
[1584516430] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.66, 7.54, 5.10
[1584516490] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.92, 7.23, 5.15
[1584516550] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.63, 8.22, 5.64
[1584516610] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.09, 8.08, 5.75
[1584517522] SERVICE ALERT: vinfra-200;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.13, 26.86, 16.98
[1584517623] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 19% free (2105 MB out of 11215 MB)
[1584517649] SERVICE ALERT: vinfra-200;MySQL Connection Time;CRITICAL;SOFT;1;CRITICAL - 11.01 seconds to connect as nagios
[1584517683] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (2187 MB out of 11215 MB)
[1584517709] SERVICE ALERT: vinfra-200;MySQL Connection Time;OK;SOFT;2;OK - 0.01 seconds to connect as nagios
[1584517743] SERVICE ALERT: vinfra-200;Swap Usage;OK;SOFT;3;SWAP OK - 21% free (2262 MB out of 11215 MB)
[1584517952] SERVICE ALERT: vinfra-200;VAULT;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1584518026] SERVICE ALERT: vinfra-200;VAULT;OK;SOFT;2;OK - Vault is initialized, unsealed, and active
[1584518128] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 14.10, 19.78, 19.47
[1584518236] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.37, 12.69, 19.13
[1584518253] SERVICE ALERT: vinfra-200;MySQL Uptime;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1584518328] SERVICE ALERT: vinfra-200;MySQL Uptime;OK;SOFT;2;OK - database is up since 40394 minutes
[1584518403] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 8% free (815 MB out of 11215 MB)
[1584518434] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1584518434] SERVICE ALERT: vinfra-200;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.81, 26.20, 22.42
[1584518463] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 10% free (1031 MB out of 11215 MB)
[1584518523] SERVICE ALERT: vinfra-200;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 11% free (1135 MB out of 11215 MB)
[1584518585] SERVICE ALERT: vinfra-200;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 8% free (826 MB out of 11215 MB)
[1584518886] SERVICE ALERT: vinfra-200;Swap Usage;OK;HARD;4;SWAP OK - 72% free (7963 MB out of 11215 MB)
[1584518988] SERVICE ALERT: vinfra-200;MySQL Uptime;CRITICAL;SOFT;1;CRITICAL - database is up since 2 minutes
[1584519048] SERVICE ALERT: vinfra-200;MySQL Uptime;CRITICAL;SOFT;2;CRITICAL - database is up since 3 minutes
[1584519108] SERVICE ALERT: vinfra-200;MySQL Uptime;CRITICAL;SOFT;3;CRITICAL - database is up since 4 minutes
[1584519168] SERVICE ALERT: vinfra-200;MySQL Uptime;WARNING;HARD;4;WARNING - database is up since 5 minutes
[1584519469] SERVICE ALERT: vinfra-200;MySQL Uptime;OK;HARD;4;OK - database is up since 10 minutes
[1584519635] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 10.62, 12.26, 18.06
[1584520475] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584520595] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 94% (5 min average)
[1584520715] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 2% (5 min average)
[1584520936] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 37.71, 28.58, 20.36
[1584524435] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1584525874] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584525994] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1584527281] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 5.5T of 7.3T
[1584527341] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 5.5T of 7.3T
[1584527401] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 5.5T of 7.3T
[1584527461] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 5.5T of 7.3T
[1584531335] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.74, 4.37, 4.82
[1584532336] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.37, 9.01, 19.92
[1584533236] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.21, 0.96, 7.80
[1584538115] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1584538235] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1584538355] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1584538955] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 11% (5 min average)
[1584540395] SERVICE ALERT: WINV-103;Memory Usage;OK;HARD;3;Memory usage: total:234403.37 MB - used: 61623.51 MB (26%) - free: 172779.86 MB (74%)
[1584554761] SERVICE ALERT: vonly-replica;/media/data Partition;OK;HARD;4;/media/data is at 80% capacity, 5.5T of 7.3T
[1584560915] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1584561035] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 60% (5 min average)
[1584573760] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584573880] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 8% (5 min average)
[1584574235] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.17, 7.69, 5.27
[1584574295] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.35, 7.60, 5.39
[1584574336] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 38.75, 27.13, 14.44
[1584574355] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.14, 7.17, 5.38
[1584574396] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 43.54, 30.91, 16.55
[1584574415] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.50, 7.31, 5.56
[1584574456] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 21.12, 26.95, 16.10
[1584574516] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.32, 22.50, 15.24
[1584574714] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1584574715] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.90, 5.02, 5.00
[1584574834] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 17% (5 min average)
[1584575015] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.56, 5.85, 5.39
[1584575075] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.28, 5.68, 5.36
[1584575135] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.05, 5.70, 5.38
[1584575195] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.21, 5.26, 5.25
[1584576016] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 80.26, 46.63, 25.88
[1584578461] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;1;/media/data is at 81% capacity! 5.5T of 7.3T
[1584578521] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;2;/media/data is at 81% capacity! 5.5T of 7.3T
[1584578581] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;SOFT;3;/media/data is at 81% capacity! 5.5T of 7.3T
[1584578641] SERVICE ALERT: vonly-replica;/media/data Partition;WARNING;HARD;4;/media/data is at 81% capacity! 5.5T of 7.3T
[1584579374] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 92% (5 min average)
[1584579494] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1584580214] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1584580295] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.99, 2.83, 4.23
[1584580334] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 13% (5 min average)
[1584580816] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.06, 2.67, 18.56
[1584581716] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 1.09, 0.60, 7.41
[1584597400] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1584597520] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 9% (5 min average)