[1586767019] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 38.85, 15.18, 8.07
[1586767079] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 15.31, 12.77, 7.69
[1586767139] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 10.06, 11.56, 7.58
[1586767199] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.98, 9.84, 7.24
[1586767499] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 16.33, 12.22, 8.63
[1586767559] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 6.54, 10.14, 8.13
[1586767619] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;3;OK - load average: 2.76, 8.41, 7.66
[1586768881] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 25.87, 13.14, 8.93
[1586768941] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 10.11, 10.95, 8.44
[1586769001] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;3;OK - load average: 4.22, 9.09, 7.95
[1586769361] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 5.14, 11.90, 9.92
[1586769421] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 2.41, 9.87, 9.34
[1586769481] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 6.61, 9.35, 9.18
[1586769541] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 6.31, 8.89, 9.04
[1586770441] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.37, 2.84, 6.39
[1586770741] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 25.81, 10.03, 7.91
[1586770801] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 34.27, 15.64, 9.99
[1586770862] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 38.41, 20.16, 11.89
[1586770919] HOST ALERT: vapps-109;UP;HARD;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1586770920] SERVICE ALERT: vapps-109;Total Processes;OK;SOFT;1;PROCS OK: 187 processes with STATE = RSZDT
[1586770921] SERVICE ALERT: vapps-109;Current Users;OK;SOFT;1;USERS OK - 0 users currently logged in
[1586770922] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.74, 22.40, 13.14
[1586770922] SERVICE ALERT: vapps-109;Swap Usage;OK;SOFT;1;SWAP OK - 100% free (24567 MB out of 24567 MB)
[1586770922] SERVICE ALERT: vapps-109;Current Load;OK;SOFT;1;OK - load average: 4.24, 1.97, 0.75
[1586770927] SERVICE ALERT: vapps-109;/mnt/data01 Partition;OK;SOFT;1;/mnt/data01 is at 1% capacity, 2.6G of 917G
[1586770931] SERVICE ALERT: vapps-109;PING;OK;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.06 ms
[1586771010] SERVICE ALERT: vapps-109;Root Partition;OK;HARD;4;/ is at 50% capacity, 20G of 43G
[1586771040] SERVICE ALERT: vapps-109;SSH;OK;SOFT;1;SSH OK - OpenSSH_7.2p2 Ubuntu-4ubuntu2.8 (protocol 2.0)
[1586771042] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.49, 8.42, 5.78
[1586771102] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 8.22, 8.40, 5.94
[1586771162] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 11.43, 9.13, 6.34
[1586771222] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 11.93, 9.92, 6.80
[1586779085] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1586779205] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586779325] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1586780554] SERVICE ALERT: WINV-111;Memory Usage;WARNING;SOFT;1;Memory usage: total:70469.53 MB - used: 56413.65 MB (80%) - free: 14055.87 MB (20%)
[1586780674] SERVICE ALERT: WINV-111;Memory Usage;WARNING;SOFT;2;Memory usage: total:70469.53 MB - used: 56448.95 MB (80%) - free: 14020.58 MB (20%)
[1586780795] SERVICE ALERT: WINV-111;Memory Usage;WARNING;HARD;3;Memory usage: total:70469.53 MB - used: 56465.96 MB (80%) - free: 14003.57 MB (20%)
[1586781395] SERVICE ALERT: WINV-111;Memory Usage;OK;HARD;3;Memory usage: total:70469.53 MB - used: 55480.78 MB (79%) - free: 14988.75 MB (21%)
[1586781722] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.61, 2.74, 4.45
[1586781722] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.64, 5.58, 17.17
[1586782622] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.38, 0.99, 7.11
[1586783882] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 18.12, 13.38, 7.60
[1586783942] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 18.53, 14.35, 8.29
[1586784002] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 17.90, 14.93, 8.87
[1586784062] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 17.81, 15.47, 9.44
[1586784995] SERVICE ALERT: WINV-111;Memory Usage;WARNING;SOFT;1;Memory usage: total:70469.53 MB - used: 56538.45 MB (80%) - free: 13931.08 MB (20%)
[1586785115] SERVICE ALERT: WINV-111;Memory Usage;WARNING;SOFT;2;Memory usage: total:70469.53 MB - used: 56717.42 MB (80%) - free: 13752.11 MB (20%)
[1586785235] SERVICE ALERT: WINV-111;Memory Usage;WARNING;HARD;3;Memory usage: total:70469.53 MB - used: 56998.47 MB (81%) - free: 13471.06 MB (19%)
[1586786204] SERVICE ALERT: WINV-108;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1586786207] HOST ALERT: WINV-108;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.3.108)
[1586786297] HOST ALERT: WINV-108;DOWN;SOFT;2;(Host check timed out after 30.03 seconds)
[1586786324] SERVICE ALERT: WINV-108;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1586786328] HOST ALERT: WINV-108;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.39 ms
[1586787035] SERVICE FLAPPING ALERT: WINV-111;Memory Usage;STARTED; Service appears to have started flapping (22.1% change >= 20.0% threshold)
[1586787035] SERVICE ALERT: WINV-111;Memory Usage;OK;HARD;3;Memory usage: total:70469.53 MB - used: 37179.14 MB (53%) - free: 33290.38 MB (47%)
[1586793393] HOST ALERT: WINV-110;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1586793456] HOST ALERT: WINV-110;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.3.110)
[1586793520] HOST ALERT: WINV-110;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.35 ms
[1586797235] SERVICE FLAPPING ALERT: WINV-111;Memory Usage;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1586811362] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 4.42, 4.51, 4.97
[1586815263] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 21.91, 13.21, 7.93
[1586815323] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 22.47, 14.93, 8.85
[1586815383] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 22.23, 16.25, 9.69
[1586815443] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 22.46, 17.38, 10.49
[1586816943] SERVICE ALERT: vonly-database;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 22.66, 22.42, 20.05
[1586818922] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 5.32, 10.63, 8.45
[1586818982] SERVICE ALERT: vapps-100;Current Load;OK;SOFT;2;OK - load average: 2.66, 8.92, 8.00
[1586821442] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 35.20, 21.51, 10.27
[1586821502] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 39.43, 25.26, 12.27
[1586821562] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 32.54, 26.11, 13.39
[1586821622] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.07, 26.42, 14.29
[1586822222] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.40, 7.36, 5.43
[1586822283] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.43, 7.20, 5.50
[1586822343] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.09, 7.42, 5.69
[1586822403] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.81, 7.27, 5.75
[1586823903] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.59, 3.78, 4.92
[1586824022] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.69, 5.50, 17.81
[1586824922] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.30, 1.16, 7.37
[1586833143] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 17.99, 18.77, 19.89
[1586833422] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1586833542] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 30% (5 min average)
[1586833662] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1586833782] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1586833902] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1586834502] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 44% (5 min average)
[1586838154] SERVICE ALERT: WINV-110;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1586838274] SERVICE ALERT: WINV-110;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)
[1586842149] SERVICE ALERT: WINV-103;CPU Load;WARNING;SOFT;1;CPU Load 83% (5 min average)
[1586842269] SERVICE ALERT: WINV-103;CPU Load;OK;SOFT;2;CPU Load 65% (5 min average)