[1583911029] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.72, 3.12, 1.50
[1583911076] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 21.48, 11.72, 5.18
[1583911089] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.91, 5.33, 2.36
[1583911136] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 14.84, 11.58, 5.54
[1583911149] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 10.09, 6.03, 2.80
[1583911196] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 14.78, 12.01, 6.06
[1583911209] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 5.86, 5.60, 2.86
[1583911256] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 15.72, 12.61, 6.63
[1583912156] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 24.66, 26.71, 19.14
[1583912269] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 12.04, 7.46, 5.54
[1583912330] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.82, 6.97, 5.48
[1583912390] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.17, 6.68, 5.48
[1583912450] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.75, 6.78, 5.59
[1583915083] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1583915203] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 29% (5 min average)
[1583917309] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.85, 10.40, 10.13
[1583917369] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.85, 10.47, 10.17
[1583917429] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.25, 10.35, 10.15
[1583917489] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.07, 10.32, 10.15
[1583918156] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.21, 15.69, 19.92
[1583919523] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1583919643] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 71% (5 min average)
[1583919656] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 29.09, 24.98, 21.10
[1583919763] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 82% (5 min average)
[1583919883] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 41% (5 min average)
[1583920003] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 81% (5 min average)
[1583920123] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 63% (5 min average)
[1583921770] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583921890] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1583922589] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.95, 9.18, 9.98
[1583922656] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.45, 15.39, 19.20
[1583927450] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.52, 3.00, 4.27
[1583927756] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.39, 1.91, 6.50
[1583928909] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 6.90, 6.43, 3.30
[1583928969] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 3.44, 5.55, 3.19
[1583929210] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583929330] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583929491] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1583930241] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.69, 5.54, 3.65
[1583930301] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 13.98, 7.26, 4.36
[1583930361] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 9.60, 7.30, 4.56
[1583930421] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.02, 6.40, 4.43
[1583930691] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 19% (5 min average)
[1583930721] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.94, 2.90, 3.44
[1583931921] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.07, 4.42, 2.86
[1583931981] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 3.66, 3.93, 2.78
[1583935659] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583935779] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1583939708] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.55, 7.36, 5.26
[1583939708] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 307 processes with STATE = RSZDT
[1583939768] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.23, 7.05, 5.28
[1583939768] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 224 processes with STATE = RSZDT
[1583939887] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 6.72, 7.12, 5.50
[1583939947] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 7.20, 7.15, 5.63
[1583940547] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 1.74, 3.53, 4.52
[1583941147] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.48, 5.34, 4.40
[1583941207] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.56, 7.27, 5.13
[1583941267] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 13.33, 8.34, 5.64
[1583941327] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 7.23, 7.59, 5.55
[1583941627] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (23.8% change >= 20.0% threshold)
[1583941627] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.49, 3.17, 4.20
[1583942527] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 15.32, 6.43, 4.10
[1583942587] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.11, 7.83, 4.74
[1583942647] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 8.32, 7.43, 4.80
[1583942707] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.88, 6.57, 4.67
[1583943007] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.32, 2.71, 3.52
[1583943706] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1583943826] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 34% (5 min average)
[1583943907] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.19, 2.20, 2.16
[1583943967] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 12.19, 4.46, 2.93
[1583944027] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 12.24, 5.96, 3.55
[1583944087] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 9.64, 6.44, 3.87
[1583944387] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.55, 2.96, 3.11
[1583950087] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1583952946] SERVICE ALERT: WINV-104;CPU Load;WARNING;SOFT;1;CPU Load 80% (5 min average)
[1583953066] SERVICE ALERT: WINV-104;CPU Load;OK;SOFT;2;CPU Load 68% (5 min average)
[1583969169] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.97, 16.05, 9.00
[1583969230] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 49.08, 23.34, 11.91
[1583969290] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 41.66, 26.16, 13.62
[1583969350] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.31, 27.50, 14.87
[1583969469] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.28, 7.40, 5.49
[1583969529] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 3.92, 6.31, 5.23
[1583969589] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 2.28, 5.43, 5.00
[1583969650] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.45, 15.49, 13.62
[1583970550] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 4.05, 4.75, 7.97
[1583970850] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.86, 7.99, 8.21
[1583970910] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 22.64, 11.39, 9.37
[1583970970] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 24.66, 14.16, 10.45
[1583971030] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1583971030] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 20.97, 14.98, 10.97
[1583971449] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.96, 6.08, 5.19
[1583971509] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.44, 5.77, 5.13
[1583971569] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.34, 6.99, 5.60
[1583971629] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.57, 6.32, 5.45
[1583971630] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.06, 28.34, 20.12
[1583976430] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.66, 13.47, 19.07
[1583976729] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.76, 4.04, 4.74
[1583977630] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.01, 0.67, 6.25
[1583978709] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 737.26 Gb (79%) - free 194.06 Gb (21%)
[1583980403] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.61, 4.74, 2.61
[1583980463] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.87, 5.17, 2.90
[1583980523] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 3.52, 4.41, 2.78
[1583982430] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583983910] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.32, 4.81, 3.07
[1583984026] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.96, 5.43, 3.39
[1583984281] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 8.76, 7.70, 4.99
[1583984341] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 8.07, 7.70, 5.16
[1583984941] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 16.10, 9.35, 6.56
[1583985243] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 2.44, 6.25, 6.17
[1583985843] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1583985843] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 1.63, 2.71, 4.37
[1583986443] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.59, 5.77, 4.41
[1583986503] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.31, 7.55, 5.11
[1583986563] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 12.45, 8.27, 5.51
[1583986623] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 7.22, 7.57, 5.45
[1583986923] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.66, 3.18, 4.13
[1583988123] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.19, 9.01, 5.25
[1583988183] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 8.51, 8.68, 5.38
[1583988243] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.20, 7.41, 5.15
[1583988303] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 2.62, 6.37, 4.93
[1583988603] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 1.21, 3.09, 3.89
[1583989503] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 14.71, 6.94, 4.28
[1583989563] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.32, 7.56, 4.67
[1583989623] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 6.94, 7.18, 4.74
[1583989683] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 3.90, 6.29, 4.58
[1583989983] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 5.17, 4.94, 4.40
[1583990398] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 10.66, 11.28, 7.57
[1583990417] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 312 processes with STATE = RSZDT
[1583990458] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 8.47, 10.59, 7.57
[1583990478] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 223 processes with STATE = RSZDT
[1583990518] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 6.08, 9.52, 7.39
[1583990578] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 3.31, 8.10, 7.04
[1583991605] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 1.86, 2.75, 4.49
[1583991905] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 8.83, 9.16, 7.00
[1583991965] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 4.73, 7.96, 6.72
[1583992025] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 2.50, 6.74, 6.38
[1583992085] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 2.13, 5.83, 6.08
[1583992385] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.89, 2.66, 4.63
[1583992685] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.97, 7.53, 6.08
[1583992745] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 4.80, 6.54, 5.83
[1583992805] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 2.22, 5.46, 5.50
[1583992865] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 1.42, 4.64, 5.21
[1583993165] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 0.26, 1.94, 3.88