[1579775030] SERVICE ALERT: vdb-102;Total Processes;OK;HARD;4;PROCS OK: 161 processes with STATE = RSZDT
[1579775030] SERVICE ALERT: vdb-102;Swap Usage;UNKNOWN;HARD;4;check_swap: Warning free space should be more than critical free space
[1579775215] SERVICE ALERT: vdb-102;Swap Usage;OK;HARD;4;SWAP OK - 1% free (0 MB out of 975 MB)
[1579776032] SERVICE ALERT: vonly-backup;PING;UNKNOWN;HARD;1;You must specify a server address or host name
[1579776147] SERVICE ALERT: vonly-backup;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.05 ms
[1579776159] SERVICE ALERT: vdb-102;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.06 ms
[1579777915] SERVICE ALERT: vdb-102;Swap Usage;CRITICAL;HARD;1;SWAP CRITICAL - 0% free (0 MB out of 975 MB)
[1579779780] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;1;WARNING - load average: 4.08, 3.98, 3.03
[1579781370] SERVICE ALERT: vdb-102;Root Partition;CRITICAL;HARD;1;NRPE: Command 'check_disk' not defined
[1579781631] SERVICE ALERT: vdb-102;/mnt/data01 Partition;CRITICAL;HARD;4;/mnt/data01 is at 96% capacity! 3.4T of 3.7T
[1579781670] SERVICE ALERT: vdb-102;Root Partition;WARNING;HARD;4;Check the capacity of a volume using df.
[1579782488] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 2.45, 2.66, 2.96
[1579782565] SERVICE ALERT: vonly-backup;Root Partition;CRITICAL;HARD;1;(Return code of 13 for service 'Root Partition' on host 'vonly-backup' was out of bounds)
[1579782877] SERVICE ALERT: vonly-backup;Root Partition;OK;HARD;4;/ is at 51% capacity, 107G of 220G
[1579783005] SERVICE ALERT: vdb-102;Root Partition;OK;HARD;4;/ is at 11% capacity, 9.4G of 91G
[1579783640] HOST ALERT: vdb-101;DOWN;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783677] SERVICE ALERT: vdb-101;/mnt/data01 Partition;CRITICAL;HARD;1;NRPE: Command 'check_volume' not defined
[1579783686] SERVICE ALERT: vdb-103;/mnt/data01 Partition;UNKNOWN;SOFT;1;NRPE: Command 'check_volume' not defined
[1579783686] HOST ALERT: vdb-103;DOWN;SOFT;1;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783700] HOST ALERT: vdb-101;DOWN;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783715] HOST ALERT: vdb-101;DOWN;SOFT;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783724] HOST ALERT: vdb-103;DOWN;SOFT;2;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783746] SERVICE ALERT: vdb-103;/mnt/data01 Partition;UNKNOWN;HARD;4;NRPE: Command 'check_volume' not defined
[1579783752] HOST ALERT: vdb-101;DOWN;SOFT;4;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783761] HOST ALERT: vdb-103;DOWN;SOFT;3;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783790] SERVICE ALERT: vdb-101;PING;CRITICAL;HARD;1;NRPE: Command 'check_ping' not defined
[1579783799] SERVICE ALERT: vdb-103;PING;UNKNOWN;HARD;1;NRPE: Command 'check_ping' not defined
[1579783812] HOST ALERT: vdb-101;DOWN;SOFT;5;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783821] HOST ALERT: vdb-103;DOWN;SOFT;4;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783827] SERVICE ALERT: vdb-101;Root Partition;CRITICAL;HARD;1;NRPE: Command 'check_disk' not defined
[1579783836] SERVICE ALERT: vdb-103;Root Partition;UNKNOWN;HARD;1;NRPE: Command 'check_disk' not defined
[1579783865] HOST ALERT: vdb-101;DOWN;SOFT;6;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783874] HOST ALERT: vdb-103;DOWN;SOFT;5;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783902] SERVICE ALERT: vdb-101;Swap Usage;CRITICAL;HARD;1;NRPE: Command 'check_swap' not defined
[1579783911] SERVICE ALERT: vdb-103;Swap Usage;UNKNOWN;HARD;1;NRPE: Command 'check_swap' not defined
[1579783925] HOST ALERT: vdb-101;DOWN;SOFT;7;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783934] HOST ALERT: vdb-103;DOWN;SOFT;6;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783940] SERVICE ALERT: vdb-101;Total Processes;CRITICAL;HARD;1;PROCS CRITICAL: 285 processes
[1579783940] SERVICE ALERT: vdb-103;Total Processes;CRITICAL;HARD;1;PROCS CRITICAL: 414 processes
[1579783985] HOST ALERT: vdb-101;DOWN;SOFT;8;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579783994] HOST ALERT: vdb-103;DOWN;SOFT;7;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579784015] HOST ALERT: vdb-101;DOWN;SOFT;9;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579784024] HOST ALERT: vdb-103;DOWN;SOFT;8;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579784052] HOST ALERT: vdb-101;DOWN;HARD;10;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579784061] HOST ALERT: vdb-103;DOWN;SOFT;9;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579784121] HOST ALERT: vdb-103;DOWN;HARD;10;(No output on stdout) stderr: execvp(/usr/local/nagios/libexec/check_local_ping, ...) failed. errno is 2: No such file or directory
[1579784240] SERVICE ALERT: vdb-101;Total Processes;OK;HARD;4;PROCS OK: 153 processes with STATE = RSZDT
[1579784241] SERVICE ALERT: vdb-103;Total Processes;OK;HARD;4;PROCS OK: 181 processes with STATE = RSZDT
[1579784252] SERVICE ALERT: vdb-101;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 88% capacity, 3.0T of 3.6T
[1579784300] SERVICE ALERT: vdb-103;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 67% capacity, 4.6T of 7.3T
[1579784395] SERVICE ALERT: vdb-101;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.04 ms
[1579784404] SERVICE ALERT: vdb-103;PING;OK;HARD;4;PING OK - Packet loss = 0%, RTA = 0.07 ms
[1579784427] SERVICE ALERT: vdb-101;Root Partition;OK;HARD;4;/ is at 11% capacity, 19G of 183G
[1579784437] SERVICE ALERT: vdb-103;Root Partition;OK;HARD;4;/ is at 30% capacity, 127G of 457G
[1579784512] SERVICE ALERT: vdb-103;Swap Usage;WARNING;HARD;4;SWAP WARNING - 93% free (1892 MB out of 2047 MB)
[1579790869] SERVICE ALERT: vdb-102;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 94% capacity! 3.3T of 3.7T
[1579791169] SERVICE ALERT: vdb-102;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 89% capacity, 3.1T of 3.7T