[1582618543] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 34.43, 19.95, 8.88
[1582618604] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 46.72, 25.79, 11.57
[1582618664] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 46.76, 29.90, 13.88
[1582618725] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 52.04, 34.35, 16.41
[1582619920] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.26, 6.53, 5.41
[1582619980] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.43, 6.14, 5.35
[1582620040] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.56, 5.80, 5.29
[1582620100] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.09, 5.42, 5.18
[1582620161] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.95, 7.06, 4.30
[1582620221] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.18, 6.75, 4.37
[1582620281] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 4.87, 6.27, 4.35
[1582620341] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.64, 5.95, 4.36
[1582620700] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.44, 3.18, 4.30
[1582620825] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 1.38, 4.69, 17.60
[1582621241] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.87, 5.79, 4.81
[1582621247] SERVICE ALERT: vonly-backup;/media/data Partition;OK;HARD;4;/media/data is at 89% capacity, 6.1T of 7.3T
[1582621465] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 14.18, 8.85, 6.02
[1582621525] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.59, 7.58, 6.09
[1582621585] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.96, 7.03, 5.99
[1582623134] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.59, 6.73, 5.36
[1582623194] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 9.76, 7.42, 5.68
[1582623254] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.17, 7.57, 5.86
[1582623282] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 41.90, 30.06, 20.79
[1582623314] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.42, 7.11, 5.80
[1582623399] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.28, 3.74, 4.61
[1582623699] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 2.98, 4.85, 5.05
[1582623759] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 2.25, 4.30, 4.84
[1582625019] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.64, 3.82, 3.36
[1582625079] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 6.61, 4.28, 3.55
[1582626092] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.13, 6.23, 4.41
[1582626092] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 259 processes with STATE = RSZDT
[1582626152] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.69, 6.44, 4.60
[1582626152] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 245 processes with STATE = RSZDT
[1582626212] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 7.28, 6.51, 4.74
[1582626272] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 5.45, 6.14, 4.72
[1582627049] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.10, 3.70, 4.84
[1582627193] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (20.4% change >= 20.0% threshold)
[1582627193] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.70, 4.25, 4.95
[1582627193] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.54, 8.68, 16.99
[1582627949] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.40, 6.40, 5.48
[1582628009] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.50, 5.98, 5.39
[1582628069] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.49, 6.13, 5.47
[1582628093] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.71, 19.79, 17.26
[1582628093] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 4.38, 5.79, 5.26
[1582628129] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 8.34, 6.74, 5.73
[1582628153] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 3.76, 5.34, 5.14
[1582628213] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 3.13, 4.88, 4.99
[1582630949] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 84% (5 min average)
[1582631069] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;2;CPU Load 84% (5 min average)
[1582631189] SERVICE ALERT: WINV-107;CPU Load;WARNING;HARD;3;CPU Load 87% (5 min average)
[1582631789] SERVICE ALERT: WINV-107;CPU Load;OK;HARD;3;CPU Load 39% (5 min average)
[1582632173] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582632629] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.37, 4.27, 4.78
[1582632893] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.30, 4.77, 17.02
[1582633112] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 255 processes with STATE = RSZDT
[1582633172] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1582633532] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1582633592] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 255 processes with STATE = RSZDT
[1582633652] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 267 processes with STATE = RSZDT
[1582633712] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 249 processes with STATE = RSZDT
[1582633793] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.17, 0.61, 6.70
[1582633973] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 11.74, 4.94, 2.98
[1582634033] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 8.66, 5.43, 3.28
[1582634093] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.93, 4.93, 3.24
[1582634926] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582635046] SERVICE ALERT: WINV-101;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1582636149] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1582636209] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1582636269] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1582636330] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes with STATE = RSZDT
[1582636390] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 254 processes with STATE = RSZDT
[1582636450] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 256 processes with STATE = RSZDT
[1582637174] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.28, 7.77, 5.84
[1582637234] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.38, 7.39, 5.83
[1582637294] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.50, 6.94, 5.77
[1582637354] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.79, 6.46, 5.68
[1582637954] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.45, 4.59, 4.98
[1582638283] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582638583] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 256 processes with STATE = RSZDT
[1582638643] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 248 processes with STATE = RSZDT
[1582638703] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582638763] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 248 processes with STATE = RSZDT
[1582638854] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 3.40, 5.76, 5.29
[1582638914] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 2.52, 5.09, 5.08
[1582638975] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 2.22, 4.54, 4.89
[1582639440] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1582639500] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1582639935] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.26, 6.19, 5.09
[1582640050] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 9.83, 7.07, 5.52
[1582640110] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 15.24, 9.16, 6.34
[1582640170] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 9.84, 8.87, 6.42
[1582640557] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582640617] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 249 processes with STATE = RSZDT
[1582640977] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582641037] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 253 processes with STATE = RSZDT
[1582641097] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1582641376] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1582641376] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.90, 4.00, 4.73
[1582641457] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 253 processes with STATE = RSZDT
[1582641517] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 253 processes with STATE = RSZDT
[1582641577] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1582641637] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 249 processes with STATE = RSZDT
[1582641676] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 5.51, 6.07, 5.45
[1582641736] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 4.38, 5.64, 5.34
[1582641796] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 3.89, 5.27, 5.23
[1582641856] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 3.65, 4.93, 5.11
[1582641937] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582641997] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 257 processes with STATE = RSZDT
[1582642057] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 254 processes with STATE = RSZDT
[1582642117] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 251 processes with STATE = RSZDT
[1582642156] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 3.40, 4.00, 4.65
[1582642417] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 248 processes with STATE = RSZDT
[1582642717] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 258 processes with STATE = RSZDT
[1582642777] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 254 processes with STATE = RSZDT
[1582642838] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1582642898] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1582642898] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 254 processes with STATE = RSZDT
[1582643798] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582644398] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 256 processes with STATE = RSZDT
[1582644458] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 247 processes with STATE = RSZDT
[1582645118] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 254 processes with STATE = RSZDT
[1582645178] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1582645238] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 254 processes with STATE = RSZDT
[1582645298] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1582645682] SERVICE ALERT: WINV-107;CPU Load;WARNING;SOFT;1;CPU Load 83% (5 min average)
[1582645802] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1582646198] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 255 processes with STATE = RSZDT
[1582646258] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1582646318] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 256 processes with STATE = RSZDT
[1582646378] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 250 processes with STATE = RSZDT
[1582646738] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 257 processes with STATE = RSZDT
[1582646828] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 248 processes with STATE = RSZDT
[1582647582] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 275 processes with STATE = RSZDT
[1582647582] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.44, 8.53, 5.21
[1582647642] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 257 processes with STATE = RSZDT
[1582647642] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 6.61, 7.93, 5.22
[1582647725] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 5.38, 7.16, 5.17
[1582647725] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 248 processes with STATE = RSZDT
[1582647785] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 4.52, 6.60, 5.10
[1582648385] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 4.10, 4.50, 4.72
[1582648385] SERVICE FLAPPING ALERT: vonly-backup;Total Processes;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1582648985] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 13.71, 7.03, 5.36
[1582649045] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 10.04, 7.29, 5.56
[1582649105] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 6.06, 6.66, 5.45
[1582649165] SERVICE ALERT: vonly-backup;Current Load;WARNING;HARD;4;WARNING - load average: 3.91, 5.94, 5.28
[1582649465] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 2.35, 3.75, 4.51
[1582652476] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.96, 7.13, 4.91
[1582652536] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 4.74, 6.35, 4.78
[1582652596] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 3.85, 5.79, 4.68
[1582653861] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 12.75, 5.65, 4.13
[1582653921] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 10.63, 6.42, 4.50
[1582653981] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 5.93, 5.86, 4.43
[1582654941] SERVICE FLAPPING ALERT: vonly-backup;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1582673409] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 24.85, 12.62, 6.44
[1582673469] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 33.03, 16.95, 8.30
[1582673529] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.09, 18.75, 9.47
[1582673589] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.77, 22.17, 11.25
[1582673889] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 2.00, 11.32, 9.89
[1582674125] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1582674189] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 3.24, 5.79, 7.83
[1582674245] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 13% (5 min average)
[1582678089] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 65.64, 30.55, 12.69
[1582678149] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 68.28, 38.05, 16.41
[1582678210] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 69.80, 44.15, 19.97
[1582678270] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 71.34, 49.99, 23.53
[1582678809] SERVICE ALERT: vdb-103;Swap Usage;CRITICAL;SOFT;1;SWAP CRITICAL - 0% free (0 MB out of 2047 MB)
[1582678869] SERVICE ALERT: vdb-103;Swap Usage;CRITICAL;SOFT;2;SWAP CRITICAL - 0% free (0 MB out of 2047 MB)
[1582678929] SERVICE ALERT: vdb-103;Swap Usage;CRITICAL;SOFT;3;SWAP CRITICAL - 0% free (0 MB out of 2047 MB)
[1582678989] SERVICE ALERT: vdb-103;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 0% free (0 MB out of 2047 MB)
[1582680523] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.33, 5.63, 5.06
[1582680583] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 5.16, 5.25, 4.96
[1582680644] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.34, 5.48, 5.05
[1582680704] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.56, 5.08, 4.94
[1582680764] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.36, 5.52, 5.10
[1582680824] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.04, 5.30, 5.05
[1582680884] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.57, 5.35, 5.08
[1582680944] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.70, 5.13, 5.02
[1582680959] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1582681079] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 12% (5 min average)
[1582681544] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.47, 4.49, 4.82
[1582682744] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.60, 6.24, 5.31
[1582682804] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.16, 5.67, 5.17
[1582682865] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.70, 5.92, 5.30
[1582682925] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.48, 4.95, 5.00
[1582683370] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.37, 5.20, 19.05
[1582684270] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.09, 0.37, 7.30
[1582698965] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1582699085] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1582699205] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1582699805] SERVICE ALERT: WINV-107;CPU Load;OK;HARD;3;CPU Load 73% (5 min average)
[1582702123] SERVICE ALERT: vdb-101;/mnt/data01 Partition;CRITICAL;HARD;4;/mnt/data01 is at 91% capacity! 3.1T of 3.6T
[1582702269] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1582702269] SERVICE ALERT: WINV-102;OS Drive Space;WARNING;SOFT;1;No data was received from host!
[1582702432] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 11% (5 min average)
[1582702432] SERVICE ALERT: WINV-102;OS Drive Space;OK;SOFT;2;c:\ - total: 931.41 Gb - used: 305.23 Gb (33%) - free 626.18 Gb (67%)
[1582702432] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.35, 5.56, 3.16
[1582702432] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 271 processes with STATE = RSZDT
[1582702492] SERVICE ALERT: vonly-backup;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 11.46, 6.76, 3.72
[1582702492] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 247 processes with STATE = RSZDT
[1582702571] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1582702571] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;3;WARNING - load average: 7.22, 6.46, 3.86
[1582702571] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1582702649] SERVICE ALERT: vonly-backup;Current Load;OK;HARD;4;OK - load average: 5.20, 5.89, 3.87
[1582702649] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 245 processes with STATE = RSZDT
[1582702691] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 90% (5 min average)
[1582702811] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1582703411] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 16% (5 min average)