[1581323172] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 6.75, 6.88, 5.12
[1581323232] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 7.23, 7.01, 5.27
[1581323292] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 7.18, 7.04, 5.39
[1581323352] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 7.78, 7.23, 5.56
[1581324852] SERVICE ALERT: vonly-database;Current Load;OK;HARD;4;OK - load average: 3.40, 4.11, 4.82
[1581327422] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 6.30, 5.88, 5.15
[1581327482] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.19, 5.95, 5.22
[1581327542] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 6.25, 5.99, 5.28
[1581327602] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 6.44, 6.06, 5.35
[1581333902] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 3.73, 4.33, 4.82
[1581336902] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 5.62, 5.61, 5.08
[1581336962] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 6.18, 5.80, 5.18
[1581337022] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 5.01, 5.53, 5.13
[1581337082] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 4.90, 5.41, 5.12
[1581337382] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 4.45, 4.87, 4.97
[1581357902] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 8.75, 5.10, 3.19
[1581357962] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 4.18, 4.46, 3.09
[1581401052] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;1;WARNING - load average: 13.47, 9.48, 6.03
[1581401112] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;2;WARNING - load average: 13.00, 10.08, 6.45
[1581401172] SERVICE ALERT: vonly-database;Current Load;WARNING;SOFT;3;WARNING - load average: 12.19, 10.39, 6.79
[1581401232] SERVICE ALERT: vonly-database;Current Load;WARNING;HARD;4;WARNING - load average: 10.46, 10.22, 6.95