[1585725462] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 27.39, 17.96, 8.77
[1585725522] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;2;CRITICAL - load average: 30.60, 20.49, 10.21
[1585725582] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 32.28, 22.56, 11.55
[1585725642] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.04, 25.12, 13.10
[1585725942] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 22.56, 24.27, 16.10
[1585726069] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.83, 8.30, 5.73
[1585726089] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 13.15, 12.32, 10.36
[1585726129] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.81, 7.71, 5.69
[1585726149] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 12.72, 12.35, 10.50
[1585726189] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.15, 7.04, 5.59
[1585726209] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 12.50, 12.38, 10.62
[1585726242] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.14, 25.65, 18.92
[1585726249] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.11, 7.04, 5.68
[1585726269] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 15.18, 13.08, 10.97
[1585727943] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585728073] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1585728130] SERVICE ALERT: WINV-102;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585728130] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585728130] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1585728131] SERVICE ALERT: WINV-102;Data Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585728203] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585728260] SERVICE ALERT: WINV-102;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1585728260] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1585728260] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout
[1585728390] SERVICE ALERT: WINV-102;Uptime;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585728390] SERVICE ALERT: WINV-102;OS Drive Space;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585728390] SERVICE ALERT: WINV-102;Memory Usage;CRITICAL;HARD;3;CRITICAL - Socket timeout
[1585731002] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 7.19, 5.93, 5.08
[1585731062] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 7.61, 6.28, 5.25
[1585731069] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.94, 29.71, 26.02
[1585731123] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 7.41, 6.51, 5.40
[1585731183] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 9.51, 7.24, 5.72
[1585734069] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 22.79, 22.20, 24.50
[1585734365] SERVICE ALERT: vonly-replica;/media/data Partition;CRITICAL;HARD;4;/media/data is at 91% capacity! 6.2T of 7.3T
[1585736583] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 2.92, 3.28, 4.49
[1585737783] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 6.50, 6.01, 5.13
[1585737843] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 6.22, 6.04, 5.20
[1585737888] SERVICE ALERT: WINV-105;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1585737903] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 6.04, 6.01, 5.24
[1585737963] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 5.41, 5.85, 5.24
[1585738008] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1585739163] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 3.13, 3.92, 4.58
[1585745560] HOST ALERT: WINV-102;DOWN;SOFT;1;(Host check timed out after 30.01 seconds)
[1585745603] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 5% (5 min average)
[1585745607] HOST ALERT: WINV-102;UP;SOFT;1;PING OK - Packet loss = 0%, RTA = 0.31 ms
[1585745790] SERVICE ALERT: WINV-102;Memory Usage;OK;HARD;3;Memory usage: total:208877.43 MB - used: 60544.16 MB (29%) - free: 148333.27 MB (71%)
[1585745790] SERVICE ALERT: WINV-102;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 4 minute(s)
[1585745793] SERVICE ALERT: WINV-102;OS Drive Space;OK;HARD;3;c:\ - total: 931.41 Gb - used: 305.40 Gb (33%) - free 626.02 Gb (67%)
[1585746049] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.55, 4.54, 4.96
[1585746131] SERVICE ALERT: WINV-102;Data Drive Space;WARNING;HARD;3;d:\ - total: 1862.64 Gb - used: 1592.42 Gb (85%) - free 270.22 Gb (15%)
[1585746942] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.14, 3.15, 15.22
[1585747842] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.06, 0.53, 6.12
[1585750149] SERVICE ALERT: vinfra-200;MySQL Connection Time;CRITICAL;SOFT;1;CHECK_NRPE STATE CRITICAL: Socket timeout after 30 seconds.
[1585750211] SERVICE ALERT: vinfra-200;MySQL Connection Time;OK;SOFT;2;OK - 0.10 seconds to connect as nagios
[1585764669] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 10.00, 9.98, 10.00
[1585766769] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.12, 10.09, 10.01
[1585766829] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;2;OK - load average: 9.85, 10.01, 9.99
[1585767489] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.20, 10.05, 10.02
[1585767549] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.09, 10.05, 10.02
[1585767609] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 10.24, 10.10, 10.03
[1585767670] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 10.01, 10.05, 10.02
[1585767970] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 9.77, 9.96, 9.98
[1585768570] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.27, 10.09, 10.03
[1585768630] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.98, 10.03, 10.01
[1585768690] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 9.91, 10.00, 10.00
[1585768750] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.09, 10.04, 10.01
[1585768810] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.08, 10.04, 10.01
[1585768870] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 9.95, 10.00, 10.00
[1585771930] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.21, 10.11, 10.03
[1585771990] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.10, 10.09, 10.03
[1585772050] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 9.80, 10.00, 10.00
[1585773310] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.08, 10.05, 10.02
[1585773370] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.99, 10.02, 10.01
[1585773430] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 9.95, 9.99, 10.00
[1585773490] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.27, 10.07, 10.02
[1585773550] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 9.97, 10.02, 10.01
[1585773610] SERVICE ALERT: vonly-replica;Current Load;OK;SOFT;3;OK - load average: 9.96, 10.00, 10.00
[1585780249] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.88, 7.09, 5.36
[1585780309] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.05, 6.43, 5.24
[1585780369] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.79, 6.48, 5.31
[1585780429] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 4.82, 5.92, 5.19
[1585780542] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 17.21, 14.03, 9.28
[1585780603] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 17.43, 14.81, 9.85
[1585780663] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 11.06, 13.45, 9.69
[1585780723] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 10.02, 12.76, 9.69
[1585781170] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 10.63, 10.78, 10.21
[1585781230] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 10.54, 10.79, 10.25
[1585781290] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 9.87, 10.55, 10.20
[1585781350] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 9.43, 10.30, 10.14
[1585781629] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.42, 3.71, 4.87
[1585781650] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.42, 9.34, 9.78
[1585781923] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.28, 2.56, 6.90
[1585784923] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 15.19, 6.34, 2.80
[1585784984] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.40, 11.31, 4.72
[1585785044] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 26.28, 13.94, 6.03
[1585785104] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 26.20, 16.10, 7.27
[1585785404] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.4% change >= 20.0% threshold)
[1585785404] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 41.51, 29.63, 15.36
[1585785550] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 11.80, 11.58, 10.19
[1585785610] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 11.74, 11.57, 10.27
[1585785670] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 11.50, 11.49, 10.32
[1585785730] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 11.90, 11.55, 10.41
[1585786429] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.88, 5.70, 5.13
[1585786489] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.30, 5.99, 5.27
[1585786549] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.83, 6.17, 5.38
[1585786609] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.19, 5.77, 5.28
[1585791104] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585791705] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.20, 14.92, 19.54
[1585792010] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.68, 2.97, 4.18
[1585792330] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.23, 8.66, 9.81
[1585792605] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.92, 1.55, 7.86