[1585811127] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1585811248] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585811368] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;HARD;3;CPU Load 100% (5 min average)
[1585811603] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 96% (5 min average)
[1585811723] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 56% (5 min average)
[1585811968] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 1% (5 min average)
[1585812105] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 20.44, 16.76, 9.37
[1585812165] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 16.84, 16.39, 9.71
[1585812225] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 15.90, 16.31, 10.11
[1585812285] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 14.62, 16.00, 10.40
[1585812430] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 15.46, 12.98, 10.32
[1585812490] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 16.82, 13.83, 10.78
[1585812550] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 16.51, 14.26, 11.12
[1585812610] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 18.14, 15.15, 11.63
[1585812710] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 9.72, 8.29, 6.16
[1585812771] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.71, 7.99, 6.19
[1585812831] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 8.16, 8.25, 6.40
[1585812891] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.90, 7.79, 6.36
[1585814985] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 36.58, 22.48, 18.96
[1585815910] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.30, 29.12, 25.30
[1585816563] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 5.72, 6.16, 5.17
[1585816623] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 4.34, 5.67, 5.06
[1585816683] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 4.54, 5.45, 5.02
[1585816743] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 4.20, 5.17, 4.95
[1585817085] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 11.39, 17.53, 19.80
[1585817385] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1585817385] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.07, 20.78, 20.32
[1585820110] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 21.31, 21.01, 23.66
[1585821121] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (4083 MB out of 21455 MB)
[1585821181] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 19% free (3921 MB out of 21455 MB)
[1585821241] SERVICE ALERT: vdb-101;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 19% free (3945 MB out of 21455 MB)
[1585821301] SERVICE ALERT: vdb-101;Swap Usage;WARNING;HARD;4;SWAP WARNING - 19% free (3900 MB out of 21455 MB)
[1585823085] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585824529] SERVICE ALERT: WINV-105;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1585824649] SERVICE ALERT: WINV-105;CPU Load;OK;SOFT;2;CPU Load 6% (5 min average)
[1585826391] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.56, 4.76, 4.99
[1585828191] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.39, 6.03, 5.19
[1585828251] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.50, 5.43, 5.03
[1585828311] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;3;OK - load average: 3.43, 4.93, 4.88
[1585828971] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.93, 5.64, 5.14
[1585829031] SERVICE ALERT: vinfra-200;Current Load;OK;SOFT;2;OK - load average: 4.41, 5.13, 4.99
[1585829091] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 6.25, 5.68, 5.20
[1585829151] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 5.78, 5.64, 5.21
[1585829211] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.60, 6.86, 5.67
[1585829271] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.55, 6.12, 5.48
[1585830471] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 4.15, 4.57, 4.94
[1585830771] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.18, 5.41, 5.17
[1585830831] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.06, 5.50, 5.22
[1585830891] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 4.38, 5.08, 5.09
[1585830951] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (22.2% change >= 20.0% threshold)
[1585830951] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.02, 5.37, 5.21
[1585831251] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 3.43, 4.35, 4.82
[1585832085] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.49, 3.11, 15.43
[1585832985] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.54, 1.20, 6.50
[1585834501] SERVICE ALERT: vdb-101;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (2072 MB out of 21455 MB)
[1585834810] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.72, 25.50, 25.08
[1585836951] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585839010] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 23.94, 24.59, 24.99
[1585839568] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1585839688] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1585839808] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 11% (5 min average)
[1585840596] SERVICE ALERT: WINV-107;Data Drive Space;CRITICAL;HARD;3;d:\ - total: 931.51 Gb - used: 846.15 Gb (91%) - free 85.36 Gb (9%)
[1585844796] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.51 Gb - used: 835.42 Gb (90%) - free 96.09 Gb (10%)
[1585849510] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 25.72, 25.58, 25.16
[1585849596] SERVICE ALERT: WINV-107;Data Drive Space;CRITICAL;HARD;3;d:\ - total: 931.51 Gb - used: 842.87 Gb (90%) - free 88.64 Gb (10%)
[1585850196] SERVICE FLAPPING ALERT: WINV-107;Data Drive Space;STARTED; Service appears to have started flapping (21.6% change >= 20.0% threshold)
[1585850196] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.51 Gb - used: 831.49 Gb (89%) - free 100.02 Gb (11%)
[1585853796] SERVICE ALERT: WINV-107;Data Drive Space;CRITICAL;HARD;3;d:\ - total: 931.51 Gb - used: 841.59 Gb (90%) - free 89.92 Gb (10%)
[1585856196] SERVICE ALERT: WINV-107;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.51 Gb - used: 834.65 Gb (90%) - free 96.86 Gb (10%)
[1585861210] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 24.08, 24.43, 24.82
[1585865796] SERVICE FLAPPING ALERT: WINV-107;Data Drive Space;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1585866585] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 19.35, 11.18, 6.45
[1585866645] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 26.38, 14.62, 7.93
[1585866705] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 29.59, 17.26, 9.23
[1585866766] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 30.83, 19.86, 10.63
[1585867066] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 9.81, 15.26, 11.41
[1585867666] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 2.32, 4.68, 7.82
[1585871566] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 30.70, 18.60, 8.71
[1585871627] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.70, 19.78, 9.73
[1585871687] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;3;CRITICAL - load average: 35.87, 23.19, 11.51
[1585871747] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 41.64, 27.23, 13.64
[1585872010] SERVICE ALERT: vonly-replica;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 27.43, 26.81, 25.03
[1585872047] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (24.9% change >= 20.0% threshold)
[1585872047] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 25.59, 24.59, 16.16
[1585872951] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.54, 5.67, 5.08
[1585873011] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.01, 5.36, 5.01
[1585873071] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 9.27, 6.41, 5.39
[1585873131] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.32, 5.79, 5.23
[1585875231] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 5.34, 4.80, 4.92
[1585877031] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 7.00, 5.72, 5.22
[1585877091] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.56, 6.01, 5.35
[1585877151] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 5.71, 5.73, 5.29
[1585877211] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.28, 5.58, 5.27
[1585877747] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1585879311] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 5.72, 4.82, 4.96
[1585879611] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 4.99, 5.29, 5.15
[1585879671] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 7.83, 6.00, 5.40
[1585879731] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 13.88, 7.91, 6.09
[1585879791] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STARTED; Service appears to have started flapping (22.0% change >= 20.0% threshold)
[1585879791] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 12.30, 8.56, 6.43
[1585880691] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 0.68, 2.55, 4.70
[1585880710] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 22.54, 23.49, 25.00
[1585881347] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.90, 1.52, 6.54
[1585885792] SERVICE FLAPPING ALERT: vinfra-200;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)