[1583136166] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 267 processes with STATE = RSZDT
[1583136242] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 334 processes with STATE = RSZDT
[1583136364] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 250 processes with STATE = RSZDT
[1583136848] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 9.33, 5.40, 4.29
[1583136908] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;2;WARNING - load average: 7.01, 5.60, 4.44
[1583136968] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;3;OK - load average: 4.09, 5.02, 4.31
[1583137063] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 23.42, 14.03, 8.77
[1583137124] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 20.67, 14.86, 9.38
[1583137184] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 17.14, 14.80, 9.70
[1583137244] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 26.65, 17.73, 11.02
[1583137544] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (23.3% change >= 20.0% threshold)
[1583137544] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 28.00, 25.19, 16.31
[1583138144] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 3.35, 14.12, 16.65
[1583138444] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 32.55, 24.41, 20.09
[1583138563] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 8.32, 6.27, 5.15
[1583138623] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 6.80, 6.23, 5.20
[1583138683] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 6.34, 6.18, 5.25
[1583138743] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 6.26, 6.38, 5.39
[1583140244] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.79, 15.13, 19.44
[1583140844] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 23.59, 21.86, 21.07
[1583141144] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 5.33, 12.27, 17.22
[1583141444] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 43.02, 26.27, 21.14
[1583142163] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;1;WARNING - load average: 12.22, 11.84, 10.46
[1583142224] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;2;WARNING - load average: 13.22, 12.16, 10.66
[1583142284] SERVICE ALERT: vonly-replica;Current Load;WARNING;SOFT;3;WARNING - load average: 13.49, 12.36, 10.82
[1583142344] SERVICE ALERT: vonly-replica;Current Load;WARNING;HARD;4;WARNING - load average: 13.23, 12.48, 10.96
[1583145163] SERVICE ALERT: WINV-107;CPU Load;CRITICAL;SOFT;1;CPU Load 99% (5 min average)
[1583145283] SERVICE ALERT: WINV-107;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1583145403] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 94% (5 min average)
[1583145523] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583145643] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 9% (5 min average)
[1583145644] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 13.34, 15.55, 19.86
[1583146544] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 39.76, 30.70, 24.11
[1583147144] SERVICE ALERT: vonly-replica;Current Load;OK;HARD;4;OK - load average: 8.53, 9.19, 9.97
[1583151043] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 2.74, 3.24, 4.57
[1583151344] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 0.31, 5.07, 19.20
[1583151463] SERVICE ALERT: vdb-103;/mnt/data01 Partition;OK;HARD;4;/mnt/data01 is at 80% capacity, 5.5T of 7.3T
[1583151463] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;1;WARNING - load average: 9.40, 9.18, 6.18
[1583151523] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;2;WARNING - load average: 8.62, 9.05, 6.32
[1583151583] SERVICE ALERT: vdb-staging;Current Load;WARNING;SOFT;3;WARNING - load average: 7.96, 8.75, 6.39
[1583151644] SERVICE ALERT: vdb-staging;Current Load;WARNING;HARD;4;WARNING - load average: 8.01, 8.59, 6.49
[1583152244] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.33, 0.53, 7.47
[1583152544] SERVICE ALERT: vdb-staging;Current Load;OK;HARD;4;OK - load average: 0.00, 1.36, 3.85
[1583154163] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;1;/mnt/data01 is at 81% capacity! 5.5T of 7.3T
[1583154223] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;2;/mnt/data01 is at 81% capacity! 5.5T of 7.3T
[1583154283] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;SOFT;3;/mnt/data01 is at 81% capacity! 5.5T of 7.3T
[1583154343] SERVICE ALERT: vdb-103;/mnt/data01 Partition;WARNING;HARD;4;/mnt/data01 is at 81% capacity! 5.5T of 7.3T
[1583157103] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;1;d:\ - total: 931.32 Gb - used: 745.51 Gb (80%) - free 185.81 Gb (20%)
[1583157223] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;SOFT;2;d:\ - total: 931.32 Gb - used: 745.65 Gb (80%) - free 185.67 Gb (20%)
[1583157343] SERVICE ALERT: WINV-110;Data Drive Space;WARNING;HARD;3;d:\ - total: 931.32 Gb - used: 745.85 Gb (80%) - free 185.47 Gb (20%)
[1583157344] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1583158025] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583158085] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583158145] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 252 processes with STATE = RSZDT
[1583158148] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;1;SWAP WARNING - 20% free (405 MB out of 2047 MB)
[1583158205] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 250 processes with STATE = RSZDT
[1583158208] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;2;SWAP WARNING - 20% free (404 MB out of 2047 MB)
[1583158268] SERVICE ALERT: vdb-103;Swap Usage;WARNING;SOFT;3;SWAP WARNING - 20% free (403 MB out of 2047 MB)
[1583158328] SERVICE ALERT: vdb-103;Swap Usage;WARNING;HARD;4;SWAP WARNING - 20% free (401 MB out of 2047 MB)
[1583159105] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583159165] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;2;PROCS OK: 246 processes with STATE = RSZDT
[1583171228] SERVICE ALERT: vdb-103;Swap Usage;CRITICAL;HARD;4;SWAP CRITICAL - 10% free (203 MB out of 2047 MB)
[1583183225] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583183285] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 253 processes with STATE = RSZDT
[1583183345] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 253 processes with STATE = RSZDT
[1583183405] SERVICE ALERT: vonly-backup;Total Processes;WARNING;HARD;4;PROCS WARNING: 253 processes with STATE = RSZDT
[1583184305] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 248 processes with STATE = RSZDT
[1583190905] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 251 processes with STATE = RSZDT
[1583190965] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 251 processes with STATE = RSZDT
[1583191025] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;3;PROCS WARNING: 251 processes with STATE = RSZDT
[1583191085] SERVICE ALERT: vonly-backup;Total Processes;OK;HARD;4;PROCS OK: 248 processes with STATE = RSZDT
[1583192144] SERVICE ALERT: vapps-100;Current Load;CRITICAL;SOFT;1;CRITICAL - load average: 32.92, 18.33, 9.68
[1583192204] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.50, 19.99, 10.83
[1583192264] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 27.54, 21.39, 11.88
[1583192286] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;1;PROCS WARNING: 252 processes with STATE = RSZDT
[1583192324] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 26.27, 22.88, 13.04
[1583192346] SERVICE ALERT: vonly-backup;Total Processes;WARNING;SOFT;2;PROCS WARNING: 252 processes with STATE = RSZDT
[1583192406] SERVICE ALERT: vonly-backup;Total Processes;OK;SOFT;3;PROCS OK: 248 processes with STATE = RSZDT
[1583192925] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.76, 4.87, 7.97
[1583193428] SERVICE ALERT: vonly-backup;Current Load;WARNING;SOFT;1;WARNING - load average: 7.49, 5.52, 2.93
[1583193488] SERVICE ALERT: vonly-backup;Current Load;OK;SOFT;2;OK - load average: 4.30, 4.98, 2.91
[1583196225] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;1;WARNING - load average: 22.76, 8.09, 3.92
[1583196285] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;2;WARNING - load average: 27.97, 12.69, 5.80
[1583196343] SERVICE ALERT: WINV-110;Data Drive Space;OK;HARD;3;d:\ - total: 931.32 Gb - used: 744.29 Gb (80%) - free 187.03 Gb (20%)
[1583196345] SERVICE ALERT: vapps-100;Current Load;WARNING;SOFT;3;WARNING - load average: 19.88, 13.25, 6.43
[1583196405] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 20.58, 14.19, 7.16
[1583197006] SERVICE FLAPPING ALERT: vapps-100;Current Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1583197006] SERVICE ALERT: vapps-100;Current Load;CRITICAL;HARD;4;CRITICAL - load average: 33.47, 29.16, 17.75
[1583197843] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;1;WARNING - load average: 5.57, 5.90, 5.04
[1583197903] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;2;WARNING - load average: 4.90, 5.69, 5.03
[1583197963] SERVICE ALERT: vinfra-200;Current Load;WARNING;SOFT;3;WARNING - load average: 7.45, 6.14, 5.22
[1583198023] SERVICE ALERT: vinfra-200;Current Load;WARNING;HARD;4;WARNING - load average: 5.98, 5.90, 5.19
[1583201623] SERVICE ALERT: vinfra-200;Current Load;OK;HARD;4;OK - load average: 1.87, 3.41, 4.53
[1583202106] SERVICE ALERT: vapps-100;Current Load;WARNING;HARD;4;WARNING - load average: 12.80, 14.07, 18.86
[1583203306] SERVICE ALERT: vapps-100;Current Load;OK;HARD;4;OK - load average: 0.56, 1.27, 6.84
[1583208106] SERVICE FLAPPING ALERT: vapps-100;Current Load;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1583213364] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1583213443] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;1;CPU Load 91% (5 min average)
[1583213484] SERVICE ALERT: WINV-101;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583213563] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 31% (5 min average)
[1583213604] SERVICE ALERT: WINV-101;CPU Load;OK;HARD;3;CPU Load 24% (5 min average)
[1583213683] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 86% (5 min average)
[1583213803] SERVICE ALERT: WINV-102;CPU Load;OK;SOFT;2;CPU Load 78% (5 min average)
[1583220523] SERVICE ALERT: WINV-102;CPU Load;WARNING;SOFT;1;CPU Load 87% (5 min average)
[1583220643] SERVICE ALERT: WINV-102;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1583220763] SERVICE ALERT: WINV-102;CPU Load;OK;HARD;3;CPU Load 6% (5 min average)